<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating Mass from Radial Density</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Let \(\rho(x)={@func@}\) represent the radial density of a disk. Calculate the mass of a disk of radius {@radius@}.</p>
<p>\(m=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Applying the Equation \(\displaystyle m=\int_0^r 2 \pi x \rho (x)dx , \) we find</p>
<p>\(\begin{align} \displaystyle m &=\int_0^r 2 \pi x ({@func@})\,dx\\
&=2 \pi \int_0^{@radius@} {@x*func@}dx \\
&=2 \pi \bigg[ ({@int(x*func,x)@}) \bigg]_0^{@radius@}\\
&={@ta@}. \end{align}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>x_power:rand_with_step(1/2,7/2,1/2);
aa: rand_with_step(1/4,9,1/4);
radius:rand_with_step(1,9,1);
func:aa*x^x_power;
ta:int(func*2*pi*x,x,0,radius);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Let \(\rho(x)={@func@}\) represent the radial density of a disk. Calculate the mass of a disk of radius {@radius@}.</p>
<p>\(m={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Wastalas</p>
<p id="description">The question considers a function that represents a the radial density of a disk, the students is required to Calculate the mass of a disk of radius at a given radius.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1284260851</deployedseed>
    <deployedseed>1686492387</deployedseed>
    <deployedseed>484771552</deployedseed>
    <deployedseed>923258739</deployedseed>
    <deployedseed>1644742187</deployedseed>
    <deployedseed>62847454</deployedseed>
    <deployedseed>114972442</deployedseed>
    <deployedseed>1735779106</deployedseed>
    <deployedseed>482183428</deployedseed>
    <deployedseed>1970021340</deployedseed>
    <deployedseed>362044066</deployedseed>
    <deployedseed>170067608</deployedseed>
    <deployedseed>1755339544</deployedseed>
    <deployedseed>1123822376</deployedseed>
    <deployedseed>2123269031</deployedseed>
    <deployedseed>632001962</deployedseed>
    <deployedseed>1644401516</deployedseed>
    <deployedseed>1280039874</deployedseed>
    <deployedseed>2120766757</deployedseed>
    <deployedseed>1294808554</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
