<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculation and Limits of a Given Function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p>
<p>Note 1: If any part of this question has no solution please simply answer <code>NO</code>. Otherwise, give your answers correct to 3 significant figures.</p>
<p> Note 2: Please remember to use radians, not degrees.</p>
<p></p>
<p>1. If \(\displaystyle f(x) = {@fx@}\), evaluate \(f({@0@})\) = [[input:ans1]] [[validation:ans1]]&nbsp;[[feedback:prt1]]</p>
<p></p>
<p>2. Fill in the table below.</p><table>
<caption></caption>
<thead>
<tr>
<th scope="col" style="text-align: center;border: 1px solid black; ">\(x\)</th>
<th scope="col" style="text-align: center; border: 1px solid black;">\(-0.5\)</th>
<th scope="col" style="text-align: center; border: 1px solid black;">\(0.5\)</th><th scope="col" style="text-align: center;border: 1px solid black; ">\(-0.25\)<br></th>
<th scope="col" style="text-align: center; border: 1px solid black;">\(0.25\)<br></th>
<th scope="col" style="text-align: center; border: 1px solid black;">\(-0.1\)<br></th>
<th scope="col" style="text-align: center; border: 1px solid black;">\(0.1\)<br></th>
</tr>
</thead>
<tbody>
<tr>
<td style="text-align: center;border: 1px solid black; ">\(f(x)\)</td>
<td style="text-align: center;border: 1px solid black; ">[[input:ans2]] [[validation:ans2]]&nbsp;[[feedback:prt2]]<br></td>
<td style="text-align: center;border: 1px solid black; ">[[input:ans3]] [[validation:ans3]]&nbsp;[[feedback:prt3]]<br></td>
<td style="text-align: center;border: 1px solid black; ">[[input:ans4]] [[validation:ans4]]&nbsp;[[feedback:prt4]]<br></td>
<td style="text-align: center;border: 1px solid black; ">[[input:ans5]] [[validation:ans5]]&nbsp;[[feedback:prt5]]<br></td>
<td style="text-align: center;border: 1px solid black; ">[[input:ans6]] [[validation:ans6]]&nbsp;[[feedback:prt6]]<br></td>
<td style="text-align: center;border: 1px solid black; ">[[input:ans7]] [[validation:ans7]]&nbsp;[[feedback:prt7]]<br></td>
</tr>
</tbody>
</table>
<p><br></p><p>3.&nbsp;&nbsp;\( \displaystyle \lim_{x\rightarrow 0} f(x) = \)&nbsp;[[input:ans8]]&nbsp; [[validation:ans8]]&nbsp;[[feedback:prt8]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b><h3>Worked Solution</h3></b></p>
<p>1. To evaluate \(f(0)\), we substitute \(x = 0\) into the function \(f(x)\): \(f(0) = \frac{\cos(0) - 1}{0} = \frac{1 - 1}{0}\)</p>
<p>However, this expression is undefined as it involves division by zero. Therefore, \(f(0)\) does not have a finite value.</p>
<p>2. Now, let's fill in the table to find \(\displaystyle \lim_{{x \to 0}} f(x)\):</p>
<table>
<caption></caption>
<thead>
<tr>
<th scope="col" style="text-align: center;border: 1px solid black; ">\(x\)</th>
<th scope="col" style="text-align: center; border: 1px solid black;">\(-0.5\)</th>
<th scope="col" style="text-align: center; border: 1px solid black;">\(0.5\)</th><th scope="col" style="text-align: center;border: 1px solid black; ">\(-0.25\)<br></th>
<th scope="col" style="text-align: center; border: 1px solid black;">\(0.25\)<br></th>
<th scope="col" style="text-align: center; border: 1px solid black;">\(-0.1\)<br></th>
<th scope="col" style="text-align: center; border: 1px solid black;">\(0.1\)<br></th>
</tr>
</thead>
<tbody>
<tr>
<td style="text-align: center;border: 1px solid black; ">\(f(x)\)</td>
<td style="text-align: center;border: 1px solid black; ">{@dispsf(ta2,3)@}<br></td>
<td style="text-align: center;border: 1px solid black; ">{@dispsf(ta3,3)@}<br></td>
<td style="text-align: center;border: 1px solid black; ">{@dispsf(ta4,3)@}<br></td>
<td style="text-align: center;border: 1px solid black; ">{@dispsf(ta5,3)@}<br></td>
<td style="text-align: center;border: 1px solid black; ">{@dispsf(ta6,3)@}<br></td>
<td style="text-align: center;border: 1px solid black; ">{@dispsf(ta7,3)@}<br></td>
</tr>
</tbody>
</table>

<br>
<p>3. The value of \(\displaystyle \lim_{{x \to 0}} f(x)\) represents the limit of the function \(f(x)\) as \(x\) approaches \(0\). </p>
<p>To find this limit, we can use either direct substitution (which may not work in this case due to division by zero) or the L'Hôpital's Rule.</p>
<p>L'Hôpital's Rule states that if the limit of a function \( \displaystyle \frac{g(x)}{h(x)}\) as \(x\) approaches a certain value is of the form \(\frac{0}{0}\) or \( \displaystyle \frac{\infty}{\infty}\), then the limit of the function can be found by taking the derivative of both \(g(x)\) and \(h(x)\) with respect to \(x\) and then evaluating the new limit.</p>
<p>In this case, we have \(\displaystyle \lim_{{x \to 0}} f(x)\) in the form \(\frac{0}{0}\), so we can use L'Hôpital's Rule:</p>
<p>\[\lim_{{x \to 0}} f(x) = \lim_{{x \to 0}} \frac{\cos(x) - 1}{x} = \lim_{{x \to 0}} \frac{-\sin(x)}{1} = -\sin(0) = 0\]</p>
<p>Thus, the value of \(\displaystyle \lim_{{x \to 0}} f(x)\) is \(0\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[fx:(cos(x)-1)/x;
ta1:"NO";
ta2:ev(fx,x=-0.5);
ta3:ev(fx,x=0.5);
ta4:ev(fx,x=-0.25);
ta5:ev(fx,x=0.25);
ta6:ev(fx,x=-0.1);
ta7:ev(fx,x=0.1);
ta8:0;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\( \displaystyle \lim_{x\rightarrow 0} f(x) ={@ta8@} \)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">NA</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Juma Zevick</p>
<p id="description"> This STACK question asks students to find the limit of a given function as x approaches a specific value. Students are required to compute the limit and determine whether it exists or not.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>string</type>
      <tans>ta1</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>dispsf(ta2,3)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>dispsf(ta3,3)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>numerical</type>
      <tans>dispsf(ta4,3)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>numerical</type>
      <tans>dispsf(ta5,3)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>numerical</type>
      <tans>dispsf(ta6,3)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>numerical</type>
      <tans>dispsf(ta7,3)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>ta8</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for crrect answert</description>
        <answertest>String</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(f(0) = \frac{\cos{(0)}-1}{0}=\frac{1-1}{0}\), which is clearly not defined as we cannot divide by \(0\). Therefore, this has no solution.&nbsp;</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>NumSigFigs</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(f(-0.5) = \frac{\cos{(-0.5)}-1}{-0.5} = {@dispsf(ta2,3)@}.\)<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>NumSigFigs</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(f(0.5) = \frac{\cos{(0.5)}-1}{0.5} = {@dispsf(ta3,3)@}.\)<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>NumSigFigs</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(f(-0.25) = \frac{\cos{(-0.25)}-1}{-0.25} = {@dispsf(ta4,3)@}.\)<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>NumSigFigs</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(f(0.25) = \frac{\cos{(0.25)}-1}{0.25} = {@dispsf(ta5,3)@}.\)<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>NumSigFigs</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(f(-0.1) = \frac{\cos{(-0.1)}-1}{-0.1} = {@dispsf(ta6,3)@}.\)<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>NumSigFigs</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>\(f(0.1) = \frac{\cos{(0.1)}-1}{0.1} = {@dispsf(ta7,3)@}.\)<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>ta8</tans>
        <testoptions>3</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>There is a limit in this case.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>String</answertest>
        <sans>ans8</sans>
        <tans>NO</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Considering the values obtained in the table, \( \displaystyle \lim_{x\rightarrow 0} f(x) = {@ta8@}\).<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispsf(ta2,3)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>dispsf(ta3,3)</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>dispsf(ta4,3)</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>dispsf(ta5,3)</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>dispsf(ta6,3)</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>dispsf(ta7,3)</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ta8</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
