<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Chain Rule and Partial Derivatives</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p></p>
<p>Given the following information</p>

<p>\(\displaystyle  w={@func_w@}\)</p>
<p>\(\displaystyle  x={@func_x@}\)</p>
<p>\(\displaystyle  y={@func_y@}\)</p>
<p>\(\displaystyle  z={@func_z@}\)</p>
<p>\(\displaystyle  p={@func_p@}\)</p>

<p> We can use Chain Rule to determine \( \displaystyle \frac{\partial{w}}{\partial{t}}\) and  \( \displaystyle \frac{\partial{w}}{\partial{s}}\) as follows; </p>

\(\displaystyle \frac{\partial{w}}{\partial{t}} = \frac{\partial{w}}{\partial{x}}\frac{\partial{x}}{\partial{p}}\frac{d{p}}{d{t}}+ \frac{\partial{w}}{\partial{y}}\frac{\partial{y}}{\partial{p}}\frac{d{p}}{d{t}}+  \frac{\partial{w}}{\partial{y}}\frac{\partial{y}}{\partial{t}}+ \frac{\partial{w}}{\partial{z}}\frac{\partial{z}}{\partial{t}}\)
<p>\(\displaystyle \quad = \)  [[input:ans1]] [[validation:ans1]] [[feedback:prt1]] </p>
<br>
\(\displaystyle \frac{\partial{w}}{\partial{s}}= \frac{\partial{w}}{\partial{y}}\frac{\partial{y}}{\partial{s}}+\frac{\partial{w}}{\partial{z}}\frac{\partial{z}}{\partial{s}}\)
<p>\(\displaystyle \quad = \)  [[input:ans2]] [[validation:ans2]] [[feedback:prt2]] </p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>Worked solution</p>
<p> We need to compute the partial derivates of the following functions as shown,</p>
\(\displaystyle \frac{\partial{w}}{\partial{x}}= {@diff_wx@}\)
\(\displaystyle\frac{\partial{x}}{\partial{p}}= {@diff_xp@}\)
\(\displaystyle \frac{\partial{p}}{\partial{t}}= {@diff_pt@}\)
\(\displaystyle \frac{\partial{w}}{\partial{y}}= {@diff_wy@}\)
\(\displaystyle \frac{\partial{y}}{\partial{p}}= {@diff_yp@}\)
\(\displaystyle \frac{\partial{y}}{\partial{t}}= {@diff_yt@}\)
\(\displaystyle \frac{\partial{w}}{\partial{z}}= {@diff_wz@}\)
\(\displaystyle \frac{\partial{z}}{\partial{t}}= {@diff_zt@}\)

<p> Then substitute them in the formulas below,</p>

\(\displaystyle \frac{\partial{w}}{\partial{t}} = \frac{\partial{w}}{\partial{x}}\frac{\partial{x}}{\partial{p}}\frac{\partial{p}}{\partial{t}}+ \frac{\partial{w}}{\partial{y}}\frac{\partial{y}}{\partial{p}}\frac{\partial{p}}{\partial{t}}+  \frac{\partial{w}}{\partial{y}}\frac{\partial{y}}{\partial{t}}+ \frac{\partial{w}}{\partial{z}}\frac{\partial{z}}{\partial{t}}\)
<p>\(\displaystyle \quad = {@ta1@}\)  </p>
<br>
\( \displaystyle \frac{\partial{w}}{\partial{s}}= \frac{\partial{w}}{\partial{y}}\frac{\partial{y}}{\partial{s}}+\frac{\partial{w}}{\partial{z}}\frac{\partial{z}}{\partial{s}}\)
<p>\(\displaystyle \quad = {@ta1@}\) </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text>/*Question variables*/
aa:rand_with_prohib(-5,5,[0]);
bb:rand_with_prohib(-5,5,[0]);
cc:rand_with_prohib(-5,5,[0]);
dd:rand_with_prohib(-5,5,[0]);
ee:rand_with_prohib(-5,5,[0]);

/*Question func*/
func_w:sqrt(x^aa+y^aa)+bb*z/y;
func_x:sin(cc*p);
func_y:bb*p+dd*t-ee*s;
func_z:t^dd/(s^aa);
func_p:bb-aa*t;


/*Worked Soln variables*/
diff_wx:diff(func_w,x);diff_xp:diff(func_x,p);diff_pt:diff(func_p,t);
diff_wy:diff(func_w,y); diff_yp:diff(func_y,p);
diff_wz:diff(func_w,z);diff_yt:diff(func_y,t);
diff_zt:diff(func_z,t);
diff_ys:diff(func_y,s);
diff_zs:diff(func_z,s);


ta1:diff_wx*diff_xp*diff_pt+ diff_wy*diff_yp*diff_wy+ diff_wy*diff_yt+ diff_wz*diff_zt;

ta2:diff_wy*diff_ys+diff_wz*diff_zs;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>1</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1: checking for correct answer ta1*/
/*node2: Computing the derivative of func w wrt to t and for going the other functions*/
mistake1:diff(func_w,t)</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x,y,z,t,u</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1: checking for correct answer ta2*/
/*node2: Computing the derivative of func w wrt to sand for going the other functions*/
mistake1b:diff(func_w,s)</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>x,y,z,p</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mistake1b</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
