<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Chain Rule Application: Partial Derivatives with Respect to Parameters</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p></p>
<p>Given the following information</p>

<p>\(\displaystyle  y={@func_y@}\)</p>
<p>\(\displaystyle  z={@func_z@}\)</p>
<p>\(\displaystyle  x={@func_x@}\)</p>
<p>\(\displaystyle  u={@func_u@}\)</p>

<p> We can use Chain Rule to compute \( z_t\) and \( z_p\) using these two formulas, </p>

<p>\( \displaystyle  z_t = \frac{\partial{z}}{\partial{t}} = \frac{\partial{z}}{\partial{x}} \frac{\partial{x}}{\partial{u}} \frac{d{u}}{d{t}} + \frac{\partial{z}}{\partial{y}} \frac{\partial{y}}{\partial{u}} \frac{d{u}}{d{t}} \)</p>

<p>\( \displaystyle  z_p = \frac{\partial{z}}{\partial{p}} = \frac{\partial{z}}{\partial{x}} \frac{\partial{x}}{\partial{p}} + \frac{\partial{z}}{\partial{y}} \frac{\partial{y}}{\partial{p}}\)</p>

<p>\(\displaystyle  z_t = \)  [[input:ans1]] [[validation:ans1]] [[feedback:prt1]] </p>

<p>\( \displaystyle  z_p = \)  [[input:ans2]] [[validation:ans2]] [[feedback:prt2]] </p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>Worked solution</p>
<br>
<p>Here is how to compute the derivatives in this problem.</p>

<p>\(\displaystyle  z_t  = \frac{\partial{z}}{\partial{x}} \frac{\partial{x}}{\partial{u}} \frac{d{u}}{d{t}} + \frac{\partial{z}}{\partial{y}} \frac{\partial{y}}{\partial{u}} \frac{d{u}}{d{t}} \)</p>

<p>\( \displaystyle \quad= \left[ {@diff_1@} \right] \cdot \left[ {@diff_2@} \right] \cdot \left[ {@diff_3@} \right] + \left[ {@diff_4@} \right] \cdot \left[ {@diff_5@} \right] \cdot \left[ {@diff_6@} \right] \)  </p>
<p>\( \displaystyle \quad= {@ta1@}\)  </p>

<p>\( \displaystyle  z_p = \frac{\partial{z}}{\partial{x}} \frac{\partial{x}}{\partial{p}} + \frac{\partial{z}}{\partial{y}} \frac{\partial{y}}{\partial{p}}\)</p>
<p>\(\displaystyle \quad = \left[ {@diff_1@} \right] \cdot \left[ {@diff_7@} \right] + \left[ {@diff_4@} \right] \cdot \left[ {@diff_8@} \right] \)   </p>
<p>\(\displaystyle \quad = {@ta2@} \)   </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text>/*Question variables*/
aa:rand_with_prohib(-5,5,[0]);
bb:rand_with_prohib(-5,5,[0]);
cc:rand_with_prohib(-5,5,[0]);
dd:rand_with_prohib(-5,5,[0]);

/*Question func*/
func_z:aa*y*sin(bb*x);
func_x:cc*u-p;
func_y:p^(bb)*u;
func_u:expand((dd-t^cc)*(cc+t));

/*Worked Soln variables*/
diff_1:diff(func_z,x);diff_2:diff(func_x,u); diff_3:diff(func_u,t);
diff_4:diff(func_z,y);diff_5:diff(func_y,u); diff_6:diff(func_u,t);
diff_7:diff(func_x,p);diff_8:diff(func_y,p);

ta1:diff_1*diff_2*diff_3+diff_4*diff_5*diff_6;

ta2:diff_1*diff_7+diff_4*diff_8;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\displaystyle  y={@func_y@}\)</p>
<p>\(\displaystyle  z={@func_z@}\)</p>
<p>\(\displaystyle  x={@func_x@}\)</p>
<p>\(\displaystyle  u={@func_u@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1: checking for correct computation of the derivative*/
/*node2: checking if the student the derivative of func_z wrt to t and ignoring the others.*/
mistake1:diff(func_z,t);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x,y,z,t,u</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have computed the derivative of \(\displaystyle  z={@func_z@}\) wrt to \(t\) alone. Note that you need to use the formula given in the question text as a guide.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1:checking for correct answer*/
/*node 2: checking if the student computed the derivative of  z=4y sin(2x) wrt to p and ignoring the others. */
mistake22:diff(func_z,p);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Diff</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>x,y,z,p</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mistake22</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You need to refer to the formula given in the question text to guide you on how to compute the second derivative, just like the first one.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
