<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Combining Logarithmic Expressions</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Express the given quantities as a single logarithmic</p>
<p>a) \({@a1@} \ln({@a2@})-\ln({@a3@})\) can be expressed as [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
       
<p>b)  \(\ln(x)+a \ln(y)-b \ln(z)\) can be expressed as [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
       
<p>c)  \(\ln(1) +x^2+ {@b2@} \ln(x) - \ln({@b1@})\) can be expressed as [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p> \(\text{(a)}\\ \begin{align} {@a1@} \ln({@a2@})-\ln({@a3@}) & \Longleftrightarrow \ln ({@a2@}^{@a1@})-\ln({@a3@})\\ & \Longleftrightarrow \ln ({@a2^a1/a3@}) \end{align} \) </p>
<p> \(\text{(b)}\\ \begin{align} \ln(x)+a \ln(y)-b \ln(z) & \Longleftrightarrow \ln (x) ( \cdot y^a)-z^b \\ & \Longleftrightarrow \ln (x)+\ln (y^a)- \ln (z^b)\\ & \Longleftrightarrow \ln \bigg(\frac{x\cdot y^a}{z^b} \bigg). \end{align} \) </p>
<p>\(\displaystyle \text{(c)}\\ \begin{align} \ln(1) +x^2+ {@b2@} \ln(x) - \ln({@b1@}) & \Longleftrightarrow x^2+\ln ({@x^b2@})- \ln({@b1@})\\  & \Longleftrightarrow \displaystyle x^2+\ln \bigg(\frac{x^{@b2@}}{{@b1@}} \bigg) \end{align}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Question Variables*/
a1:rand(5)+2;
a2:rand(5)+1;
a3:rand(5)+1;
func_1:a1*ln(a2)-ln(a3);
ta1:logcontract(func_1);

func_2:ln(x)+a*ln(y)-b*ln(z);
declare(a,integer);
declare(b,integer);
ta2:logcontract(func_2);

b1:rand([sin(x),cos(x),tan(x)]);
trig_1:if b1=sin(x) then csc(x) elseif b1=cos(x) then sec(x) else cot(x);
b2:1/(rand(6)+2);
func_3: ln(1) + b2*ln(x) - ln(b1);
ta3:x^2+ln(x^b2*trig_1);
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>a) \(\displaystyle {@a1@} \ln({@a2@})-\ln({@a3@}) \Longleftrightarrow   {@ta1@}\)</p>    
<p>b)  \(\displaystyle\ln(x)+a \ln(y)-b \ln(z) \Longleftrightarrow   {@ta2@}\)</p>      
<p>c)  \(\displaystyle\ln(1) + {@b2@} \ln(x) - \ln({@b1@}) \Longleftrightarrow   {@ta3@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">https://docs.google.com/document/d/1PBa24xyFQMwNxb7CrGB0EUtQ2BZMXZGezjcxD0Jx3Wg/edit#heading=h.kog50oi25ikx/p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">[Reviewer 2 Name]</p>
<p id="description">This question involves combining multiple logarithmic expressions into a single logarithmic expression using the properties of logarithms. Each expression requires applying properties such as the quotient rule, product rule, and power rule to condense the given expressions into a simplified form.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1471366207</deployedseed>
    <deployedseed>2063304786</deployedseed>
    <deployedseed>199377837</deployedseed>
    <deployedseed>202377577</deployedseed>
    <deployedseed>801198327</deployedseed>
    <deployedseed>1516477260</deployedseed>
    <deployedseed>1289687139</deployedseed>
    <deployedseed>855172669</deployedseed>
    <deployedseed>1826143672</deployedseed>
    <deployedseed>168473200</deployedseed>
    <deployedseed>678326890</deployedseed>
    <deployedseed>246038888</deployedseed>
    <deployedseed>7298531</deployedseed>
    <deployedseed>648524891</deployedseed>
    <deployedseed>643651988</deployedseed>
    <deployedseed>21492514</deployedseed>
    <deployedseed>1845219977</deployedseed>
    <deployedseed>1297031511</deployedseed>
    <deployedseed>2047247649</deployedseed>
    <deployedseed>1221923249</deployedseed>
    <deployedseed>158196859</deployedseed>
    <deployedseed>1600553019</deployedseed>
    <deployedseed>658983399</deployedseed>
    <deployedseed>1780241819</deployedseed>
    <deployedseed>1700517726</deployedseed>
    <deployedseed>309743043</deployedseed>
    <deployedseed>2003150523</deployedseed>
    <deployedseed>742192019</deployedseed>
    <deployedseed>1101459121</deployedseed>
    <deployedseed>99594061</deployedseed>
    <deployedseed>2133778038</deployedseed>
    <deployedseed>1445960426</deployedseed>
    <deployedseed>249829345</deployedseed>
    <deployedseed>99930692</deployedseed>
    <deployedseed>1339863081</deployedseed>
    <deployedseed>669682056</deployedseed>
    <deployedseed>1482341507</deployedseed>
    <deployedseed>815654101</deployedseed>
    <deployedseed>14497759</deployedseed>
    <deployedseed>457054695</deployedseed>
    <deployedseed>1836973385</deployedseed>
    <deployedseed>1838913495</deployedseed>
    <deployedseed>58247884</deployedseed>
    <deployedseed>85603871</deployedseed>
    <deployedseed>50683513</deployedseed>
    <deployedseed>128638021</deployedseed>
    <deployedseed>451656730</deployedseed>
    <deployedseed>1765464941</deployedseed>
    <deployedseed>484931455</deployedseed>
    <deployedseed>1735495941</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
