<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Comparison between the functions f(x) and g(x)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Compare the functions \(f(x)={@fx@}\) and \(g(x)={@gx@}\) using the graph below.</p>

<p>
[[jsxgraph width="500px" height="500px"]]  
            var board = JXG.JSXGraph.initBoard(divid, { 
            boundingbox: [-2, 5, 2, -2],
            keepaspectratio: false,
            showCopyright: false,
            axis: true,
            showNavigation: false
        });


// Define functions
function f(x) {
    return Math.pow(x, {#power#}); // f(x)=x^10 
}

function g(x) {
    return Math.exp(x); // g(x)=e^x
}

// Plot function curves
//var curveF = board.create('functiongraph', [f, -2,  {#power#}], { strokeColor: 'blue', strokeWidth: 2, name: '\\(f(x) = x^{10}\\)' });
//var curveG = board.create('functiongraph', [g, -5,  {#power#}], { strokeColor: 'red', strokeWidth: 2, name: '\\(g(x) = e^x\\)' });
// Plot function curves
var curveF = board.create('functiongraph', [f, -2,  {#power#}], {
    strokeColor: 'blue',
    strokeWidth: 2,
    name: 'f(x) = x^{ {#power#}}',
    withLabel: true,
    label: {
        position: 'rt',
        offset: [-10, 10]
    }
});
var curveG = board.create('functiongraph', [g, -2,  {#power#}], {
    strokeColor: 'red',
    strokeWidth: 2,
    name: 'g(x) = e^x',
    withLabel: true,
    label: {
        position: 'rt',
        offset: [10, 10]
    }
});
// Find intersection points
var intersection1 = board.create('intersection', [curveF, curveG, 0], { name: 'Intersection 1', color: 'green', size: 4 });
var intersection2 = board.create('intersection', [curveF, curveG, 1], { name: 'Intersection 1', color: 'purple', size: 4 });

// Dotted lines for intersection 1
var xLine1 = board.create('line', [[intersection1.X(), 0], intersection1], {
    strokeColor: 'green',
    dash: 2,
    straightFirst: false,
    straightLast: false
});

var yLine1 = board.create('line', [[0, intersection1.Y()], intersection1], {
    strokeColor: 'green',
    dash: 2,
    straightFirst: false,
    straightLast: false
});

// Dotted lines for intersection 2
var xLine2 = board.create('line', [[intersection2.X(), 0], intersection2], {
    strokeColor: 'purple',
    dash: 2,
    straightFirst: false,
    straightLast: false
});

var yLine2 = board.create('line', [[0, intersection2.Y()], intersection2], {
    strokeColor: 'purple',
    dash: 2,
    straightFirst: false,
    straightLast: false
});
// Calculate coordinates of intersection points
var coords1 = "\\((" + intersection1.coords.usrCoords[1].toFixed(2) + ", " + intersection1.coords.usrCoords[2].toFixed(2) + ")\\)";
var coords2 = "\\((" + intersection2.coords.usrCoords[1].toFixed(2) + ", " + intersection2.coords.usrCoords[2].toFixed(2) + ")\\)";

// Update the names of intersection points to include coordinates
intersection1.setAttribute({ name: ''  });
intersection2.setAttribute({ name: ''  });
 
[[/jsxgraph]]
</p>
<p>Based on the graph, when does the graph of \(g\) surpass the graph of \(f\)?</p>
<!--<p>Observing from the graph, when does the graph of \(g\) surpass the graph of \(f\)?</p>-->
<p><strong>Hint:</strong> Give your answer in the form \((x, y)\) and correct to \(2\) decimal places.</p>
<p>\(\big( x, y \big)=\) \(\big(\) [[input:ans1]] \(,\) [[input:ans2]] \(\big)\) [[validation:ans1]]  [[validation:ans2]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p> \(f(x)={@fx@}\) and \(g(x)={@gx@}\) as shown on the graph below.</p>
<p>
[[jsxgraph width="500px" height="500px"]]  
            var board = JXG.JSXGraph.initBoard(divid, { 
            boundingbox: [-2, 5, 2, -2],
            keepaspectratio: false,
            showCopyright: false,
            axis: true,
            showNavigation: false
        });


// Define functions
function f(x) {
    return Math.pow(x, {#power#}); // f(x)=x^10 
}

function g(x) {
    return Math.exp(x); // g(x)=e^x
}

// Plot function curves
//var curveF = board.create('functiongraph', [f, -2,  {#power#}], { strokeColor: 'blue', strokeWidth: 2, name: '\\(f(x) = x^{10}\\)' });
//var curveG = board.create('functiongraph', [g, -5,  {#power#}], { strokeColor: 'red', strokeWidth: 2, name: '\\(g(x) = e^x\\)' });
// Plot function curves
var curveF = board.create('functiongraph', [f, -2,  {#power#}], {
    strokeColor: 'blue',
    strokeWidth: 2,
    name: 'f(x) = x^{ {#power#}}',
    withLabel: true,
    label: {
        position: 'rt',
        offset: [-10, 10]
    }
});
var curveG = board.create('functiongraph', [g, -2,  {#power#}], {
    strokeColor: 'red',
    strokeWidth: 2,
    name: 'g(x) = e^x',
    withLabel: true,
    label: {
        position: 'rt',
        offset: [10, 10]
    }
});
// Find intersection points
var intersection1 = board.create('intersection', [curveF, curveG, 0], { name: 'Intersection 1', color: 'green', size: 4 });
var intersection2 = board.create('intersection', [curveF, curveG, 1], { name: 'Intersection 2', color: 'purple', size: 4 });

// Dotted lines for intersection 1
var xLine1 = board.create('line', [[intersection1.X(), 0], intersection1], {
    strokeColor: 'green',
    dash: 2,
    straightFirst: false,
    straightLast: false
});

var yLine1 = board.create('line', [[0, intersection1.Y()], intersection1], {
    strokeColor: 'green',
    dash: 2,
    straightFirst: false,
    straightLast: false
});

// Dotted lines for intersection 2
var xLine2 = board.create('line', [[intersection2.X(), 0], intersection2], {
    strokeColor: 'purple',
    dash: 2,
    straightFirst: false,
    straightLast: false
});

var yLine2 = board.create('line', [[0, intersection2.Y()], intersection2], {
    strokeColor: 'purple',
    dash: 2,
    straightFirst: false,
    straightLast: false
});
// Calculate coordinates of intersection points
var coords1 = "\\((" + intersection1.coords.usrCoords[1].toFixed(2) + ", " + intersection1.coords.usrCoords[2].toFixed(2) + ")\\)";
var coords2 = "\\((" + intersection2.coords.usrCoords[1].toFixed(2) + ", " + intersection2.coords.usrCoords[2].toFixed(2) + ")\\)";

// Update the names of intersection points to include coordinates
intersection1.setAttribute({ name: ' ' + coords1 });
intersection2.setAttribute({ name: '' + coords2 });
 
[[/jsxgraph]]
</p>

<p>Observing from the graph, the graph of \(g\) surpasses the graph of \(f\) at \( \big( {@x_cord@},{@y_cord@} \big)\)]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>power:rand(14)+4;
fx:x^power;
gx:e^x;
/* Define the intersection points */
 point: if power = 4 then [1.43, 4.18]
    elseif power = 5 then [1.30, 3.65]
    elseif power = 6 then [1.23, 3.41]
    elseif power = 7 then [1.18, 3.27]  
    elseif power = 8 then [1.16, 3.18] 
    elseif power = 9 then [1.13, 3.11]  
    elseif power = 10 then [1.12, 3.06] 
    elseif power = 11 then [1.11, 3.02]
    elseif power = 12 then [1.10, 2.99] 
    elseif power = 13 then [1.09, 2.97] 
    elseif power = 14 then [1.08, 2.95] 
    elseif power = 15 then [1.07, 2.93]
    elseif power = 16 then [1.07, 2.91]
    elseif power = 17 then [1.06, 2.90];
x_cord:point[1];
y_cord:point[2];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \(f(x)={@fx@}\) and \(g(x)={@gx@}\) </p>
<p>Surpassing point is at \( (x,y)= \big( {@x_cord@},{@y_cord@} \big)\)]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p><p>
</p><p id="concept">https://docs.google.com/document/d/1Aq5zI6mExrgmLzgAHhfPRJEzQ8nW2d-_aLYPoWpg8rw/edit#heading=h.eddp73mos3v3</p>
<p id="reviewer">Emmaculate odhiambo</p>
<p id="reviewer">[2]</p>
<p id="description">Given functions f(x) and g(x), students are required to compare the two functions by graphing them and show when the graph \(g\) surpasses \(f\). </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>x_cord</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>y_cord</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct reading of x coordinate</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>x_cord</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for correct reading of y coordinate</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>y_cord</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks for the value of x entered twice</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ans2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You entered the same answer twice. The point has two distinct coordinates</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>checks for the value of y entered twice</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ans1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You entered the same answer twice. The point has two distinct coordinates</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1884342440</deployedseed>
    <deployedseed>82248245</deployedseed>
    <deployedseed>1382018075</deployedseed>
    <deployedseed>1159139251</deployedseed>
    <deployedseed>1392210523</deployedseed>
    <deployedseed>1434283205</deployedseed>
    <deployedseed>1265226300</deployedseed>
    <deployedseed>1077138126</deployedseed>
    <deployedseed>411377584</deployedseed>
    <deployedseed>1931319399</deployedseed>
    <deployedseed>2069175710</deployedseed>
    <deployedseed>1368294195</deployedseed>
    <deployedseed>662941526</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>x_cord</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>y_cord</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-6-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
