<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Derivatives of the six trig inverse functions</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the derivative of the following function:</p>
<p>\(\quad \displaystyle f(x)={@fx@}\)</p>
<p>\(\quad \displaystyle f'(x)=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>The derivative of the function \(  \displaystyle f(x)={@fx@}\) with respect to \(x\) is:</p>
<p>\( \qquad  \displaystyle f'(x)={@diff_fx@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>trig:rand([asin,acos,atan,acot,asec,acsc]);
coeff:rand_with_prohib(-5,5,[0]);
fx:trig(coeff*x);
diff_fx : diff(fx, x);
ta:diff_fx;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\quad \displaystyle f(x)={@fx@}\)</p>
<p>\(\quad \displaystyle f'(x)={@diff_fx@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">[Godfrey Wabwire]</p>
<p id="reviewer">[Juma Zevick]</p>
<p id="description">Derivatives of the six trig inverse functions </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct derivative */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct derivative </description>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>41704532</deployedseed>
    <deployedseed>765195316</deployedseed>
    <deployedseed>657401001</deployedseed>
    <deployedseed>1817115699</deployedseed>
    <deployedseed>366564831</deployedseed>
    <deployedseed>853184935</deployedseed>
    <deployedseed>1814483954</deployedseed>
    <deployedseed>75450095</deployedseed>
    <deployedseed>995407985</deployedseed>
    <deployedseed>139646577</deployedseed>
    <deployedseed>1684940590</deployedseed>
    <deployedseed>871914946</deployedseed>
    <deployedseed>766112923</deployedseed>
    <deployedseed>2079782330</deployedseed>
    <deployedseed>221723978</deployedseed>
    <deployedseed>786665747</deployedseed>
    <deployedseed>1623659311</deployedseed>
    <deployedseed>1547967329</deployedseed>
    <deployedseed>513091791</deployedseed>
    <deployedseed>1416999681</deployedseed>
    <deployedseed>1132914302</deployedseed>
    <deployedseed>1704913004</deployedseed>
    <deployedseed>566966573</deployedseed>
    <deployedseed>1757464497</deployedseed>
    <deployedseed>1815485581</deployedseed>
    <deployedseed>558467173</deployedseed>
    <deployedseed>1955577042</deployedseed>
    <deployedseed>1037919942</deployedseed>
    <deployedseed>1749739059</deployedseed>
    <deployedseed>977258098</deployedseed>
    <deployedseed>503540803</deployedseed>
    <deployedseed>1205727692</deployedseed>
    <deployedseed>800277446</deployedseed>
    <deployedseed>1617293813</deployedseed>
    <deployedseed>1913037506</deployedseed>
    <deployedseed>1831386616</deployedseed>
    <deployedseed>322418409</deployedseed>
    <deployedseed>804259950</deployedseed>
    <deployedseed>273428066</deployedseed>
    <deployedseed>2084032302</deployedseed>
    <deployedseed>1479757898</deployedseed>
    <deployedseed>1368153484</deployedseed>
    <deployedseed>1905198001</deployedseed>
    <deployedseed>578113550</deployedseed>
    <deployedseed>684312012</deployedseed>
    <deployedseed>443485430</deployedseed>
    <deployedseed>1444603316</deployedseed>
    <deployedseed>1067552990</deployedseed>
    <deployedseed>360232943</deployedseed>
    <deployedseed>558795306</deployedseed>
    <deployedseed>803546496</deployedseed>
    <deployedseed>1593169189</deployedseed>
    <deployedseed>1812710528</deployedseed>
    <deployedseed>1555695874</deployedseed>
    <deployedseed>423979111</deployedseed>
    <deployedseed>234120311</deployedseed>
    <deployedseed>2110309241</deployedseed>
    <deployedseed>838649797</deployedseed>
    <deployedseed>1346744239</deployedseed>
    <deployedseed>949219823</deployedseed>
    <deployedseed>600933746</deployedseed>
    <deployedseed>1112792886</deployedseed>
    <deployedseed>1613229196</deployedseed>
    <deployedseed>2031880298</deployedseed>
    <deployedseed>20784162</deployedseed>
    <deployedseed>1520657298</deployedseed>
    <deployedseed>474826832</deployedseed>
    <deployedseed>1028524079</deployedseed>
    <deployedseed>963908585</deployedseed>
    <deployedseed>1653950885</deployedseed>
    <deployedseed>1652475960</deployedseed>
    <deployedseed>1893976959</deployedseed>
    <deployedseed>1823762627</deployedseed>
    <deployedseed>91632781</deployedseed>
    <deployedseed>531271290</deployedseed>
    <deployedseed>1769755771</deployedseed>
    <deployedseed>1724366419</deployedseed>
    <deployedseed>806608350</deployedseed>
    <deployedseed>1440767533</deployedseed>
    <deployedseed>35800807</deployedseed>
    <deployedseed>1659866655</deployedseed>
    <deployedseed>917905331</deployedseed>
    <deployedseed>541081710</deployedseed>
    <deployedseed>280317543</deployedseed>
    <deployedseed>179277017</deployedseed>
    <deployedseed>149849137</deployedseed>
    <deployedseed>907491659</deployedseed>
    <deployedseed>2111877961</deployedseed>
    <deployedseed>1541003729</deployedseed>
    <deployedseed>1149496658</deployedseed>
    <deployedseed>2024927057</deployedseed>
    <deployedseed>1226925071</deployedseed>
    <deployedseed>1007320893</deployedseed>
    <deployedseed>1751962559</deployedseed>
    <deployedseed>1067360248</deployedseed>
    <deployedseed>767757867</deployedseed>
    <deployedseed>867304889</deployedseed>
    <deployedseed>427504484</deployedseed>
    <deployedseed>1391664568</deployedseed>
    <deployedseed>975527651</deployedseed>
    <deployedseed>570590005</deployedseed>
    <deployedseed>1567430234</deployedseed>
    <deployedseed>295677082</deployedseed>
    <deployedseed>1450119073</deployedseed>
    <deployedseed>1046322402</deployedseed>
    <deployedseed>115576250</deployedseed>
    <deployedseed>992892061</deployedseed>
    <deployedseed>862805869</deployedseed>
    <deployedseed>372019980</deployedseed>
    <deployedseed>1814266603</deployedseed>
    <deployedseed>462976447</deployedseed>
    <deployedseed>1461659341</deployedseed>
    <deployedseed>298114374</deployedseed>
    <deployedseed>1936560859</deployedseed>
    <deployedseed>15735437</deployedseed>
    <deployedseed>1147339801</deployedseed>
    <deployedseed>1180351829</deployedseed>
    <deployedseed>892128718</deployedseed>
    <deployedseed>1435617091</deployedseed>
    <deployedseed>1754731363</deployedseed>
    <deployedseed>463991308</deployedseed>
    <deployedseed>1898681915</deployedseed>
    <deployedseed>1640502045</deployedseed>
    <deployedseed>600102862</deployedseed>
    <deployedseed>1984173495</deployedseed>
    <deployedseed>1860887207</deployedseed>
    <deployedseed>1428825169</deployedseed>
    <deployedseed>701332306</deployedseed>
    <deployedseed>943207398</deployedseed>
    <deployedseed>53108267</deployedseed>
    <deployedseed>714579539</deployedseed>
    <deployedseed>79357800</deployedseed>
    <deployedseed>1067117293</deployedseed>
    <deployedseed>1212509782</deployedseed>
    <deployedseed>1514261749</deployedseed>
    <deployedseed>1914945569</deployedseed>
    <deployedseed>915287248</deployedseed>
    <deployedseed>1595526489</deployedseed>
    <deployedseed>1766461413</deployedseed>
    <deployedseed>373148248</deployedseed>
    <deployedseed>114562685</deployedseed>
    <deployedseed>1525957086</deployedseed>
    <deployedseed>198786347</deployedseed>
    <deployedseed>1902576910</deployedseed>
    <deployedseed>1348409479</deployedseed>
    <deployedseed>880490356</deployedseed>
    <deployedseed>1169397525</deployedseed>
    <deployedseed>211846692</deployedseed>
    <deployedseed>238677967</deployedseed>
    <deployedseed>1036034928</deployedseed>
    <deployedseed>359743347</deployedseed>
    <deployedseed>436999794</deployedseed>
    <deployedseed>1621574934</deployedseed>
    <deployedseed>241907111</deployedseed>
    <deployedseed>808235352</deployedseed>
    <deployedseed>1786766856</deployedseed>
    <deployedseed>779781494</deployedseed>
    <deployedseed>1055537450</deployedseed>
    <deployedseed>119737041</deployedseed>
    <deployedseed>700223193</deployedseed>
    <deployedseed>526555558</deployedseed>
    <deployedseed>1741805018</deployedseed>
    <deployedseed>1831046912</deployedseed>
    <deployedseed>1413706264</deployedseed>
    <deployedseed>115005376</deployedseed>
    <deployedseed>297523195</deployedseed>
    <deployedseed>734965035</deployedseed>
    <deployedseed>725788479</deployedseed>
    <deployedseed>440935207</deployedseed>
    <deployedseed>273601365</deployedseed>
    <deployedseed>1151319032</deployedseed>
    <deployedseed>1206515513</deployedseed>
    <deployedseed>330309647</deployedseed>
    <deployedseed>304200355</deployedseed>
    <deployedseed>1800140049</deployedseed>
    <deployedseed>217221584</deployedseed>
    <deployedseed>555939437</deployedseed>
    <deployedseed>342650575</deployedseed>
    <deployedseed>1401335658</deployedseed>
    <deployedseed>1617246269</deployedseed>
    <deployedseed>618664203</deployedseed>
    <deployedseed>1486740039</deployedseed>
    <deployedseed>880643835</deployedseed>
    <deployedseed>1700760498</deployedseed>
    <deployedseed>146921185</deployedseed>
    <deployedseed>135889462</deployedseed>
    <deployedseed>1660397252</deployedseed>
    <deployedseed>1070095967</deployedseed>
    <deployedseed>680403084</deployedseed>
    <deployedseed>1973449913</deployedseed>
    <deployedseed>1015889934</deployedseed>
    <deployedseed>234387851</deployedseed>
    <deployedseed>1794482616</deployedseed>
    <deployedseed>1301440585</deployedseed>
    <deployedseed>1320622299</deployedseed>
    <deployedseed>551218783</deployedseed>
    <deployedseed>288905449</deployedseed>
    <deployedseed>1572493629</deployedseed>
    <deployedseed>211372073</deployedseed>
    <deployedseed>948702125</deployedseed>
    <deployedseed>2080209512</deployedseed>
    <deployedseed>2024603783</deployedseed>
    <deployedseed>1294993723</deployedseed>
    <deployedseed>1227089342</deployedseed>
    <deployedseed>1177402349</deployedseed>
    <deployedseed>1291460822</deployedseed>
    <deployedseed>2000829527</deployedseed>
    <deployedseed>1305689963</deployedseed>
    <deployedseed>572931904</deployedseed>
    <deployedseed>92717452</deployedseed>
    <deployedseed>1299752538</deployedseed>
    <deployedseed>22838017</deployedseed>
    <deployedseed>1057759422</deployedseed>
    <deployedseed>866683100</deployedseed>
    <deployedseed>1463567120</deployedseed>
    <deployedseed>1628239806</deployedseed>
    <deployedseed>884509006</deployedseed>
    <deployedseed>1673727114</deployedseed>
    <deployedseed>2126323931</deployedseed>
    <deployedseed>1475829945</deployedseed>
    <deployedseed>932684773</deployedseed>
    <deployedseed>1848994721</deployedseed>
    <deployedseed>1060626495</deployedseed>
    <deployedseed>1090267567</deployedseed>
    <deployedseed>1967861765</deployedseed>
    <deployedseed>917572229</deployedseed>
    <deployedseed>174336693</deployedseed>
    <deployedseed>926894036</deployedseed>
    <deployedseed>601638838</deployedseed>
    <deployedseed>1717123147</deployedseed>
    <deployedseed>1737025204</deployedseed>
    <deployedseed>427901422</deployedseed>
    <deployedseed>1629780995</deployedseed>
    <deployedseed>1920224389</deployedseed>
    <deployedseed>1549009559</deployedseed>
    <deployedseed>954941590</deployedseed>
    <deployedseed>1566079100</deployedseed>
    <deployedseed>1625588214</deployedseed>
    <deployedseed>409997225</deployedseed>
    <deployedseed>2127714053</deployedseed>
    <deployedseed>1660864045</deployedseed>
    <deployedseed>1017755384</deployedseed>
    <deployedseed>451805605</deployedseed>
    <deployedseed>242247032</deployedseed>
    <deployedseed>2108700223</deployedseed>
    <deployedseed>158824005</deployedseed>
    <deployedseed>2058947179</deployedseed>
    <deployedseed>1943263157</deployedseed>
    <deployedseed>1379036183</deployedseed>
    <deployedseed>1097307312</deployedseed>
    <deployedseed>1595320461</deployedseed>
    <deployedseed>1861874786</deployedseed>
    <deployedseed>1672265344</deployedseed>
    <deployedseed>326411112</deployedseed>
    <deployedseed>352510856</deployedseed>
    <deployedseed>218704853</deployedseed>
    <deployedseed>2122661031</deployedseed>
    <deployedseed>36068818</deployedseed>
    <deployedseed>1720832585</deployedseed>
    <deployedseed>2037317147</deployedseed>
    <deployedseed>921703731</deployedseed>
    <deployedseed>110767509</deployedseed>
    <deployedseed>1686432753</deployedseed>
    <deployedseed>1136729028</deployedseed>
    <deployedseed>1013986099</deployedseed>
    <deployedseed>1769558308</deployedseed>
    <deployedseed>432982805</deployedseed>
    <deployedseed>189464082</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
