<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Derivative of a Function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the derivative of \(y =\displaystyle {@fx@}\).</p>
<p><em>Note: To enter the natural logarithm, \(\ln(x)\), you should write <code>log(x)</code>.</em></p>
<p>\(\displaystyle y' = \) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3> Worked Solution</h3>

<p>Here we need to use logarithmic differentiation. We first take natural logarithms on both sides and use the rules of logarithms to simplify so that we can differentiate implicitly.</p>

<p style="margin-left: 30px">\(\begin{align*}
y &= ({@ins@})^{\frac{1}{x}} \\
\ln(y) &= \ln(({@ins@})^{\frac{1}{x}}) \\
&= \frac{1}{x} \ln({@ins@})
\end{align*}\)</p>

<br>

<p>Differentiating implicitly on both sides,</p>

<p style="margin-left: 30px">\(\begin{align*}
\displaystyle \frac{1}{y} y' &= {@-1/x^2@} \ln({@ins@}) + {@1/x@} \frac{{@2*aa*x@}}{{@ins@}} \\
&= {@-1/x^2*log(ins) + 1/x*2*aa*x/ins@}
\end{align*}\)</p>

<br>

<p>Substituting back the expression for \(y\) we have,</p>

<p style="margin-left: 30px">\(\displaystyle \frac{1}{({@ins@})^{\frac{1}{x}}} y' = {@-1/x^2*log(ins) + 1/x*2*aa*x/ins@}\)</p>

<br>

<p>Therefore, \(\displaystyle y' =&nbsp;({@ins@})^{\frac{1}{x}}  \left( {@-1/x^2*log(ins) + 1/x*2*aa*x/ins@}\right) \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(4)+2;
bb:rand(8)+1;
ins:aa*x^2+bb;
fx:ins^(1/x);
ta:fx*(-1/x^2*log(ins) + 1/x*2*aa*x/ins);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>The derivative of \( y = {@fx@}\) is \(  {@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">This question the student to find the derivative of the function \(y\) with respect to \(x\) when \(y\) is defined as the given mathematical expression.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1:Checks for correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1676325515</deployedseed>
    <deployedseed>777354629</deployedseed>
    <deployedseed>1693356509</deployedseed>
    <deployedseed>271276803</deployedseed>
    <deployedseed>1084067829</deployedseed>
    <deployedseed>558623153</deployedseed>
    <deployedseed>1705339334</deployedseed>
    <deployedseed>401534680</deployedseed>
    <deployedseed>2005644275</deployedseed>
    <deployedseed>2117969710</deployedseed>
    <deployedseed>1582394307</deployedseed>
    <deployedseed>797177687</deployedseed>
    <deployedseed>1456497264</deployedseed>
    <deployedseed>1352690767</deployedseed>
    <deployedseed>1881965640</deployedseed>
    <deployedseed>1879668100</deployedseed>
    <deployedseed>29650994</deployedseed>
    <deployedseed>1568291823</deployedseed>
    <deployedseed>1387366026</deployedseed>
    <deployedseed>68839071</deployedseed>
    <deployedseed>443330379</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
