<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Derivative of log and trig function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p> Find the  derivatives for the following function
<ol>
<li>
<p> \(f(x)={@quest_expr@}\)</p>
<p>\( f'(x) =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]] </p>
</li>

<li>
<p> \( g(y)= {@quest_exprresion@}\)</p>
<p> \(g'(y)= \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]] </p>
</li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<ol>
<li>
<p>Find the derivatives for the following function.</p>
<p>\(f(x)={@quest_expr@}\)</p>
<p>\(f'(x)={@coef_cos@} \cdot [{@cos(a*x)@}]'+{@coef_sin@} \cdot [{@sin(b*x)@}]'\)</p>
<p>\( f'(x) = {@coef_cos@} \cdot \left[ {@diff(cos(a*x),x)@} \right]+{@coef_sin@} \cdot \left[ {@diff(sin(b*x),x)@} \right] \) </p>
<p>\( f'(x)= \displaystyle {@ta1@}\)</p>
</li>

<li>
<p> Find the derivative for the following function.</p>
<p>\(  g(y)= \displaystyle {@quest_exprresion@}\)</p>
<p> We use the formula \(\displaystyle [\ln(u(y))]'= \displaystyle \frac{u'(y)}{u(y)} \)</p>
<p>Here \(u(y)={@u_y@}\) and \(u'(y)= \displaystyle {@diff(u_y,y)@}\)</p>
<p>Therefore \(g'(y)= \displaystyle {@ta2@} \)</p>
</li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*part a*/
coef_cos:rand(2)+2;
coef_sin:rand(4)+2;
a: rand_with_step(1,5,1);
b: rand_with_prohib(1,5,[a]);
quest_expr:coef_cos*cos(a*x)+coef_sin*sin(b*x);
ta1: diff(quest_expr,x);

/*part b*/
u_y: rand(3)*y^2+rand_with_step(-5,5,1)*y+rand_with_step(-5,5,1);
quest_exprresion: ln(u_y);
ta2: diff(quest_exprresion,y);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(f(x)={@quest_expr@}\)</p>
<p>\( f'(x)= {@ta1@}\)</p>
<p> \( g(y)= {@quest_exprresion@}\)</p>
<p>\(g'(y)= {@ta2@} \)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">[Godfrey Wabwire]</p>
<p id="reviewer">[Juma Zevick]</p>
<p id="description">Deffirentaiting a log and trig function</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Check correct answers*/
/*Node 2: Instead of derivative, the student does anti-derivative*/
mistake_1: integrate(quest_expr,x);
/*Node 3: Used the variable y, instead of x*/
mistake_2: ev(ta1,[x=y]);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check correct answers</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Instead of derivative, the student does anti-derivative</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>You did the anti-derivative, instead of the derivative. Check the worked solution!</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Used the variable y, instead of x</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>You used the variable \(y\) instead of using \(x\), which should not the case</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Check correct answers*/
/*Node 2: Instead of derivative, the student does anti-derivative*/
mistake_1: integrate(quest_exprresion,x);
/*Node 3: Used the variable y, instead of x*/
mistake_2: ev(ta2,[x=y]);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check correct answers</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Instead of derivative, the student does anti-derivative</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mistake_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text>You did the anti-derivative, instead of the derivative. Check the worked solution!</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Used the variable y, instead of x</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text>You used the variable \(y\) instead of using \(x\), which should not the case</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1687622038</deployedseed>
    <deployedseed>1104966487</deployedseed>
    <deployedseed>948740387</deployedseed>
    <deployedseed>1019569081</deployedseed>
    <deployedseed>1015081870</deployedseed>
    <deployedseed>1362573860</deployedseed>
    <deployedseed>1088662560</deployedseed>
    <deployedseed>1966590639</deployedseed>
    <deployedseed>1438072150</deployedseed>
    <deployedseed>230856598</deployedseed>
    <deployedseed>1272961582</deployedseed>
    <deployedseed>836036246</deployedseed>
    <deployedseed>1950871834</deployedseed>
    <deployedseed>1632026290</deployedseed>
    <deployedseed>742482</deployedseed>
    <deployedseed>1907987268</deployedseed>
    <deployedseed>148451919</deployedseed>
    <deployedseed>1232108177</deployedseed>
    <deployedseed>1136846067</deployedseed>
    <deployedseed>1096212683</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
