<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Determining the Existence of a Limit for a Fractional Function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[</p>Determine whether the limit of the function below exist or not.</p>
<p>Let \(h(x,y)= \frac{@expr_numerator@}{@expr_denominator@} \)</p>
<p><b>Recall: </b>&nbsp; \( \lim_{(x,y) \longrightarrow (a,b)} \frac{f(x,y)}{g(x,y)} \) ( \( a, b \in \mathbb{R} \) ) exists when 
 \( \lim_{(x,y) \longrightarrow (a,b)} g(x, y) \ne 0 \).</p>

<br> Do you think that \( \lim_{(x,y) \longrightarrow ({@value_of_x@}, {@value_of_y@})} h(x, y) \) exists?.
<button id="btn1yes" type="button" style="font-size:16px" onclick="reveal('1');">Yes</button>
 <button id="btn1no" type="button" style="font-size:16px" onclick="hide('1');">No</button> [[feedback:prtyn1]]
<div id="input1" style="display:none;">[[input:ans_yn1]] [[validation:ans_yn1]]</div>
<br>
<div id="example1" style="display:none;">
    For:  \( \lim_{(x,y) \longrightarrow ({@value_of_x@}, {@value_of_y@})} \left({@expr_numerator@} \right) = \) [[input:ans1]] [[validation:ans1]] and \( \lim_{(x,y) \longrightarrow ({@value_of_x@}, {@value_of_y@})} \left({@expr_denominator@}\right) = \)[[input:ans2]] [[validation:ans2]].
    <br> From this we can conclude \( \lim_{(x,y) \longrightarrow ({@value_of_x@}, {@value_of_y@})} h(x, y)= \)  [[input:ans3]]  [[validation:ans3]]
<br> [[feedback:prt1]]
</div>



<script>
    function reveal(s) {
        document.querySelector("#input" + s).querySelector("input").value = "true";
        document.querySelector("#example" + s).style.display = "block";
        document.querySelector("#btn" + s + "yes").style["font-size"] = "24px";
        document.querySelector("#btn" + s + "no").style["font-size"] = "16px";
    }

    function hide(s) {
        document.querySelector("#input" + s).querySelector("input").value = "false";
        document.querySelector("#example" + s).style.display = "none";
        document.querySelector("#btn" + s + "yes").style["font-size"] = "16px";
        document.querySelector("#btn" + s + "no").style["font-size"] = "24px";
    }
</script>

<script>
    window.onload = function() {
        for (const s of ['1', '2']) {
            if (document.querySelector("#input" + s).querySelector("input").value == "true") {
                reveal(s);
            } else if (document.querySelector("#input" + s).querySelector("input").value == "false") {
                hide(s)
            }
        }
    };
</script>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(5)+1;bb:rand(5)+1;cc:rand(5)+1;
expr_numerator: aa*x^2*y+bb*x*y;
expr_denominator: 5*x*y^2+cc*y;

value_of_x: rand_with_step(-5,5,1);
value_of_y: rand_with_step(-5,5,1);
ta_1: at(expr_numerator, [x=value_of_x, y=value_of_y]);
ta_2: at(expr_denominator, [x=value_of_x, y=value_of_y]);
ta_3: ta_1/ta_2;
 
ta_yn1: if ta_2=0 then false else true;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@a@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans_yn1</name>
      <type>algebraic</type>
      <tans>ta_yn1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta_3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtyn1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_yn1</sans>
        <tans>ta_yn1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtyn1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtyn1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>48</deployedseed>
  </question>
</quiz>
