<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Determining  Arc Length-Q2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Determine the length of  the curve whose equation is given by \( \displaystyle {@expr@}\),  for \({@lower_y@}\leq y\leq{@upper_y@}\). Give your answer correct to two decimal places.</p>
<p>\(L=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Let's take the derivative of the given function x with respect to y, and then plug it into the ds formula.</p>
<p>\(\frac{dx}{dy}={@coef1*2@}({@const@}+y)\)</p>
<p>
\(
\begin{align}
ds &= \sqrt{1+[\frac{dx}{dy}]^{2}}dy \\
  &= \sqrt{1+[{@coef1*2@}({@const@}+y)]^{2}}dy \\
  &= \sqrt{1+{@(coef1*2)^2@}({@const@}+y)^{2}}dy \\
\end{align}
\)
</p>
<p>So, the integral giving the arc length is,</p>
<p>\[L= \int ds=\int_{{@lower_y@}}^{{@upper_y@}}\sqrt{1+{@(coef1*2)^2@}({@const@}+y)^{2}}dy\]</p>
<p>To evaluate this integral, we need to use a trig substitution</p>
<p>\(y={@1/(2*coef1)@}\tan\theta-{@const@}   \qquad \rightarrow dy={@diff(1/(2*coef1)*tan(theta)-const)@}\)
<p>\(y={@lower_y@}: \qquad {@const+lower_y@}={@1/(2*coef1)@}\tan\theta\) \(\qquad \tan\theta= {@1/(1/(2*coef1))*(const+lower_y)@}\) \(\rightarrow\) \( \theta\) = \(\tan^{-1}( {@1/(1/(2*coef1))*(const+lower_y)@})\) =\({@dd@}\)</p>
<p>\(y={@upper_y@}: \qquad {@const+upper_y@}={@1/(2*coef1)@}\tan\theta\) \(\qquad \tan\theta= {@1/(1/(2*coef1))*(const+upper_y)@}\) \(\rightarrow\)\(\ \theta\) = \(\tan^{-1}( {@1/(1/(2*coef1))*(const+upper_y)@})\) =\({@bb@}\)</p>
<p>So we have  is in the range \({@dd@} \leq \theta \leq {@bb@}\) . This means we are in the first quadrant, then a  secant is positive.</p>
<p>Now let’s deal with the square root,</p>
<p>
\(
\begin{align}
\sqrt{1+{@(coef1*2)^2@}({@const@}+y)^{2}} &  =  \sqrt{1+{@(coef1*2)^2@}({@1/(2*coef1)@}\tan(\theta))^{2}} \\
& = \sqrt{1+\tan^{2}(\theta)} \\
&= \sqrt{\sec^{2}(\theta)} \\
&=|\sec\theta| \\
&=\sec(\theta) \quad \text{since, } \sec\theta > 0 \\
\end{align}
\)
</p>
<p>Therefore,</p>
<p>
\(
\begin{align}
L&= \int_{@lower_y@}^{@upper_y@}\sqrt{1+{@(coef1*2)^2@}({@const@}+y)^{2}}dy\\
&={@1/(2*coef1)@} \int_{@dd@}^{@bb@}\sec^{3}(\theta) d(\theta)\\ 
&={@ta_1@}\\
\end{align}
\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Defining the question variables*/ 
coef1:rand(9)+1;/*coeffient of the function*/
const:rand(4)+1;/* constant inside the function*/
lower_y:rand(5)+2;/*lower lmit of y*/
upper_y:rand(3)+lower_y+2;/*upper lmit of y*/
aa:1/(2*coef1); /*coeff of tan*/
functy: coef1*(const+y)^(2);/*function of y*/
expr:f(y)=coef1*(const+y)^(2); /*the function*/
dx:rhs(diff(expr)/del(y)); /*dx/dy*/
ds:sqrt(1+(dx)^2);
func:y=(1/(2*coef1))*tan(theta)-const;   /*function of y*/
diff_func:diff(rhs(func),theta); /*derivative of y*/
/*new limits of theta*/
theta1:ev(func,y=lower_y);
theta_1:solve(theta1,theta)[1];
rhs_theta_1:rhs(theta_1);   dd: rhs_theta_1;     
theta2:ev(func,y=upper_y);
theta_2:solve(theta2,theta)[1];
rhs_theta_2:rhs(theta_2);   bb:rhs_theta_2;
inty:int(ds,y);

/*Model answers*/
arc_length:decimalplaces(float(int(ds,y,lower_y,upper_y)),2);
ta:float(int(sec(theta)^3,theta,rhs_theta_1,rhs_theta_2));
ta_1:decimalplaces(aa*ta,2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>The length of  the curve whose equation is given by \( \displaystyle {@expr@}\),  for \({@lower_y@}\leq y\leq{@upper_y@}\) is</p>
<p>\(L={@ta_1@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">[Maria Sayuni]</p>
<p id="reviewer">[Juma Zevick]</p>
<p id="description">The question focuses on finding the arc length of a given function over a defined random range</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta_1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1 checks for correct answer*/
/*Node 2: mistake:Integrating the expression instead of differentiating*/
mistake_2:rhs((int(expr,y)));
/*Node 3: mistake:   Interchanging the limits of integration*/
mistake_1:int(ds,yrhs_y2,rhs_y1);
mistake_2:decimalplaces(float(mistake_1),3);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Integrating the expression instead of differentiating</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have integrated the  the function instead of differentiating</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Interchanging the limits of integration</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your computation looks okay for the definite integral, however it looks like you have interchanged the limits of integration.
Take a look at the worked out solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1183191486</deployedseed>
    <deployedseed>558608838</deployedseed>
    <deployedseed>2040764187</deployedseed>
    <deployedseed>899651164</deployedseed>
    <deployedseed>468916668</deployedseed>
    <deployedseed>1020826846</deployedseed>
    <deployedseed>503474338</deployedseed>
    <deployedseed>493803561</deployedseed>
    <deployedseed>340437088</deployedseed>
    <deployedseed>1275383394</deployedseed>
    <deployedseed>637631599</deployedseed>
    <deployedseed>529799800</deployedseed>
    <deployedseed>852251361</deployedseed>
    <deployedseed>1691773612</deployedseed>
    <deployedseed>1889193743</deployedseed>
    <deployedseed>956780318</deployedseed>
    <deployedseed>1948651691</deployedseed>
    <deployedseed>1887977735</deployedseed>
    <deployedseed>1869034675</deployedseed>
    <deployedseed>1999555361</deployedseed>
    <deployedseed>994006062</deployedseed>
    <deployedseed>1074307805</deployedseed>
    <deployedseed>644085781</deployedseed>
    <deployedseed>616666666</deployedseed>
    <deployedseed>496294173</deployedseed>
    <deployedseed>66815297</deployedseed>
    <deployedseed>753237326</deployedseed>
    <deployedseed>761504351</deployedseed>
    <deployedseed>1378612558</deployedseed>
    <deployedseed>1658508899</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
