<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Determining the Existence of a Limit for a Fractional Function (i)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[</p>Determine whether the limit of the function below exist or not.</p>

<br> \(\displaystyle \lim_{x \longrightarrow {@aa@}} {@fx@}\). If the limit doesn't exist click NO as your answer.


<br>
<button id="btn1yes" type="button" style="font-size:16px" onclick="reveal('1');">Yes</button>
 <button id="btn1no" type="button" style="font-size:16px" onclick="hide('1');">No</button> [[feedback:prtyn1]]
<div id="input1" style="display:none;">[[input:ans1]] [[validation:ans1]]</div>
<br>
<div id="example1" style="display:none;">
<script>
    function reveal(s) {
        document.querySelector("#input" + s).querySelector("input").value = "true";
        document.querySelector("#example" + s).style.display = "block";
        document.querySelector("#btn" + s + "yes").style["font-size"] = "24px";
        document.querySelector("#btn" + s + "no").style["font-size"] = "16px";
    }

    function hide(s) {
        document.querySelector("#input" + s).querySelector("input").value = "false";
        document.querySelector("#example" + s).style.display = "none";
        document.querySelector("#btn" + s + "yes").style["font-size"] = "16px";
        document.querySelector("#btn" + s + "no").style["font-size"] = "24px";
    }
</script>

<script>
    window.onload = function() {
        for (const s of ['1', '2']) {
            if (document.querySelector("#input" + s).querySelector("input").value == "true") {
                reveal(s);
            } else if (document.querySelector("#input" + s).querySelector("input").value == "false") {
                hide(s)
            }
        }
    };
</script>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>

<p>Remember that you need to test values from both sides. You will see that from the left, values tend towards \(-1\) and from the right they tend to \(1\), and therefore there is no limit.<br></p><p>We can also look at simplification of the function. In order to do so we need to consider the following two cases:</p><p>1. Values to the left of {@aa@}. In this case, the value of {@x-aa@} is negative and the function becomes \(\displaystyle&nbsp;f(x) = \frac{{@x-aa@}}{-({@x-aa@})} = -1\) so we can see that from the left the function tends to \(-1\).</p><p>2. Values to the right of {@aa@}. In this case, the value of {@x-aa@} is positive and the function becomes \(\displaystyle&nbsp;f(x) = \frac{{@x-aa@}}{{@x-aa@}} = 1\) so we can see that from the left the function tends to \(1\).</p><p>Therefore there is no limit.<br><br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*the limit will always not exist for my first function*/
pm1:rand_with_prohib(-1,1,[0]);
aa:pm1*(rand(9)+1);
fx:(x-aa)/abs(x-aa);
ta: false;


/*the limit will always  exist for this function
a1:rand(10)+1;
bb:rand(6)+1;
cc:rand(8)+1;
dd:rand(9);
expr_numerator: a1*x^2*y+bb*x*y;
expr_denominator: cc*x^2+dd;

value_of_x: rand_with_prohib(-5,5,[0]);
value_of_y: rand_with_step(-5,5,1);
ta_1: at(expr_numerator, [x=value_of_x, y=value_of_y]);
ta_2: at(expr_denominator, [x=value_of_x, y=value_of_y]);
ta_3: ta_1/ta_2;

question_test: a1;
ta_yn1: if oddp(question_test)=true then true else false ;*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<br> \(\displaystyle \lim_{x \longrightarrow {@aa@}} {@fx@}  \quad {@ta@} \quad \text{the limit does not exist}\).]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">NA</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Santiago Borio</p>
<p id="description">This question asks the student to determine whether the limit of a given function exists. The function is in the form of a fraction, and the student is required to assess the existence of the limit based on the limit of the denominator. The question provides a clear prompt and interactive buttons to guide the student's response. </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i> Your answer is partially correct.</span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prtyn1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct choice</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtyn1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtyn1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>755116261</deployedseed>
    <deployedseed>1927765598</deployedseed>
    <deployedseed>913278478</deployedseed>
    <deployedseed>1210176802</deployedseed>
    <deployedseed>968616666</deployedseed>
    <deployedseed>893148435</deployedseed>
    <deployedseed>6456648</deployedseed>
    <deployedseed>658255146</deployedseed>
    <deployedseed>1947526186</deployedseed>
    <deployedseed>636408627</deployedseed>
    <deployedseed>1350247240</deployedseed>
    <deployedseed>1094482609</deployedseed>
    <deployedseed>1828682662</deployedseed>
    <deployedseed>2138971111</deployedseed>
    <deployedseed>810643518</deployedseed>
    <deployedseed>51215110</deployedseed>
    <deployedseed>1769933586</deployedseed>
    <deployedseed>1545947979</deployedseed>
    <deployedseed>1971398046</deployedseed>
    <deployedseed>1870149712</deployedseed>
    <deployedseed>133997331</deployedseed>
    <deployedseed>855622788</deployedseed>
    <deployedseed>1043236946</deployedseed>
    <deployedseed>179774311</deployedseed>
    <deployedseed>968914077</deployedseed>
    <deployedseed>1330531472</deployedseed>
    <deployedseed>1020920546</deployedseed>
    <deployedseed>92509256</deployedseed>
    <deployedseed>1072534725</deployedseed>
    <deployedseed>540043058</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prtyn1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtyn1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
