<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Differentiation by Power rule</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Given that \( f(x)={@f(x)@}\)</p>
<p>Find:</p></p>
<ol class="HELM_parts">
<li>\(f({@p1@})=\) [[input:ans1]] [[validation:ans1]] </li>
<p>[[feedback:prt1]] </p>
<li>\(f'(x)=\) [[input:ans2]] [[validation:ans2]] </li>
<p>[[feedback:prt2]] </p>
<li>\(f'({@p2@})=\) [[input:ans3]] [[validation:ans3]] </li>
<p>[[feedback:prt3]] </p>
<li>\(f''(x)=\) [[input:ans4]] [[validation:ans4]] </li>
<p>[[feedback:prt4]] </p>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked  Solution</h3>
<p>Given that \( f(x)={@f(x)@}\)</p>
<ol class="HELM_parts">
  <li>
    <p>To find \(f({@p1@})\), we simply need to substitute \(x\) with {@p1@} in the function \(f(x)\):</p>
    <p>\(f({@x@}) = {@f(x)@}\)</p>
    <p>\( \displaystyle \begin{align*} 
              f({@p1@}) &= {@cf1@}({@p1@})^{@pwr1@} + {@cf2@}({@p1@})^2 + {@cf3@}({@p1@}) + {@cf4@} \\
                                &=  {@cf1*(p1^pwr1)@} + {@cf2*(p1^2)@} + {@cf3*p1@} + {@cf4@} \\
                                 &= {@ta1@} \end{align*} \)</p>
    <p>So, \(f({@p1@}) = {@ta1@}\).</p>
  </li>
  <li>
    <p>To find \(f'(x)\), we'll need to take the derivative of the function \(f(x)\) with respect to \(x\), using Power rule</p>
     <p>The power rule is as follows:</p>
    <p>If \( \displaystyle f(x) = ax^n\), where \(a\) and \(n\) are constants, then the derivative \(f'(x)\) is given by:</p>
     <p>\( \displaystyle f'(x) = n \cdot ax^{n-1}\)</p>
  <p>\(f(x) = {@f(x)@}\)</p>
<p> \( \displaystyle \begin{align*}  f'(x) &= \frac{d}{dx}(f(x)) \\
                           &= \frac{d}{dx}({@f(x)@}) \\
                           &= {@ta2@} \end{align*} \)</p>
  </li>
  <li>
    <p>To find \(f'({@p2@})\), substitute \(x\) with {@p2@} in the derivative \(f'(x)\):</p>
    <p>\(f'({@x@}) = {@ta2@}\) </p>
    <p>\(f'({@p2@}) = {@ta3@}\) </p>
<p>So, \(f'({@p2@}) = {@ta3@}\).</p>
  </li>
<li>
<p> The second derivative \(f''(x)\):</p>
<p> \( \displaystyle \begin{align*} f''(x) &= \frac{d}{dx}\left({@ta2@}\right) \\
                    &= {@ta4@} \end{align*} \)</p>
</li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>p1:rand_with_step(-5,5,1);
p2:rand_with_prohib(-5,5,[p1]);
pwr1:rand_with_step(-90,90,3)/(rand(100)+3);
cf1:(p1*2);
cf2:(5*p2);
cf3:(p2*5);
cf4:(rand(10)+10);
expr:cf1*x^pwr1 +cf2*x^2 + cf3*x + cf4;
f(x) :=expr;  
ta1: subst(p1,x,expr);        /* Calculate f(4) */ 
ta2: diff(f(x), x); /* Calculate the derivative f'(x) */ 
ta3:subst(p2,x,ta2);         /* Calculate f'(0) */
ta4: diff(ta2, x); /* Calculate the derivative f''(x) */ </text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \( f(x)={@f(x)@}\)</p>
<p>\(f({@p1@})={@ta1@}\)</p>
<p>\(f'(x)={@ta2@}\)</p>
<p>\(f'({@p2@})={@ta3@}\)</p>
<p>\(f''(x)={@ta4@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="concept">N/A</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer">[Reviewer 2 Name]</p>
<p id="description">Using method of differentiation (Power Rule)</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>Diff</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>573483482</deployedseed>
    <deployedseed>1429028231</deployedseed>
    <deployedseed>127336112</deployedseed>
    <deployedseed>2125480380</deployedseed>
    <deployedseed>695424552</deployedseed>
    <deployedseed>952222253</deployedseed>
    <deployedseed>785591059</deployedseed>
    <deployedseed>1606055314</deployedseed>
    <deployedseed>1658909987</deployedseed>
    <deployedseed>2077891625</deployedseed>
    <deployedseed>101886313</deployedseed>
    <deployedseed>1254526161</deployedseed>
    <deployedseed>1871957205</deployedseed>
    <deployedseed>672738566</deployedseed>
    <deployedseed>1936695861</deployedseed>
    <deployedseed>180474880</deployedseed>
    <deployedseed>298492154</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
