<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Differentiating-exponential-and-logarithmic-functions</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the derivative of \(f(x) = e^{x^2+{@aa@}\cdot x} + {@bb@}^{{@cc@}\cdot x} + \ln(x^2 +{@dd@} )\).</p><p>recall that \(\ln(ax)\) is entered as log(\(ax\)) and \(e^{ax}\) is entered as %e^(\(ax\))</p>
<p>\(f'(x) = \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>\(f(x) = e^{x^2+{@aa@}\cdot x} + {@bb@}^{{@cc@}\cdot x} + \ln(x^2 +{@dd@} )\)<br></p>
<p>\(\begin {align} 
f'(x) &= (2\cdot x + {@aa@}) \cdot e^{x^2+{@aa@}\cdot x} + {@bb@}^{{@cc@}\cdot x}\cdot \ln({@cc@}\cdot x) \cdot {@cc@} + \frac{1}{x^2 +{@dd@}} \cdot (2\cdot x)\\
&= (2\cdot x + {@aa@}) \cdot e^{x^2+{@aa@}\cdot x} + {@cc@}\cdot {@bb@}^{{@cc@}\cdot x}\cdot \ln({@cc@}\cdot x) + \frac{2\cdot x}{x^2 +{@dd@}}\\
\end{align}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(5)+2;
bb:rand(4)+3;
cc:rand(3)+2;
dd:rand(8)+2;
fx:exp(x^2+aa*x)+bb^(cc*x)+log(x^2+dd);
dfx:diff(fx,x);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(f(x) = e^{x^2+{@aa@}\cdot x} + {@bb@}^{{@cc@}\cdot x} + \ln(x^2 +{@dd@} )\).</p>
<p>\(f'(x) ={@dfx@} \)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">NA</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">Differentiating-exponential-and-logarithmic-functions</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>dfx</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description> checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>dfx</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>877182462</deployedseed>
    <deployedseed>18780472</deployedseed>
    <deployedseed>988407642</deployedseed>
    <deployedseed>595249906</deployedseed>
    <deployedseed>2069077894</deployedseed>
    <deployedseed>1050340414</deployedseed>
    <deployedseed>1787334430</deployedseed>
    <deployedseed>1845348075</deployedseed>
    <deployedseed>1007659917</deployedseed>
    <deployedseed>972254749</deployedseed>
    <deployedseed>1537381122</deployedseed>
    <deployedseed>639720389</deployedseed>
    <deployedseed>1808222183</deployedseed>
    <deployedseed>666390898</deployedseed>
    <deployedseed>806305509</deployedseed>
    <deployedseed>404235441</deployedseed>
    <deployedseed>1719756814</deployedseed>
    <deployedseed>1415321338</deployedseed>
    <deployedseed>450649811</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>dfx</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
