<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Directional Derivatives_Number 1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Consider the function \(f(x,y)\) given below,</p>
<p>\(\displaystyle f(x,y)={@func@}\)</p>  
<p>Calculate the gradient function \(\nabla f (x,y)\) </p>
<p>\(\nabla f (x,y)=\) \(\displaystyle ( \) [[input:ans1]]\(,\)[[input:ans2]]\(\displaystyle ) \)</p>
<p>[[validation:ans1]][[validation:ans2]][[feedback:prt1]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>What we need to do here is to find \(\displaystyle \frac{d}{dy}f(x,y)\) and \(\displaystyle \frac{d}{dx}f(x,y)\)</p>
<p>\(\displaystyle \frac{d}{dx}f(x,y)={@gradient_func1@}\)</p>
<p>\(\displaystyle\frac{d}{dy}f(x,y)={@gradient_func2@}\)</p>
<p>\(\displaystyle \nabla f (x,y)={@gradient_func1@},{@gradient_func2@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>power_x:rand(5)+1;
coef_x:rand(4)+1;
coef_y:rand(4)+1;
func:(power_x*x^2)*sec(coef_x*x)-(power_x*x^2)/(coef_y*y^3);
gradient_func1:diff((func),x);
gradient_func2:diff((func),y);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\displaystyle f(x,y)={@func@}\)</p>  
<p>\(\displaystyle \nabla (f x,y)={@gradient_func1@},{@gradient_func2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer">Salvatore Mercuri</p>
<p id="description">This question involves Calculating the gradient function \(\triangledown f (x,y)\) given a function \(f(x,y)\) </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>gradient_func1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>gradient_func2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1 checks for correct answer*/
/*Node 2 mistake: Differentiating with respect to y*/
mistake_2:gradient_func2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for gradient_func1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>gradient_func1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for gradient_func2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>gradient_func2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It seems you have interchanged the inputs. Remember you are supposed to compute \( \frac {df(x,y)}{dx} and insert the result here for the second input.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1 checks for correct answer*/
/*Node 2 mistake:Differentiating with respect to x*/
mistake_2:gradient_func1;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for gradient_func2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>gradient_func2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for gradient_func1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>gradient_func1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It seems you have interchanged the inputs. Remeber you are supposed to compute \( \frac {df(x,y)}{dy} and insert the result here for the second input.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1294885624</deployedseed>
    <deployedseed>1747900244</deployedseed>
    <deployedseed>1014230337</deployedseed>
    <deployedseed>1387136994</deployedseed>
    <deployedseed>34187588</deployedseed>
    <deployedseed>1407492655</deployedseed>
    <deployedseed>1941445354</deployedseed>
    <deployedseed>453174704</deployedseed>
    <deployedseed>1385469310</deployedseed>
    <deployedseed>156980517</deployedseed>
    <deployedseed>1757775645</deployedseed>
    <deployedseed>1020057829</deployedseed>
    <deployedseed>731161588</deployedseed>
    <deployedseed>1095569642</deployedseed>
    <deployedseed>1539477249</deployedseed>
    <deployedseed>1491790012</deployedseed>
    <deployedseed>1313753252</deployedseed>
    <deployedseed>304430430</deployedseed>
    <deployedseed>690729722</deployedseed>
    <deployedseed>9121937</deployedseed>
    <deployedseed>828849237</deployedseed>
    <deployedseed>1038218958</deployedseed>
    <deployedseed>807074292</deployedseed>
    <deployedseed>1350502657</deployedseed>
    <deployedseed>1081825170</deployedseed>
    <deployedseed>1169585414</deployedseed>
    <deployedseed>995558548</deployedseed>
    <deployedseed>2137042866</deployedseed>
    <deployedseed>1364522522</deployedseed>
    <deployedseed>736562056</deployedseed>
    <deployedseed>2099030902</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>gradient_func1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>gradient_func2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
