<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Domain and Range of a Function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Find the domain and range of \(f(x)={@fx@}\).</p><p></p><p>Syntax Note:</p><p></p><ul><li>Type R for all real numbers<br></li><li>The interval \((2,5)\) would be written as \(2&lt;x\) and \(x&lt;5\)&nbsp;<br></li><li>The interval \([-1,2)\) would be written as \(-1&lt;=x\) and \(x&lt;2\)&nbsp;</li><li>The interval \((-\infty ,3)\) would be written as \(x&lt;3\)</li><li>The interval \([7,\infty)\) would be written as \(x&gt;=7\)</li></ul><p></p><p><br></p><p>Domain = [[input:ans1]] [[validation:ans1]]&nbsp;[[feedback:prt1]]</p><p>Range = [[input:ans2]] [[validation:ans2]]&nbsp;[[feedback:prt2]]<br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[bb:rand(6)+1;
fx:sqrt(bb^2-x^2);
ta1:x>=-bb and x<=bb;
ymin:0;
ymax:bb;
ta2:y>=ymin and y<=ymax;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>aa,bb,cc,pm1,pm2</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>To calculate the domain of \(f(x)\) we need to find the values of \(x\) for which the quadratic inside the square root is non-negative as we cannot calculate the square root of a negative number. To do this, we need to find the roots of the quadratic, which can be done with the quadratic formula or by&nbsp;factorising. In this case, \({@bb^2-x^2@} = {@(bb-x)*(bb+x)@}\) and the roots are {@-bb@} and {@bb@}. We now note that as it is a negative quadratic, the values between the roots will be positive. Hence, the domain is the set \([{@-bb@},{@bb@}]\), which can be expressed as \(x\geq {@-bb@}\) and \(x\leq {@bb@}\).<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>To calculate the range of this function we need to consider the minimal and maximal values that can be obtained by analysing the quadratic inside the square root. As it is a negative quadratic, the maximum&nbsp;needs to be calculated, which in this case is when \(x=0\) and the maximum value of \(f(x)\) is \(f(0) = {@bb@}\). The minimal value will happen when the quadratic equals 0, and the minimum&nbsp;value of \(f(x)\) is \(0\). Therefore the range of \(f\) is the interval \([0,{@bb@}]\), which can be written as \(y\geq 0\) and \(y\leq {@bb@}\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
