<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Domain-of-sum-of-roots</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Find the domain&nbsp; of \(f(t)={@fx@}\).</p><p></p><p>Syntax Note:</p><p></p><ul><li>Type R for all real numbers<br></li><li>The interval \((2,5)\) would be written as \(2&lt;x\) and \(x&lt;5\)&nbsp;<br></li><li>The interval \([-1,2)\) would be written as \(-1&lt;=x\) and \(x&lt;2\)&nbsp;</li><li>The interval \((-\infty ,3)\) would be written as \(x&lt;3\)</li><li>The interval \([7,\infty)\) would be written as \(x&gt;=7\)</li></ul><p></p><p><br></p><p>Domain = [[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[simp:false;
aa:ev(rand(10)+7,simp);
bb:ev(rand(5)+2,simp);
fx:sqrt(aa-t)+sqrt(bb+t);
xmin:-bb;
xmax:aa;
ta1:t>=xmin and t<=xmax;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]<br></p>]]></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>aa,bb</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>x</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>t</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>When looking at the domain of a square root function we&nbsp;analyse&nbsp;when the expression inside it would be non-negative. Here we need to look at \(\sqrt{{@aa-t@}}\) first, and we see that \({@aa-t@}\) is non-negative when \(t\leq {@aa@}\). We then look at&nbsp;\(\sqrt{{@bb+t@}}\), and we see that \({@bb+t@}\) is non-negative when \(t\geq {@-bb@}\). Since both conditions need to be true simultaneously the domain is the interval \([{@-bb@},{@aa@}]\), which is entered as \(t \geq {@-bb@}\) and \(t \leq {@aa@}\).<br></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
