<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Equation Solving and Root Existence in an Interval</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the equation \({@fx@} = {@aa@}\)</p>

<br>

<p>(a) Find the value of \({@fx@}\) when \(x = {@cc@}\)</p>
<p style="margin-left: 30px">[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>

<br>

<p>(b) Find the value of \({@fx@}\) when \(x = {@dd@}\)</p>
<p style="margin-left: 30px">[[input:ans2]] [[validation:ans2]]</p>
<p>[[feedback:prt2]]</p>

<br>

<p>(c) Determine whether there exists at least one root for the equation \({@fx2@} = 0\) in the interval \(({@cc@},{@dd@})\).</p>
<p style="margin-left: 30px">[[input:ans3]] [[validation:ans3]]</p>
<p>[[feedback:prt3]]</p>

<br>

<p>(d) State the name of the theorem that justifies your answer to part (c).<p>
<p style="margin-left: 30px">[[input:ans4]] [[validation:ans4]]</p>
<p>[[feedback:prt4]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>

<p>For the first two parts we simply need to evaluate \({@fx@}\) at \(x = {@cc@}\) and \(x = {@dd@}\).</p>
<p>(a) &emsp; \({@bb@} \cdot {@if cc=0 then "" else "("@}{@cc@}{@if cc=0 then "" else ")"@}^5 - 
{@if cc=0 then "" else "("@}{@cc@}{@if cc=0 then "" else ")"@}^3 = {@ta1@}\)</p>
<p>(b) &emsp; \({@bb@} \cdot {@dd@}^5 - {@dd@}^3 = {@ta2@}\)</p>

<br>

<p>(c) We can use the information from parts (a) and (b) to evaluate the value of the expression given at the end points of the interval,</p>
<p style="margin-left: 30px">At \(x = {@cc@}\), we have \({@fx2@} = {@ta1 - aa@}\)</p>
<p style="margin-left: 30px">At \(x = {@dd@}\), we have \({@fx2@} = {@ta2 - aa@}\)</p>

<br>

<p>Since the function {@fx2@} is <b>continuous</b>, and since there is a <b>change of sign</b> over the interval \({@oo(cc,dd)@}\), we can conclude that there exists a root of the equation in the given interval.</p>

<br>

<p>(d) The name of the theorem which justifies our conclusion in part (c) is <b>The Intermediate Value Theorem</b>.]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[bb:rand(5)+2;
fx:bb*x^5-x^3;
cc:rand([-2,-1,0]);
dd:rand([1,2,3]);

ta1:ev(fx,x=cc);
ta2:ev(fx,x=dd);
aa:rand_with_prohib(ta1+1, ta2-1, []);
fx2:bb*x^5-x^3-aa;
ma3: [
  ["Yes, there exists a root", true],
  ["No, there does not exist a root", false],
  ["We cannot determine if a root exists from the information given", false]
];

incorr_thms: rand_selection([
  ["The Substitution Rule", false],
  ["The Maximum-Minimum Theorem", false],
  ["The First Derivative Test", false],
  ["The Second Derivative Test", false],
  ["The Power Rule", false],
  ["The Chain Rule", false],
  ["The Product Rule", false],
  ["The Quotient Rule", false]
], 3);
ma4: append(
  [["The Intermediate Value Theorem", true]],
  incorr_thms
);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>The equation \({@fx@} = {@aa@}\).</p>
<p>Value of \({@fx@}\) when \(x = {@cc@}: \, {@ta1@} \).</p>
<p>Value of \({@fx@}\) when \(x = {@dd@}: \, {@ta2@} \).</p>]]></text>
    </questionnote>
    <questiondescription format="markdown">
      <text><![CDATA[<p id="author">Godfrey Wabwirep>
<p id="concept">N/A</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">This question involves solving an equation and analyzing the existence of a root within a given interval. Students are asked to find specific values of the equation for different input values, determine if a root exists within the interval, and identify the theorem that justifies their conclusion.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ma3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>ma4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct string answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct string answer</description>
        <answertest>String</answertest>
        <sans>ans3</sans>
        <tans>mcq_correct(ma3)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct string answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct string answer</description>
        <answertest>String</answertest>
        <sans>ans4</sans>
        <tans><![CDATA["The Intermediate Value Theorem"]]></tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>108653658</deployedseed>
    <deployedseed>707309059</deployedseed>
    <deployedseed>517242333</deployedseed>
    <deployedseed>1927783709</deployedseed>
    <deployedseed>103429016</deployedseed>
    <deployedseed>1563781686</deployedseed>
    <deployedseed>238562766</deployedseed>
    <deployedseed>646603489</deployedseed>
    <deployedseed>69422948</deployedseed>
    <deployedseed>652907568</deployedseed>
    <deployedseed>914498222</deployedseed>
    <deployedseed>1100403838</deployedseed>
    <deployedseed>617333156</deployedseed>
    <deployedseed>374850080</deployedseed>
    <deployedseed>2041672003</deployedseed>
    <deployedseed>1970119785</deployedseed>
    <deployedseed>86884142</deployedseed>
    <deployedseed>447348960</deployedseed>
    <deployedseed>619233729</deployedseed>
    <deployedseed>1864001773</deployedseed>
    <deployedseed>2013139010</deployedseed>
    <deployedseed>663865134</deployedseed>
    <deployedseed>1410226758</deployedseed>
    <deployedseed>1885447977</deployedseed>
    <deployedseed>765255740</deployedseed>
    <deployedseed>526065018</deployedseed>
    <deployedseed>921236558</deployedseed>
    <deployedseed>1232222177</deployedseed>
    <deployedseed>1683878653</deployedseed>
    <deployedseed>254231256</deployedseed>
    <deployedseed>117624659</deployedseed>
    <deployedseed>1619066049</deployedseed>
    <deployedseed>2142366076</deployedseed>
    <deployedseed>445827105</deployedseed>
    <deployedseed>746423550</deployedseed>
    <deployedseed>56847250</deployedseed>
    <deployedseed>516152084</deployedseed>
    <deployedseed>809966902</deployedseed>
    <deployedseed>502226059</deployedseed>
    <deployedseed>569041889</deployedseed>
    <deployedseed>1104783030</deployedseed>
    <deployedseed>962969243</deployedseed>
    <deployedseed>1850556363</deployedseed>
    <deployedseed>483839846</deployedseed>
    <deployedseed>137869577</deployedseed>
    <deployedseed>562833059</deployedseed>
    <deployedseed>402844513</deployedseed>
    <deployedseed>742850979</deployedseed>
    <deployedseed>19871824</deployedseed>
    <deployedseed>917844837</deployedseed>
    <deployedseed>1812505553</deployedseed>
    <deployedseed>1262485530</deployedseed>
    <deployedseed>1746705082</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(ma3))</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>first(mcq_correct(ma4))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
