<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Evaluate the definite Integral of a trigonometric Functions</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Evaluate the following definite integral:</p>
<p>\[ \int^{@rdm3@}_{@rdm4@} {@p@} d{@v@}. \]</p>

<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To solve the integral \(\int {@p@} d{@v@}\) we can use integration by parts. </p>
<p>The formula for integration by parts is:</p>
<p>\[\int u dv = uv - \int v du\]</p> 

<p>First, lets find the indefinite integral </p>

<p>\[\begin{array}{rcl}\displaystyle 
 \int {@p@} d{@v@} &\displaystyle =& \displaystyle {@ppp1@}\left[{@int(ppp2,v)@}\right] - \int \left[{@int(ppp2,v)@}\right]{@diff(ppp1,v)@} d{@v@} \\
   &\displaystyle =& \displaystyle {@ppp1*int(ppp2,v)@} - \int {@int(ppp2,v)*diff(ppp1,v)@} d{@v@} \\
   &\displaystyle =& \displaystyle {@ppp1*int(ppp2,v)-int((int(ppp2,v)*diff(ppp1,v)),v)@}
 \end{array}\] </p>

<p>Now let us substitute back the limits of integration, </p>

<p>\[\begin{array}{rcl}\displaystyle 
 \int^{@rdm3@}_{@rdm4@} {@p@} d{@v@} &\displaystyle =&\displaystyle  F({@rdm3@}) - F({@rdm4@}) \\ 
   &\displaystyle =&\displaystyle  {@subst(rdm3,v,q)@} - {@subst(rdm4,v,q)@} \\
   &\displaystyle =&\displaystyle  {@int(p,v,rdm4,rdm3)@}
 \end{array}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>v:x;
rdm2:-2-rand(2);
rdm3:rand([pi/(4*rdm2),pi/(2*rdm2),pi/(2*rdm2),pi/(rdm2),2*pi/(4*rdm2)]);
rdm4:0;
ppp1:v;
ppp2:cos(rdm2*v);
p:ppp1*ppp2;
q:int(p,v);
ta:int(p,v,rdm4,rdm3);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>The definite integral:</p>
<p>\[ \int^{@rdm3@}_{@rdm4@} {@p@} d{@v@}={@int(p,v,rdm4,rdm3)@} \]</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">[Juma Zevick]</p>
<p id="description">The question focuses on Evaluating the definite Integral of a trigonometric function</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>int</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */
/*Node 2: Computing indefinite integral*/
mistake1: int(p,v)+c;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Giving the answer as a variable </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Remember, this is a definite integration problem so you need to substitute back the limits of integration.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2070934317</deployedseed>
    <deployedseed>65991505</deployedseed>
    <deployedseed>1298288302</deployedseed>
    <deployedseed>1853963510</deployedseed>
    <deployedseed>1266628242</deployedseed>
    <deployedseed>1127084853</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v,rdm4,rdm3),simp)</value>
      </testinput>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ev(int(p,v),simp)</value>
      </testinput>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>p</value>
      </testinput>
    </qtest>
  </question>
</quiz>
