<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Evaluating the Convergence and Value of an exponetial Improper Integral</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Evaluate the integral below and determine if it converges or diverges. 
</p>
<p>Note: \( \infty\) is entered as <b>inf</b>  \( -\infty\) is entered as <b>minf</b> and  \( e\) is entered as <b>%e</b>.</p>
<p>\(\displaystyle \int_{{@lower_lm@}}^{{@upper_lm@}} {@num*%e@}^{ \displaystyle {@(-x^power/den)@}} \,dx=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>Therefore, the integral is [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3> 
<p>Recall that we can’t evaluate the integral as long as there is an infinity in one of the limits. Therefore, we’ll need to replace infinity with a parameter \( t \) as follows:</p>
 <p>The given improper integral is</p> 
<p> \(\displaystyle I = \int_{@lower_lm@}^{@upper_lm@} {@fx@} \, dx. \) </p>
 <p>We will first evaluate whether the integral converges or diverges by finding the limit of the integral as \( t \to \infty \):</p> <p> \(\displaystyle I = \int_{@lower_lm@}^{t} {@fx@} \, dx. \) </p>

 <p>The antiderivative of \(\displaystyle {@fx@} \) is:</p>
 <p> \(\displaystyle \int {@fx@} \, dx = {@indef_int@} \) </p>
 <p>Now, evaluate the integral from \( {@lower_lm@} \) to \( t \):</p>
 <p> \(\displaystyle I(t) = \left[ {@indef_int@} \right]_{@lower_lm@}^t = {@indef_val@} \) </p> 
<p>Take the limit as \( t \to \infty \):</p> <p> \(\displaystyle I = \lim_{t \to \infty} \left( {@indef_val@} \right) = {@def_integral@} \) </p>
 <p>The improper integral is {@mcq_correct(ta_4)[1]@}, and its exact value is \( \boxed{{@def_integral@}} \).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Define the constants and variables */
num: rand(5)+2; /* Coefficient of e^-x/4 */
power: 1; /* The function e^-x/4 does not depend on x^power */
upper_lm: inf; /* Upper limit of integration */
lower_lm: rand(8)+1; /* Lower limit of integration */
limit_pt: t; /* Variable for upper limit of integration */
den: rand(10)+4; /* Denominator from the exponent in e^-x/4 */
fx : num * exp(-x^power/den); /* Define the function */

/* Compute the indefinite integral */
indef_int: int(fx, x);

/* Compute the definite integral from 0 to b */
upper: at(indef_int, x = limit_pt);
lower: at(indef_int, x = lower_lm);
indef_val: upper - lower;

/* Compute the definite integral from 0 to infinity */
def_integral: limit(indef_val, limit_pt, inf);

/* Model answers */
ta_1: def_integral;
ta_2: if ta_1 = def_integral then "Convergent" elseif ta_1 = inf then "Divergent" else "Divergent";
ta_3: if ta_2 = "Convergent" then "Divergent" else "Convergent"; /* Inverse logic for clarity */
ta_4: random_permutation([[ta_2, true], [ta_3, false], ["Divergent and convergent", false]]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Evaluate the integral below and determine if it converges or diverges.</p>
<p>\(\displaystyle\int_{{@lower_lm@}}^{{@upper_lm@}}{@fx@}\,dx={@ta_1@}\)</p>
<p>Therefore, the integral is {@mcq_correct(ta_4)[1]@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="concept">N/A</p>
<p id="reviewer">1</p>
<p id="reviewer">2</p>
<p id="description">The question focuses on evaluating an Improper Integral</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta_4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1:checks for correct answers*/

/*node 2:Calculating the integral as F(a)-F(b) instead of F(b)-F(a)*/
mistake_1:int(fx,upper_lm, lower_lm) ;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for the first mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You calculated the value of the integral in a wrong format \(F(a)-F(b)\),instead of \(F(b)-F(a)\)</p>
<p>Hint: Check the worked out solution</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_4)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1485135548</deployedseed>
    <deployedseed>469900350</deployedseed>
    <deployedseed>403840166</deployedseed>
    <deployedseed>671643018</deployedseed>
    <deployedseed>1923060838</deployedseed>
    <deployedseed>1433365102</deployedseed>
    <deployedseed>2069750315</deployedseed>
    <deployedseed>725301087</deployedseed>
    <deployedseed>1888079340</deployedseed>
    <deployedseed>2002634990</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta_4))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
