<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Evaluating an Integral by the Fundamental Theorem of Calculus</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Use Fundamental Theorem of Calculus to evaluate</p>
<p>\(\displaystyle \int _{@lower_limit@}^{@upper_limit@}{@quest_expr@}\,dt\)</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Recall: \( \displaystyle \int x^n dx=\frac{x^{n+1}}{n+1}+c\)</p>
<p></p>
<p>Using this rule we'll have:</p>
<p>\(\begin{align}
\displaystyle 
\int _{@lower_limit@}^{@upper_limit@}{@quest_expr@}&=\left [ {@soln1@} \right ]_{@lower_limit@}^{@upper_limit@}\\
                                                                                &= {@ta1@}\\
\end{align}
\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Fundamental Theorem*/
lower_limit:rand(2)-4;
upper_limit:rand(2)+4;
power:rand(4)+2;
constant:rand(8)+4;
coef_t:rand(3)+6;
quest_expr:coef_t*t^power-constant;
/*worked solution*/
soln1:int(quest_expr,t);
question_solution:int(quest_expr,t,lower_limit,upper_limit);
soln3:subst(upper_limit,t,soln1);
soln4:subst(lower_limit,t,soln1);
s5:soln4-soln3;
ta1:float(question_solution);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Use Fundamental Theorem of Calculus to evaluate \(\displaystyle \int _{@lower_limit@}^{@upper_limit@}{@quest_expr@}\,dt={@ta1@}\)</p>
]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept"></p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer">Juma Zevick</p>
<p id="description">Use the second part of the Fundamental Theorem of Calculus to evaluate \(\displaystyle \int_{@lower_limit@}^{@upper_limit@}\)</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1063302115</deployedseed>
    <deployedseed>956180064</deployedseed>
    <deployedseed>1776855056</deployedseed>
    <deployedseed>1022475186</deployedseed>
    <deployedseed>688605592</deployedseed>
    <deployedseed>1152786535</deployedseed>
    <deployedseed>785897936</deployedseed>
    <deployedseed>1938480792</deployedseed>
    <deployedseed>1528068901</deployedseed>
    <deployedseed>803671090</deployedseed>
    <deployedseed>1314649398</deployedseed>
    <deployedseed>62687662</deployedseed>
    <deployedseed>415029901</deployedseed>
    <deployedseed>596066509</deployedseed>
    <deployedseed>851494990</deployedseed>
    <deployedseed>518031879</deployedseed>
    <deployedseed>16328286</deployedseed>
    <deployedseed>2101636167</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
