<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Evaluating limits of functions (a and b)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Evaluate the value of the following limit:</p>
<p><strong>Note:</strong> Give your answer in fraction form or correct to \(2\) decimal places</p>
<p>\(\displaystyle \lim_{{@var_letter@}\rightarrow {@val_limit@}}{@eqn@} =\)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
[[ if test='eqn = eqn_1' ]]
    \(\begin{align*}
 \lim_{{@var_letter@}\rightarrow {@val_limit@}}{@eqn@} &= \lim_{{@var_letter@}\rightarrow {@val_limit@}}{@eqn@} \cdot \frac{{@sqrt(var_letter^2+var_2)+var_1@}}{{@sqrt(var_letter^2+var_2)+var_1@}} \\
& = \lim_{{@var_letter@}\rightarrow {@val_limit@}} \frac{\left ({@sqrt(var_letter^2+var_2)-var_1@}\right )  \left ( {@sqrt(var_letter^2+var_2)+var_1@}\right )}{{@den_eqn@} \left ({@sqrt(var_letter^2+var_2)+var_1@}\right )}\\
&  = \lim_{{@var_letter@}\rightarrow {@val_limit@}} \frac{\left ( {@sqrt(var_letter^2+var_2)@}\right )^{2}-{@var_1@}^2}{{@den_eqn@} \left ({@sqrt(var_letter^2+var_2)+var_1@}\right )}\\
& = \lim_{{@var_letter@}\rightarrow {@val_limit@}} \frac{{@var_letter@}^{2}}{{@den_eqn@} \left ({@sqrt(var_letter^2+var_2)+var_1@}\right )}\\
& = \frac{1}{ \left ({@sqrt(var_letter^2+var_2)+var_1@}\right )}\\
& = \frac{1}{ {@sqrt(val_limit^2+var_2)+var_1@}}\\
&= {@ta_1@}\\
\end{align*}
\)
[[ else ]]
   \(\begin{align*}
\lim_{{@var_letter@} \to {@val_limit@}} {@eqn@} &= \frac{{@num_eqn@}}{{@den_eqn_1@}} \\
&= {@sol_1@} \\
&= {@sol_2@} \\
&= {@ta_2@}\\
\end{align*}
\)
[[/ if]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>var_1:rand(10)+1;
var_2:var_1^2;
val_limit:rand(3);
var_letter:rand([x,y,t,z,r]);
eqn_1:(sqrt(var_letter^2+var_2)-var_1)/(var_letter^2);
eqn_2:(var_letter-var_1)/(var_letter^2-var_2);
eqn:rand([eqn_1,eqn_2]);

num_eqn:num(eqn);
den_eqn:denom(eqn); 
den_eqn_limit: at(den_eqn, [var_letter=val_limit]);
den_eqn_1:factor(den_eqn);
sol_1:num_eqn/den_eqn_1;/*the simplified form of the equation*/
sol_a:subst(val_limit,var_letter,eqn_1);
sol_b:1/(sqrt(val_limit^2+var_2)+var_1)
ta_1:decimalplaces(sol_b,2)
sol_2:subst(val_limit,var_letter,sol_1);
ta_2:decimalplaces(sol_2,2);
if eqn = eqn_1 then ta: sol_b else ta:sol_2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\displaystyle \lim_{{@var_letter@}\rightarrow {@val_limit@}}{@eqn@} ={@ta@}\)]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">https://docs.google.com/document/d/1iWHYSZIwx9SW_6nJsyzWwdMU1g0uqRrcF44y5MjZYBY/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">[Reviewer 2 Name]</p>
<p id="description">This question requires students to explores the behavior of the functions as they approach specific points and determine the values of the limits or identify if they exist</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer in fraction form</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer correct to 2 dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1434819900</deployedseed>
    <deployedseed>288900256</deployedseed>
    <deployedseed>840334860</deployedseed>
    <deployedseed>1115504099</deployedseed>
    <deployedseed>1213968782</deployedseed>
    <deployedseed>2082374001</deployedseed>
    <deployedseed>572142121</deployedseed>
    <deployedseed>1630881190</deployedseed>
    <deployedseed>1935686480</deployedseed>
    <deployedseed>177737704</deployedseed>
    <deployedseed>397918832</deployedseed>
    <deployedseed>211416568</deployedseed>
    <deployedseed>1132438657</deployedseed>
    <deployedseed>382271984</deployedseed>
    <deployedseed>1847231148</deployedseed>
    <deployedseed>2040898589</deployedseed>
    <deployedseed>1951872172</deployedseed>
    <deployedseed>1933961311</deployedseed>
    <deployedseed>1734199789</deployedseed>
    <deployedseed>1406638382</deployedseed>
    <deployedseed>261690622</deployedseed>
    <deployedseed>347662827</deployedseed>
    <deployedseed>1144074703</deployedseed>
    <deployedseed>869694230</deployedseed>
    <deployedseed>733483270</deployedseed>
    <deployedseed>333789131</deployedseed>
    <deployedseed>472031181</deployedseed>
    <deployedseed>820992491</deployedseed>
    <deployedseed>628261833</deployedseed>
    <deployedseed>1872672486</deployedseed>
    <deployedseed>1843792891</deployedseed>
    <deployedseed>1441120271</deployedseed>
    <deployedseed>2076177559</deployedseed>
    <deployedseed>2020925039</deployedseed>
    <deployedseed>1918764022</deployedseed>
    <deployedseed>42235934</deployedseed>
    <deployedseed>1002547848</deployedseed>
    <deployedseed>1260083406</deployedseed>
    <deployedseed>324737634</deployedseed>
    <deployedseed>312784570</deployedseed>
    <deployedseed>1883316163</deployedseed>
    <deployedseed>1740230167</deployedseed>
    <deployedseed>949067327</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
