<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Evaluate the Definite integral by partial fraction</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>1. Write the expression below as a sum of its partial fractions.</p>
<p>\( \displaystyle {@question_expression@}=\) [[input:ans1]]+[[input:ans2]]+[[input:ans3]] <br></p>
<p>[[validation:ans1]][[validation:ans2]][[validation:ans3]] [[feedback:prt1]]</p>
<p>2. Using the partial fractions above to compute the integral&nbsp;\( \displaystyle \int_{{@lower_limit@}}^{@upper_limit@} {@question_expression@}dx\).&nbsp;<br><br><br></p>
<p>\( \displaystyle \int_{{@lower_limit@}}^{@upper_limit@} {@question_expression@}dx =\) [[input:ans4]][[validation:ans4]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><span><strong>Worked solution</strong></span></p>
<p>In this case the denominator is already factored and so we go straight to the form of the partial fraction for the integrand.</p>

<p style=""> To write \( \displaystyle {@question_expression@}\) as a sum of its partial fractions we need to follow the steps below:</p>
<p style="">\[{@question_expression@} ={@part_frac_A@}+{@part_frac_B@}+{@part_frac_C@}\]</p>
<p style="">Setting the numerators equal gives:</p>
<p style="">\[{@expr@}\]</p>
<p style="">Getting the coefficients A, B and C we have</p>
<p style="">Let \(x={@-bb@}:\quad {@expr1@} \implies {@expB@}\). </p>
<p style="">Let \(x={@-aa@}:\quad {@expr2@} \implies {@expA@}\).</p>
<p style="">Let \(x={@-cc@}:\quad {@expr3@} \implies {@expC@}\). </p>
<p style="">The partial fraction of the integrand is then,</p>
<p style="">\[{@expr4@}\]</p>
<p style="">Evaluating the integral,</p>


<p><span style="font-size: 0.9375rem;">\[\displaystyle \int_{{@lower_limit@}}^{@upper_limit@} {@question_expression@}dx= \int_{{@lower_limit@}}^{@upper_limit@} {@part_frac_1@}\,dx+&nbsp;\int_{{@lower_limit@}}^{@upper_limit@}{@part_frac_2@}\,dx+&nbsp;\int_{{@lower_limit@}}^{@upper_limit@}{@part_frac_3@}dx\]</span><br></p>

<p></p>\[\iff \displaystyle \int_{{@lower_limit@}}^{@upper_limit@} {@question_expression@}dx={@int_ta1@}+{@int_ta2@}+{@int_ta3@}\]<br>Therefore,<br>\[\displaystyle \int_{{@lower_limit@}}^{@upper_limit@} {@question_expression@}dx={@ta@}\]<br><p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Random variables for the question*/
aa:rand(5); /*constant in the first part of the partial fraction*/
bb:rand_with_prohib(1,5,[aa]);/*constant in the second part of the partial fraction*/
cc:rand_with_prohib(1,5,[aa,bb]);/*constant in the second part of the partial fraction*/
zz_1:rand_with_prohib(1,6,[0]);/* numerator_1*/
zz_2:rand_with_prohib(1,6,[0,zz_1]);/* numerator_2*/
zz_3:rand_with_prohib(1,6,[0,zz_1,zz_2]);/* numerator_3*/
part_frac_1:zz_1/(x+aa);/*part one of the partial fraction*/
part_frac_2:zz_2/(x+bb);/*part two of the partial fraction*/
part_frac_3:zz_3/(x+cc);/*part three of the partial fraction*/
question_expression: factor(part_frac_1+part_frac_2+part_frac_3);
lower_limit:rand_with_step(1,4,1);
upper_limit:rand_with_step(5,9,1);

/*model answer*/
ta1:part_frac_1;
ta2:part_frac_2;
ta3:part_frac_3;
ta4:ta1+ta2+ta3;


/*Worked out solution codes */
part_frac_A:A/(x+aa);
part_frac_B:B/(x+bb);
part_frac_C:C/(x+cc);
numerator_question_expression:num(question_expression);
expr:A*(x+bb)*(x+cc)+B*(x+aa)*(x+cc)+C*(x+aa)*(x+bb)=numerator_question_expression;
expr1:ev(expr,x=-bb);
expB:expr1/ev((x+aa)*(x+cc),x=-bb);
expr2:ev(expr,x=-aa);
expA:expr2/ev((x+bb)*(x+cc),x=-aa);
expr3:ev(expr,x=-cc);
expC:expr3/ev((x+aa)*(x+bb),x=-cc);

expr4: question_expression=ta1+ta2+ta3;
int_ta1:int(ta1,x,lower_limit,upper_limit);
int_ta2:int(ta2,x,lower_limit,upper_limit);
int_ta3:int(ta3,x,lower_limit,upper_limit);
ta:int(question_expression,x,lower_limit,upper_limit);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\( \displaystyle {@question_expression@}={@ta1@}+{@ta2@}+{@ta3@}\)</p>
<p>\( \displaystyle \int_{{@lower_limit@}}^{@upper_limit@} {@question_expression@}dx ={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Emmaculate</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">The question focuses on Evaluating a definite Integral by partial fraction and using the partial fraction to find value of the Integral</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/*Node 1: Check answer is correct and in correct form*/
sans_full: ans1+ans2+ans3; /*Using EqualComAss test*/

/*Node 2: check if the answer is equivalent to the expression given*/
sans_full: ans1+ans2+ans3; /*Using Algebraic test*/
/*Node 3: Identify one incorrect answer */
sans_set:{ans1,ans2,ans3};
ta_set:{ta1,ta2,ta3};
correct:intersect(sans_set,ta_set);
incorrect:setdifference(sans_set,ta_set);
any_correct: if 1<=cardinality(correct) and cardinality(correct)<=2 then 1 else 0;]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>EqualComAss</answertest>
        <sans>sans_full</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks if the answer is equivalent to the expression given</description>
        <answertest>AlgEquiv</answertest>
        <sans>sans_full</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">Your answer is equivalent to {@question_expression@}. However, it is not written in full partial fractions.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks to Identify one incorrect answer </description>
        <answertest>AlgEquiv</answertest>
        <sans>any_correct</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;"></p><p></p><p></p><p>The following answer(s) is (are) correct: {@correct@}. However, the following answer(s) is (are) not correct: {@incorrect@}.</p><p>You can verify this by adding your answers together:</p><p>\(\displaystyle {@ans1@}+{@ans2@}+{@ans3@} = {@factor(sans_full)@} \neq {@question_expression@} \).</p><br><br><p></p><p></p><br><p></p>]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;"></p><p></p><p></p><p>None of your answers are correct.</p><p>You can verify this by adding your answers together:</p><p>\(\displaystyle {@ans1@}+{@ans2@}+{@ans3@} = {@factor(sans_full)@} \neq {@question_expression@} \).</p><br><br><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1 checks for correct answer*/

/*node 2 mistake:Interchanging the limits of integration, however, with the correct definite integral*/
mistake_2:int(question_expression,x,upper_limit,lower_limit);

/*node 3 mistake:Computing indefinite integral with constant*/
mistake_3:int(question_expression,x)+c;

/*node 4 mistake:Computing indefinite integral without the constant*/
mistake_4:int(question_expression,x);

/*node 5 mistake:Computing definite integral then add constant*/
mistake_5:int(question_expression,x,lower_limit,upper_limit)+c;

/*node 6:  using the wrong part_frac expressions written in (a) and computing its integral correctly*/
question_expression_6:ans1+ans2+ans3;
mistake_6:int(question_expression_6,x,lower_limit,upper_limit);
/*node 7:  using the wrong part_frac expressions written in (a), computing its integral correctly and add a constant*/
question_expression_7:ans1+ans2+ans3;
mistake_7:int(question_expression_7,x,lower_limit,upper_limit)+c;
/*node8: Giving the answer in sum of partial fractions form without completing the integral*/
mistake_8:ans1+ans2+ans3;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for first mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your computation looks okay for the definite integral, however it looks like you have interchanged the limits of integration.</p>
<p>Recall that for definite integral,&nbsp;<span style="font-size: 0.9375rem;">\( \displaystyle \int_{a}^{b}f(x)dx= F(b)-F(a),\) where \(F\) is the antiderivative of f.</span></p>
<p>Take a look at the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks for the second mistake</description>
        <answertest>Int</answertest>
        <sans>ans4</sans>
        <tans>mistake_3</tans>
        <testoptions>x</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have computed an indefinite integral and not failed to add the&nbsp; constant of integration.&nbsp;</p><p><span style="font-size: 0.9375rem;">Take a look at the worked solution below:</span></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>checks for the third mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mistake_4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p> You seem to have computed an indefinite integral and failed to add the constant of integration.</p>
<p>Take a look at the worked solution below:</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>checks for the fourth mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mistake_5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer looks okay but you don't need to add the constant.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt2-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description>checks for the fifth mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mistake_6</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have used the incorrect partial fractions and computed the definite integral correctly.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt2-6-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description>checks for sixth mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mistake_7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-7-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">You seem to have used the incorrect partial fractions, computed the definite integral correctly and added a constant.<br></p><p dir="ltr" style="text-align: left;">Recall that when computing a definite integral, no constant is included.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>7</falsenextnode>
        <falseanswernote>prt2-7-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <description>checks for seventh mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mistake_8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-8-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p dir="ltr" style="text-align: left;">You seem to have submitted the computed partial fractions above without finding the actual integral of the partial fractions as was intended in question (2).<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-8-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>116292304</deployedseed>
    <deployedseed>1968948341</deployedseed>
    <deployedseed>1110945243</deployedseed>
    <deployedseed>1054820811</deployedseed>
    <deployedseed>1056792939</deployedseed>
    <deployedseed>265079445</deployedseed>
    <deployedseed>1762455460</deployedseed>
    <deployedseed>905516863</deployedseed>
    <deployedseed>529257980</deployedseed>
    <deployedseed>1268249061</deployedseed>
    <deployedseed>154174134</deployedseed>
    <deployedseed>1279146257</deployedseed>
    <deployedseed>261585535</deployedseed>
    <deployedseed>17489177</deployedseed>
    <deployedseed>549759992</deployedseed>
    <deployedseed>94137896</deployedseed>
    <deployedseed>1300843071</deployedseed>
    <deployedseed>1744536361</deployedseed>
    <deployedseed>564277745</deployedseed>
    <deployedseed>1292624543</deployedseed>
    <deployedseed>366803082</deployedseed>
    <deployedseed>585682610</deployedseed>
    <deployedseed>992813527</deployedseed>
    <deployedseed>1214413981</deployedseed>
    <deployedseed>1674991464</deployedseed>
    <deployedseed>1968425935</deployedseed>
    <deployedseed>1298813964</deployedseed>
    <deployedseed>1386174787</deployedseed>
    <deployedseed>1898028660</deployedseed>
    <deployedseed>1995400841</deployedseed>
    <deployedseed>1433617956</deployedseed>
    <deployedseed>1070310761</deployedseed>
    <deployedseed>1992412063</deployedseed>
    <deployedseed>1588481161</deployedseed>
    <deployedseed>1535183259</deployedseed>
    <deployedseed>1157456693</deployedseed>
    <deployedseed>687579992</deployedseed>
    <deployedseed>1067045334</deployedseed>
    <deployedseed>543245119</deployedseed>
    <deployedseed>584025393</deployedseed>
    <deployedseed>1661581137</deployedseed>
    <deployedseed>1760912662</deployedseed>
    <deployedseed>1461283420</deployedseed>
    <deployedseed>977193592</deployedseed>
    <deployedseed>1006966896</deployedseed>
    <deployedseed>723456566</deployedseed>
    <deployedseed>2139626217</deployedseed>
    <deployedseed>1173107097</deployedseed>
    <deployedseed>1249567409</deployedseed>
    <deployedseed>1537354051</deployedseed>
    <deployedseed>569036537</deployedseed>
    <deployedseed>1760663647</deployedseed>
    <deployedseed>760093528</deployedseed>
    <deployedseed>91938972</deployedseed>
    <deployedseed>766181145</deployedseed>
    <deployedseed>531248308</deployedseed>
    <deployedseed>202010441</deployedseed>
    <deployedseed>533109351</deployedseed>
    <deployedseed>1697582012</deployedseed>
    <deployedseed>1827663212</deployedseed>
    <deployedseed>1774997666</deployedseed>
    <deployedseed>313517481</deployedseed>
    <deployedseed>87847900</deployedseed>
    <deployedseed>1753851776</deployedseed>
    <deployedseed>437337921</deployedseed>
    <deployedseed>979165667</deployedseed>
    <deployedseed>1859408501</deployedseed>
    <deployedseed>1476429664</deployedseed>
    <deployedseed>1486678088</deployedseed>
    <deployedseed>1531376405</deployedseed>
    <deployedseed>378432354</deployedseed>
    <deployedseed>1727404999</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
