<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Evaluating an Indefinite Integral by Partial Fractions 2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Evaluate the integral \(\displaystyle\int{@question_expression@}dx\).</p>
<p><em>Note: First we use long division since the degree of the numerator is greater than the degree of the denominator. By long division, we can identify the polynomial part and that of the remainder fraction. </em></p>
<p>Polynomial part: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]] </p>
<p>Remainder fraction part: [[input:ans2]] [[validation:ans2]] </p>
<p>\(\displaystyle\int{@question_expression@}dx=\) [[input:ans4]][[validation:ans4]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To evaluate the following integral \(\displaystyle\int{@question_expression@}dx\), we need two consecutive steps.</p>
<p><strong><span>First step:</span> The integrand has to be expressed as a sum of its partial fractions.</strong></p>
<p>To write \( \displaystyle {@question_expression@} \) as a sum of its partial fractions we need to follow the steps below:</p>
<p>First use long division to write \( \displaystyle {@question_expression@}\) as a sum of a polynomial and a rational function. That is</p>
<p>\[{@question_expression@} ={@ta_3@}+{@expr5@}\].</p>
<p>Next is to write the rational function\( \displaystyle {@expr5@}\) as a sum of partial fractions. That is \( \displaystyle{@expr5@}={@part_frac_A@}+{@part_frac_B@}\) </p>
<p>Setting the numerators equal gives:</p>
<p>\[{@expr@}\]</p>
<p>Getting the values of A and B, we have to equate the coefficients of \(x\) and the constants in the equation above.</p>
<p>This should give \[{@expr5@}={@ta_1@}+{@ta_2@}\].</p>
<p>The partial fraction of the integrand is then,</p>
<p>\[{@question_expression@}={@ta_1a@}\]</p>
<p><strong>Second step: Evaluating the integral.</strong></p>
<p>From the step 1, we have:</p>
<p>\[ \displaystyle\int {@question_expression@}dx = \displaystyle\int{@ta_1@}\,dx +  \displaystyle\int{@ta_2@}\,dx + \displaystyle\int{@ta_3@}\,dx \]</p>
<p>\[ \iff \displaystyle\int {@question_expression@}dx = {@int_part1@}+ {@int_part2@}+ {@int_part3@}+c\]</p>
<p>Therefore,</p>
<p>\[ \displaystyle\int {@question_expression@}dx = {@ta4@}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Defining the question*/
aa:2+ rand(3);
bb: 1+rand(2);
const_1:1+rand(5);/*The first constant coefficient*/
const_2:1+rand(5);/*The second constant*/
const_3:1+rand(4);/*The third constant*/
const_4:1+rand(8);/*The fourth constant*/
zz1:rand_with_prohib(-7,7,[0]);/* numerator_1*/
zz2:rand_with_prohib(-7,7,[0]);/* numerator_2*/
zz3:rand_with_prohib(-7,7,[0]);/* numerator_2*/
part_frac_1:zz1/x;/*The first Partial fraction*/
part_frac_2:zz2/(x+const_1*const_2);/*The second Partial fraction*/
part_frac_3: x^aa+zz3*x^bb+const_3;/*The third Partial fraction*/
question_expression: factor(part_frac_1+part_frac_2+part_frac_3);

/*model answer question 1*/
ta_1:part_frac_1;
ta_2:part_frac_2;
ta_3:part_frac_3;
ta_4: ta_1+ta_2+ta_3;
ta_1a: partfrac(question_expression,x);

/*Worked solution codes for question 1 */
part_frac_A:A/x;
part_frac_B:B/(const_1*x+const_2);
part_frac_C:C;
expr5: factor(ta_1+ta_2);
numerator_question_expression:num(expr5);
expr: A*(const_1*x+const_2)+B*x=numerator_question_expression;

/*Find A and B/*
expr1:ev(expr,x=0);
expA:expr1/ev(((const_1*x+const_2)),x=0);
expr2:ev(expr,x=-const_3);
expC:expr2/ev((x*(const_1*x+const_2)),x=-const_3);
coef_B:ev((-const_2/const_1),simp);
expr3b:ev(expr,x=coef_B);
expB:expr3b/ev(x*(x+const_3),x=(-const_2/const_1));

expa:question_expression=ta_4;

/*model answer part b*/
int_part1:int(ta_1,x);
int_part2:int(ta_2,x);
int_part3:int(ta_3,x);
ta1: ta_3;
ta2: expr5;
ta4:int_part1+int_part2+int_part3+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \(\displaystyle\int{@question_expression@}dx\).</p>
<p> First we use long division since the degree of the numerator is greater than the degree of the denominator. By long division, the polynomial part is \({@ta1@}\) and the remainder fraction part is \({@ta2@}\)</p>
<p>\(\displaystyle\int{@question_expression@}dx={@ta4@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">The question focuses on finding partial fraction from a given function</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>30</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer */
/*Node 2: Checks for the incorrect answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the incorrect answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1:checking for correct answers*/

/*node 2:forgetting the constant of integration*/
mistake_2:int_part1+int_part2+int_part3;

/*node 3:Giving the answer in partial fraction form without completing the integral*/
mistake_3:ta_1+ta_2+ta_3;
/*node 4:Computing correct integral with wrong answers in part (1)*/
question_4:ans1+ans2+ans3;
int_question1:int(question_4,x)+c;
mistake_4:int_question1;
/*node 5:Computing correct integral with wrong answers in part (1) and forgetting the constant of integration*/
int_question2:int(question_4,x);
mistake_5:int_question2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checking for correct answers</description>
        <answertest>Int</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>forgetting the constant of integration</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p> </p>
<p>You seem to have to calculated the integral correctly but failed to add the constant of integration.</p>
<p> </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Giving the answer in partial fraction form without completing the integral</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mistake_3</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have submitted the computed partial fractions above without finding the actual integral of the partial fractions as was intended in (2).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Computing correct integral with wrong answers in part (1)</description>
        <answertest>Int</answertest>
        <sans>ans4</sans>
        <tans>mistake_4</tans>
        <testoptions>x</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p style="text-align: left;" dir="ltr">You seem to have used the wrong answers in (1) and correctly computed the integral with the wrong answers.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Computing correct integral with wrong answers in part (1) and forgetting the constant of integration</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mistake_5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p style="text-align: left;" dir="ltr">You seem to have used the wrong answers in (1) and correctly computed the integral with the wrong answer but failed to add the constant of integration.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>662082801</deployedseed>
    <deployedseed>1346184093</deployedseed>
    <deployedseed>435903118</deployedseed>
    <deployedseed>2002137319</deployedseed>
    <deployedseed>1130039361</deployedseed>
    <deployedseed>1885644658</deployedseed>
    <deployedseed>1385900969</deployedseed>
    <deployedseed>427031026</deployedseed>
    <deployedseed>1365273997</deployedseed>
    <deployedseed>1733873179</deployedseed>
    <deployedseed>1768778806</deployedseed>
    <deployedseed>175884057</deployedseed>
    <deployedseed>1248860732</deployedseed>
    <deployedseed>1607741241</deployedseed>
    <deployedseed>443674508</deployedseed>
    <deployedseed>554457403</deployedseed>
    <deployedseed>1116806218</deployedseed>
    <deployedseed>1088479984</deployedseed>
    <deployedseed>599649999</deployedseed>
    <deployedseed>1548674450</deployedseed>
    <deployedseed>1447465301</deployedseed>
    <deployedseed>1588126702</deployedseed>
    <deployedseed>1985775474</deployedseed>
    <deployedseed>2124327243</deployedseed>
    <deployedseed>1911077452</deployedseed>
    <deployedseed>1905257492</deployedseed>
    <deployedseed>2046379358</deployedseed>
    <deployedseed>1318446793</deployedseed>
    <deployedseed>1425361039</deployedseed>
    <deployedseed>1111118395</deployedseed>
    <deployedseed>111522123</deployedseed>
    <deployedseed>1457003707</deployedseed>
    <deployedseed>1536789533</deployedseed>
    <deployedseed>339351611</deployedseed>
    <deployedseed>1972906631</deployedseed>
    <deployedseed>953530958</deployedseed>
    <deployedseed>1095522697</deployedseed>
    <deployedseed>1809975976</deployedseed>
    <deployedseed>563465212</deployedseed>
    <deployedseed>1760116667</deployedseed>
    <deployedseed>386002870</deployedseed>
    <deployedseed>1056182447</deployedseed>
    <deployedseed>1660337025</deployedseed>
    <deployedseed>867534711</deployedseed>
    <deployedseed>1765053797</deployedseed>
    <deployedseed>426904047</deployedseed>
    <deployedseed>951226266</deployedseed>
    <deployedseed>1061303436</deployedseed>
    <deployedseed>394192891</deployedseed>
    <deployedseed>153042990</deployedseed>
    <deployedseed>902969587</deployedseed>
    <deployedseed>1642156175</deployedseed>
    <deployedseed>910563066</deployedseed>
    <deployedseed>117978450</deployedseed>
    <deployedseed>2088953413</deployedseed>
    <deployedseed>1075351622</deployedseed>
    <deployedseed>1297548912</deployedseed>
    <deployedseed>1428762872</deployedseed>
    <deployedseed>1187829209</deployedseed>
    <deployedseed>2120535214</deployedseed>
    <deployedseed>840818076</deployedseed>
    <deployedseed>173492972</deployedseed>
    <deployedseed>1054205068</deployedseed>
    <deployedseed>429165428</deployedseed>
    <deployedseed>1826971454</deployedseed>
    <deployedseed>1584871249</deployedseed>
    <deployedseed>972343622</deployedseed>
    <deployedseed>476812907</deployedseed>
    <deployedseed>761785627</deployedseed>
    <deployedseed>1987412610</deployedseed>
    <deployedseed>690555717</deployedseed>
    <deployedseed>1424027135</deployedseed>
    <deployedseed>1477997331</deployedseed>
    <deployedseed>1021484817</deployedseed>
    <deployedseed>1888834193</deployedseed>
    <deployedseed>1047236876</deployedseed>
    <deployedseed>1385655051</deployedseed>
    <deployedseed>1064882636</deployedseed>
    <deployedseed>1849890286</deployedseed>
    <deployedseed>480596465</deployedseed>
    <deployedseed>23210415</deployedseed>
    <deployedseed>1471683434</deployedseed>
    <deployedseed>1609413846</deployedseed>
    <deployedseed>1471474977</deployedseed>
    <deployedseed>302498922</deployedseed>
    <deployedseed>1259232977</deployedseed>
    <deployedseed>65756085</deployedseed>
    <deployedseed>935706159</deployedseed>
    <deployedseed>16427282</deployedseed>
    <deployedseed>1276031376</deployedseed>
    <deployedseed>2097657006</deployedseed>
    <deployedseed>436751313</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
