<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Evaluating limits of functions (c)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Evaluate the folowing limit:</p>
<p> \(\displaystyle \lim_{x\rightarrow 0} = {@fx@}=\)[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>(a) When we try to evaluate \(f(0)\), both the numerator and the denominator of the fraction evaluate to \(0\). 
<br>Therefore, \(\displaystyle  f(0) = \frac{0}{0}\), which is <b>undefined</b>.</p>

<br>

<p>(b) We evaluate \(f  \displaystyle  (x) = {@fx@}\) at the relevant points to give the following results, rounded to 2 decimal places.

<br>

<p hidden>{@stackfltfmt:"~,2f";@}</p>
<table style="margin-left: 30px">
<tr>
  <td style="text-align: center; border: 1px solid gray; width:60px;">\(x\)</th>
  <td style="text-align: center; border: 1px solid gray; width:100px;">\(-0.25\)</th>
  <td style="text-align: center; border: 1px solid gray; width:100px;">\(-0.1\)</th>
  <td style="text-align: center; border: 1px solid gray; width:100px;">\(-0.01\)</th>
  <td style="text-align: center; border: 1px solid gray; width:100px;">\(0.01\)</th>
  <td style="text-align: center; border: 1px solid gray; width:100px;">\(0.1\)</th>
  <td style="text-align: center; border: 1px solid gray; width:100px;">\(0.25\)</th>
</tr>
<tr style="height:40px">
  <td style="text-align: center; border: 1px solid gray; ">\(f(x)\)</td>
  <td style="text-align: center; border: 1px solid gray; ">{@ta2@}</td>
  <td style="text-align: center; border: 1px solid gray; ">{@ta3@}</td>
  <td style="text-align: center; border: 1px solid gray; ">{@ta4@}</td>
  <td style="text-align: center; border: 1px solid gray; ">{@ta5@}</td>
  <td style="text-align: center; border: 1px solid gray; ">{@ta6@}</td>
  <td style="text-align: center; border: 1px solid gray; ">{@ta7@}</td>
</tr>
</table>

<br>

<p>(c) We use the approximations determined in part (b). 
<br>As \(x \rightarrow 0\) from both the positive and negative direction, the value of \(f(x)\) approaches {@aa@}.
<br>Therefore, \( \displaystyle \lim_{x\rightarrow 0} f(x) = {@aa@}\)</p>

<p>This is particularly clear if we consider the graph of \(y = f(x)\)</p>
<p style="margin-left: 30px">{@plot(
    fx,
    [x,-3*%pi/aa,3*%pi/aa],
    [y, -aa, aa+2],
    [box, false],
    [plottags, false]
)@}</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[aa:rand(5)+3;
fx:sin(aa*x)/x;

ma1:random_permutation([
  [0, true, "\\(\\text{undefined}\\)"],
  [1, false, "\\(0\\)"],
  [2, false, "\\(1\\)"],
  [3, false, "\\( \\infty \\)"],
  [4, false, "\\(-\\infty\\)"]
]);

ta2:decimalplaces(ev(fx, x = -0.25),2);
ta3:decimalplaces(ev(fx, x = -0.1),2);
ta4:decimalplaces(ev(fx, x = -0.01),2);
ta5:decimalplaces(ev(fx, x = 0.01),2);
ta6:decimalplaces(ev(fx, x = 0.1),2);
ta7:decimalplaces(ev(fx, x = 0.25),2);

ta8:aa;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(f(x) = {@fx@}\)</p>
<p>\(\lim_{x\rightarrow 0} f(x) = {@aa@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey Wabwire </p>
<p id="concept">NA</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">This question involves evaluating a given function at a specific value, filling in a table of values for the function, and finding the limit of the function as x approaches 0. The question requires providing answers correct to 3 significant figures.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta8</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>0</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>When we try to evaluate \(f(0)\), both the numerator and the denominator of the fraction evaluate to \(0\). 
<br>Therefore, \(f(0) = \frac{0}{0}\), which is <b>undefined</b>.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1313167693</deployedseed>
    <deployedseed>1187307701</deployedseed>
    <deployedseed>685727649</deployedseed>
    <deployedseed>1835028593</deployedseed>
    <deployedseed>1373401442</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ma1))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
