<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Evaluating an Integral by Mean Value Theorem</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the average value of the function  \(f(x)={@quest_expr@}\) and find the value of \(c\)  that satisfies the mean value theorem for \(f(x)\) in the interval \([{@value_a@},{@value_b@}]\).</p>
<p>Average value \(=\)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p></p>
<p>\(c=\)[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To find the average value of the function \(f(x)={@quest_expr@}\) over the interval \([{@value_a@},{@value_a@}],\)  you can use the following formula:</p>

<p>Average value \(=\displaystyle \frac{1}{b-a}\int_a^b f(x)dx\)</p>
<p>In this case \( ,a={@value_a@}\) and \(b={@value_b@} \) </p>

<p>Average value \(=\frac{1}{{@value_b@}-{@value_a@}}\displaystyle \int _{{@value_a@}}^{{@value_b@}} ({@quest_expr@}) dx\)</p>
<p>Intergrating \({@quest_expr@}\) with respect to \(x\) we get;</p>
<p>Average value \(=\frac{1}{{@value_b-value_a@}}  \left [  {@func_1@} \right ]_{@value_a@}^{@value_b@} \)</p>
<p>Evaluate this expression by plugging in \({@value_b@}\) and \({@value_a@}\) for \(x\):</p>
<p>Average value \(=\frac{1}{{@value_b-value_a@}}  \left [  {@constant@} \cdot {@value_b@}-{@coeff_2/2@} \cdot {@value_b@}^{2} \right ] \)</p>
<p>\(\qquad\qquad \qquad= \frac{1}{{@value_b-value_a@}}  \left [  {@constant*value_b@}-{@(coeff_2/2)*value_b^2@} \right ]\)</p>
<p>\(\qquad\qquad \qquad=\frac{1}{{@value_b-value_a@}}  \left [  {@(constant*value_b)-((coeff_2/2)*value_b^2)@} \right ]\)</p>
<p> \(\qquad\qquad \qquad={@ta1@}\)</p>
<p>Now, let's find the value of \(c\) suvh that \(f(c)\) equals the average value of the function over the interval \([{@value_a@},{@value_a@}]\)</p>
<p>\({@func_3@}={@ta1@}\)<p>
<p>\({@-coeff_2*c@}={@ta1@}-{@constant@}\)<p>
<p>\({@-coeff_2*c@}={@ta1-constant@}\)<p>
<p>\( \qquad c={@ta2@}\)<p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Question variables*/
constant:rand(5)+5;
coeff_2:rand(4)+2;
quest_expr:constant-coeff_2*x;
value_a:0;
value_b:rand(5)+1;

func_1:int(quest_expr,x);
func_2:int(quest_expr,x,value_a,value_b);
ta1:ratsimp(((1/value_b)*func_2));
func_3:subst(c,x,quest_expr);
ta2:rhs(solve(func_3=ta1)[1]);
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(f(x)={@quest_expr@}\) </p>
<p>Average value \(f(c)={@ta1@}\)</p>
<p>\(c={@ta2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept"></p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">Juma Zevick</p>
<p id="description">Finding the average value of the function \(f(x)={@quest_expr@}\) over the interval \([{@value_a@},{@value_a@}]\) and find \(c\) such that \(f(c)\) equals the average value of the function over ([{@value_a@},{@value_a@}]\).</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2110282285</deployedseed>
    <deployedseed>1473940864</deployedseed>
    <deployedseed>1642242965</deployedseed>
    <deployedseed>554738734</deployedseed>
    <deployedseed>799142408</deployedseed>
    <deployedseed>722093573</deployedseed>
    <deployedseed>157717927</deployedseed>
    <deployedseed>1371754897</deployedseed>
    <deployedseed>284750814</deployedseed>
    <deployedseed>552078089</deployedseed>
    <deployedseed>1624852362</deployedseed>
    <deployedseed>673834349</deployedseed>
    <deployedseed>543381488</deployedseed>
    <deployedseed>1415505286</deployedseed>
    <deployedseed>1090307945</deployedseed>
    <deployedseed>571657051</deployedseed>
    <deployedseed>981914185</deployedseed>
    <deployedseed>1945896676</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
