<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Find the function  f(x) form dydx</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the function \( f(x) \) given \( \displaystyle {@fx@} \)</p>
<p>\( f(x) =\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Given \( \displaystyle \frac{dy}{dx} = {@expr@} \),</p>

<p>Integrating both sides with respect to \( x \),</p>

<p>\( \displaystyle  y ={@ta@} \)</p>

<p>where \( C \) is the constant of integration.</p>

<p>Therefore, \(  \displaystyle f(x) = {@ta@} \) is the function \( f(x) \) that satisfies \(  \displaystyle \frac{dy}{dx} = {@expr@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/* Define the differential equation */
dff_operator: dy/dx;

/* Solve the differential equation */
power:rand(100)+4;
expr:x^power;
fx: dff_operator=expr;

ta:int(expr,x)+C;
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Find the function \( f(x) \) given \( \displaystyle{@fx@} \)</p>
<p>\( f(x)={@ta@} \)</p>

]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1565845234</deployedseed>
    <deployedseed>787147143</deployedseed>
    <deployedseed>377083245</deployedseed>
    <deployedseed>842128283</deployedseed>
    <deployedseed>1596009419</deployedseed>
    <deployedseed>1235332443</deployedseed>
    <deployedseed>1973405240</deployedseed>
    <deployedseed>79596523</deployedseed>
    <deployedseed>2103930736</deployedseed>
    <deployedseed>1553515705</deployedseed>
    <deployedseed>184657497</deployedseed>
    <deployedseed>2042097002</deployedseed>
    <deployedseed>1294741516</deployedseed>
    <deployedseed>1459153988</deployedseed>
    <deployedseed>1178891104</deployedseed>
    <deployedseed>1172529332</deployedseed>
    <deployedseed>774761577</deployedseed>
    <deployedseed>996211272</deployedseed>
    <deployedseed>1495146991</deployedseed>
    <deployedseed>1185222621</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
