<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the initial deposit given the equation of the investment account.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>After a certain number of years, the value of an investment account is represented by the equation</p>
 <p>\[\displaystyle {@P@}\left(1+\frac{{@annual_rate@}}{{@n@}}\right)^{{@N@}}\]</p>
<p>What was the initial deposit made to the account?</p>
<p>  \( \mbox{Initial Principal} =\mbox{N}\$ \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p> \( \mbox{The initial deposit is } \mbox{N}\${@P@} \) </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/* Define the initial principal */
P: (rand(10000)+50)*10;

/* Define the annual interest rate */
annual_rate: decimalplaces((rand(5)+4)/100,2);

/* Number of times interest is compounded per year */
n: 12;

/* Total number of compounding periods (10 years) */
t: rand(10)+10;
N: n * t;

/* Calculate the monthly interest rate */
monthly_rate: annual_rate / n;

/* Define the formula for compound interest */
V: P * (1 + monthly_rate)^N;

/* Calculate the value */
V_value: decimalplaces(float(ev(V)),2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\[\displaystyle {@P@}\left(1+\frac{{@annual_rate@}}{{@n@}}\right)^{{@N@}}\]</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="concept">https://docs.google.com/document/d/1kGxoRJBOzT1MCttOxfxB7yk5ZsqXZc6AcDrWF9s627Y/edit</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">2</p>
<p id="description">Given the equation of the investment account, students are required to caculate the initial deposit made to the account. </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>P</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>P</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>867816937</deployedseed>
    <deployedseed>1725350167</deployedseed>
    <deployedseed>1980621984</deployedseed>
    <deployedseed>1907822376</deployedseed>
    <deployedseed>1872809096</deployedseed>
    <deployedseed>1985557222</deployedseed>
    <deployedseed>1876855555</deployedseed>
    <deployedseed>1629829808</deployedseed>
    <deployedseed>1049202082</deployedseed>
    <deployedseed>381823195</deployedseed>
    <deployedseed>168489825</deployedseed>
    <deployedseed>5238946</deployedseed>
    <deployedseed>1966872216</deployedseed>
    <deployedseed>1055619504</deployedseed>
    <deployedseed>1401482715</deployedseed>
    <deployedseed>1895188276</deployedseed>
    <deployedseed>554899371</deployedseed>
    <deployedseed>1541646447</deployedseed>
    <deployedseed>2074327411</deployedseed>
    <deployedseed>356062635</deployedseed>
    <deployedseed>1493172235</deployedseed>
    <deployedseed>618583994</deployedseed>
    <deployedseed>417439478</deployedseed>
    <deployedseed>924895607</deployedseed>
    <deployedseed>1247716242</deployedseed>
    <deployedseed>1223202395</deployedseed>
    <deployedseed>1287075453</deployedseed>
    <deployedseed>406452711</deployedseed>
    <deployedseed>1201000532</deployedseed>
    <deployedseed>324553945</deployedseed>
    <deployedseed>191787231</deployedseed>
    <deployedseed>1169735624</deployedseed>
    <deployedseed>1989951360</deployedseed>
    <deployedseed>1610615112</deployedseed>
    <deployedseed>459453091</deployedseed>
    <deployedseed>802569468</deployedseed>
    <deployedseed>57008521</deployedseed>
    <deployedseed>2042762754</deployedseed>
    <deployedseed>54414860</deployedseed>
    <deployedseed>376669916</deployedseed>
    <deployedseed>2144511934</deployedseed>
    <deployedseed>1120134063</deployedseed>
    <deployedseed>1547357342</deployedseed>
    <deployedseed>1620323471</deployedseed>
    <deployedseed>569363398</deployedseed>
    <deployedseed>1096499551</deployedseed>
    <deployedseed>536552645</deployedseed>
    <deployedseed>1347869073</deployedseed>
    <deployedseed>283827963</deployedseed>
    <deployedseed>1767950274</deployedseed>
    <deployedseed>79792990</deployedseed>
    <deployedseed>486284732</deployedseed>
    <deployedseed>578719518</deployedseed>
    <deployedseed>361580802</deployedseed>
    <deployedseed>1745915003</deployedseed>
    <deployedseed>192423965</deployedseed>
    <deployedseed>580440677</deployedseed>
    <deployedseed>1916113906</deployedseed>
    <deployedseed>867836291</deployedseed>
    <deployedseed>1370808721</deployedseed>
    <deployedseed>1724591661</deployedseed>
    <deployedseed>2141129707</deployedseed>
    <deployedseed>1728202415</deployedseed>
    <deployedseed>1461683976</deployedseed>
    <deployedseed>513437556</deployedseed>
    <deployedseed>1999276806</deployedseed>
    <deployedseed>1242405663</deployedseed>
    <deployedseed>1416673297</deployedseed>
    <deployedseed>555918570</deployedseed>
    <deployedseed>1742319658</deployedseed>
    <deployedseed>720176699</deployedseed>
    <deployedseed>1177105200</deployedseed>
    <deployedseed>1902214855</deployedseed>
    <deployedseed>567298363</deployedseed>
    <deployedseed>2069611122</deployedseed>
    <deployedseed>1029333586</deployedseed>
    <deployedseed>703251171</deployedseed>
    <deployedseed>1912565330</deployedseed>
    <deployedseed>1530142560</deployedseed>
    <deployedseed>188286101</deployedseed>
    <deployedseed>458973262</deployedseed>
    <deployedseed>1800044941</deployedseed>
    <deployedseed>1355117488</deployedseed>
    <deployedseed>423552631</deployedseed>
    <deployedseed>650960617</deployedseed>
    <deployedseed>1094729241</deployedseed>
    <deployedseed>882434032</deployedseed>
    <deployedseed>406193681</deployedseed>
    <deployedseed>1322872468</deployedseed>
    <deployedseed>1082803655</deployedseed>
    <deployedseed>1522825004</deployedseed>
    <deployedseed>1845084434</deployedseed>
    <deployedseed>449947613</deployedseed>
    <deployedseed>286172478</deployedseed>
    <deployedseed>669806946</deployedseed>
    <deployedseed>653749844</deployedseed>
    <deployedseed>764507445</deployedseed>
    <deployedseed>512763505</deployedseed>
    <deployedseed>729378603</deployedseed>
    <deployedseed>83062812</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>P</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
