<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Find the derivative of the logarithmic function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Find the derivative of the function \( \displaystyle f(x) = {@t1@} + {@t2@} + {@t3@} + {@t4@}\).</p>
<p>\(f'(x) = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Here, we will differentiate one term at a time for simplicity.</p>
<p>If \( \displaystyle y = {@t1@}\), we have \(y' = 0\), since this is a constant function.</p>
<p>If \( \displaystyle y = {@t2@}\), then \(\ln y = \ln({@bb@}^x) = x\cdot \ln({@bb@})\).</p>
<p> Differentiating both sides implicitly we get \( \displaystyle \frac{1}{y}\cdot y' = \ln({@bb@})\) and so \(y' = y\cdot \ln({@bb@}) = {@t2@} \cdot \ln({@bb@})\).</p>
<p>If \( \displaystyle y = x^{@cc@}\), then \(y' = {@cc*x^(cc-1)@}\).</p>
<p>If \( \displaystyle y= x^{1/x}\), then \(\ln(y) = \ln(x^{1/x}) = \frac{1}{x}\cdot \ln(x)\).</p>
<p> Differentiating implicitly on both sides we get \( \displaystyle \frac{1}{y}\cdot y' = {@-1/x^2@} \cdot \ln(x) + {@1/x@}\cdot {@1/x@} = {@-1/x^2*log(x)+1/x^2@}\). </p>
<p>Therefore, \( \displaystyle y' = y\cdot \big( {@-1/x^2*log(x)+1/x^2@}\big) = {@t4@}\cdot \big( {@-1/x^2*log(x)+1/x^2@}\big) \).</p>
<p>Therefore, \( \displaystyle f'(x) = {@t2@} \cdot \ln({@bb@}) + {@cc*x^(cc-1)@} + {@t4@}\cdot \big( {@-1/x^2*log(x)+1/x^2@}\big) \).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand([2,3]);
bb:rand([2,3,4,5]);
cc:rand([2,3,4]);
t1:pi^aa;
t2:bb^x;
t3:x^cc;
t4:x^(1/x);
fx:t1+t2+t3+t4;
ta:t2*log(bb)+cc*x^(cc-1)+t4*(-1/x^2*log(x)+1/x^2);
ta:diff(fx,x);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p></p><p>Find the derivative of the function \( \displaystyle f(x) = {@t1@} + {@t2@} + {@t3@} + {@t4@}\).</p>
<p>\(f'(x) = {@ta@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="concept">N/A</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">Find the derivative of the logarithmic function.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>Diff</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>110039515</deployedseed>
    <deployedseed>1780520019</deployedseed>
    <deployedseed>243529721</deployedseed>
    <deployedseed>1033043773</deployedseed>
    <deployedseed>2108104416</deployedseed>
    <deployedseed>570112489</deployedseed>
    <deployedseed>179917912</deployedseed>
    <deployedseed>1288999799</deployedseed>
    <deployedseed>259845745</deployedseed>
    <deployedseed>2050554470</deployedseed>
    <deployedseed>1401647110</deployedseed>
    <deployedseed>1741143325</deployedseed>
    <deployedseed>2079958621</deployedseed>
    <deployedseed>1947727146</deployedseed>
    <deployedseed>1852366965</deployedseed>
    <deployedseed>1943772162</deployedseed>
    <deployedseed>1614628315</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
