<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Find Volume of the cylinder- Q1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Use the method of cylinders to determine the volume of the solid obtained by rotating the region bounded by {@eq1@}, the \(x\)-axis and the \(y\)-axis about the \(x\)-axis.</p>
<p>Here is a sketch of the bounded region with the axis of rotation shown.</p>
<p>
[[jsxgraph]]
        var board = JXG.JSXGraph.initBoard(divid, { 
            boundingbox: [-1, 6, 6, -1],
            keepaspectratio: true,
            showCopyright: false,
            axis: true,
            showNavigation: true
        });
    // Plot functions only up to the mutual intersections
        var function1 = board.create('functiongraph', ['{#func#}', 0, {#xx#}], { name: '{#simplify(x=((y-root)^2)/coeff_x)#}', withlabel: true, strokeColor: 'blue' });
        //var function2 = board.create('functiongraph', ['(x-1)^2+1', 0, 2]);
        // Note: This function goes backwards, i.e. from 3 to -3.
        //var function4 = board.create('functiongraph', ['(x/3)^2-1', 3, -3]);

        // You can copy this function and then use it like in the examples below.
        // element_list is a list of elements defining the outline of the function.
        // The elements can be FunctionGraphs or Pairs of coordinates (encoding points)
        // properties is a dictionary of properties for the generated function.
        function make_filled_curve(board, element_list, properties) {
            var curve = board.create('curve', [[], []], properties);
            curve.updateDataArray = function () {
                this.dataX = [];
                this.dataY = [];
                for (const element of element_list) {
                    if (Array.isArray(element)) {
                        // Add the point to the curve
                        this.dataX.push(element[0]);
                        this.dataY.push(element[1]);
                    } else {
                        // Copy all points from function
                        this.dataX = this.dataX.concat(element.points.map(p => p.usrCoords[1]));
                        this.dataY = this.dataY.concat(element.points.map(p => p.usrCoords[2]));
                    }
                }
                // Close the loop
                this.dataX.push(this.dataX[0]);
                this.dataY.push(this.dataY[0]);
            };
            board.update();
        }

        // Plot different filled functions
        //make_filled_curve(board, [function1, function2, function4], { strokeWidth: 3, fillColor: 'yellow', fillOpacity: 0.3 });
        make_filled_curve(board, [function1, [0, 0]], { strokeWidth: 1, fillColor: 'green', fillOpacity: 0.3 });
        //make_filled_curve(board, [function1, [0, 3], [-2, 3]], { strokeWidth: 3, fillColor: 'red', fillOpacity: 0.3 });

       
    [[/jsxgraph]] 
</p>
<p>\(V=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Recall that,</p>
<p>Using the method of cylinders, the volume of the solid generated by a region bounded by the \(y\)- axis, and the horizontal lines \(y=a\) and \(y=b\) which is revolved about the \(x\)-axis(horizontal axis of revolution) is given by </p>
<p>\(V=2{@%pi@}\displaystyle\int_{@a@}^{@b@}yf(y)dy\) and radius \(r=y\)</p>
<p>For this problem, we have \(a=0\), \(b={@root@}\), \(x=f(y)={@f_y@}\) and radius \(r=y\)</p>
<p>Therefore,</p>
<p>
\(
\begin{align}
  V &amp;= 2{@%pi@}\displaystyle\int_0^{@root@}y\cdot{@f_y@}dy\\
  &amp;=2{@%pi@}\displaystyle\int_0^{@root@}\left({@bb@} \right)\,dy\\
  &amp;= 2{@%pi@} \left[ {@cc@} \right]_0^{@root@}\\
  &amp;={@ta@}
\end{align}
\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>qq/*Defining the question*/
coeff_x:rand(3)+1;
root:rand(5)+1;
f_y:((y-root)^2)/coeff_x;
eq1:coeff_x*x=(y-root)^2;
func:rhs(solve(eq1,y)[1]);
bb: expand(y*f_y);
cc:int(bb,y);
area_y: 2*%pi*y*f_y; /*The area of the cylinder*/
volume_y: int(area_y,y,0,root); /*The volume of the solid*/
x_1:ev(eq1,x,y=0);
xx:rhs(solve(x_1)[1]);
x_2:ev(eq1,y,x=0);
yy:rhs(solve(x_2)[1]);

/*Model answer*/
ta:volume_y;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>{@eq1@}, the \(x\)-axis and the \(y\)-axis about the \(x\)-axis \(V= {@volume_y@}\)</p>
<p>
[[jsxgraph]]
        var board = JXG.JSXGraph.initBoard(divid, { 
            boundingbox: [-1, 10, 6, -1],
            keepaspectratio: true,
            showCopyright: false,
            axis: true,
            showNavigation: true
        });
    // Plot functions only up to the mutual intersections
        var function1 = board.create('functiongraph', ['{#func#}', 0, {#xx#}], { name: 'x=(y-2)^2', withlabel: true, strokeColor: 'blue' });
        //var function2 = board.create('functiongraph', ['(x-1)^2+1', 0, 2]);
        // Note: This function goes backwards, i.e. from 3 to -3.
        //var function4 = board.create('functiongraph', ['(x/3)^2-1', 3, -3]);

        // You can copy this function and then use it like in the examples below.
        // element_list is a list of elements defining the outline of the function.
        // The elements can be FunctionGraphs or Pairs of coordinates (encoding points)
        // properties is a dictionary of properties for the generated function.
        function make_filled_curve(board, element_list, properties) {
            var curve = board.create('curve', [[], []], properties);
            curve.updateDataArray = function () {
                this.dataX = [];
                this.dataY = [];
                for (const element of element_list) {
                    if (Array.isArray(element)) {
                        // Add the point to the curve
                        this.dataX.push(element[0]);
                        this.dataY.push(element[1]);
                    } else {
                        // Copy all points from function
                        this.dataX = this.dataX.concat(element.points.map(p => p.usrCoords[1]));
                        this.dataY = this.dataY.concat(element.points.map(p => p.usrCoords[2]));
                    }
                }
                // Close the loop
                this.dataX.push(this.dataX[0]);
                this.dataY.push(this.dataY[0]);
            };
            board.update();
        }

        // Plot different filled functions
        //make_filled_curve(board, [function1, function2, function4], { strokeWidth: 3, fillColor: 'yellow', fillOpacity: 0.3 });
        make_filled_curve(board, [function1, [0, 0]], { strokeWidth: 1, fillColor: 'green', fillOpacity: 0.3 });
        //make_filled_curve(board, [function1, [0, 3], [-2, 3]], { strokeWidth: 3, fillColor: 'red', fillOpacity: 0.3 });

       
    [[/jsxgraph]] 
</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">[Juma Zevick]</p>
<p id="description">The question focuses on finding the volume of a cylinder using a respective formula of finding it's vol</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1:checks for correct answers*/

/*node 2:Integrating the indefinite integral without the limits.*/
mistake_1:int(bb,y);
/*node 3:Interchanging the limits of integration, however, with the correct definite integral*/
mistake_2:%pi*int(bb,y,2,0);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answers</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Integrating the indefinite integral without the limits.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Interchanging the limits of integration, however, with the correct definite integral</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1933669707</deployedseed>
    <deployedseed>19346350</deployedseed>
    <deployedseed>371158560</deployedseed>
    <deployedseed>1598238477</deployedseed>
    <deployedseed>734404062</deployedseed>
    <deployedseed>1861098003</deployedseed>
    <deployedseed>1223090752</deployedseed>
    <deployedseed>2092200365</deployedseed>
    <deployedseed>1785023670</deployedseed>
    <deployedseed>360924071</deployedseed>
    <deployedseed>1886286339</deployedseed>
    <deployedseed>1235876219</deployedseed>
    <deployedseed>1124759593</deployedseed>
    <deployedseed>175852723</deployedseed>
    <deployedseed>260135189</deployedseed>
    <deployedseed>1628281915</deployedseed>
    <deployedseed>1881147442</deployedseed>
    <deployedseed>570966163</deployedseed>
    <deployedseed>2135906708</deployedseed>
    <deployedseed>1796068066</deployedseed>
    <deployedseed>1821598320</deployedseed>
    <deployedseed>1825421481</deployedseed>
    <deployedseed>640184538</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
