<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the Domain and Value of the Logarithmic Function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the function \(f(x)=\log (x)\):</p>
<p>i. Is \(x={@valx@}\) in the domain of  \(f(x)\)? [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>ii. \(\text{Dom}(f)=\) [[input:ans2]]. Therefore, the domain is [[input:ans3]] </p>
<p></p>
<p> [[validation:ans2]][[validation:ans3]][[feedback:prt2]][[feedback:prt3]]</p>

<p><em><b>Note:</b> To enter intervals of real line, you should write the relevant text in the table below.</em></p>

<table width="250" cellspacing="1" cellpadding="0" border="1" style="margin-left: 30px">
  <tr>
    <th style="text-align: center;">Intervals</th>
    <th style="text-align: center;">Input Syntax</th>
  </tr>
  <tr>
    <td style="text-align: center;">\( [a,b] \)</td>
    <td style="text-align: center;"><code>cc(a,b)</code></td>
  </tr>
  <tr>
    <td style="text-align: center;">\( (a,b) \)</td>
    <td style="text-align: center;"><code>oo(a,b)</code></td>
  </tr>
  <tr>
    <td style="text-align: center;">\( [a,b) \)</td>
    <td style="text-align: center;"><code>co(a,b)</code></td>
  </tr>
  <tr>
    <td style="text-align: center;">\( (a,b] \)</td>
    <td style="text-align: center;"><code>oc(a,b)</code></td>
  </tr>
 <tr>
    <td style="text-align: center;">\( (a, \infty) \)</td>
    <td style="text-align: center;"><code>oo(a,inf)</code></td>
  </tr>
 <tr>
    <td style="text-align: center;">\( (-\infty, b] \)</td>
    <td style="text-align: center;"><code>oc(-inf,b)</code></td>
  </tr>
 <tr>
    <td style="text-align: center;">\( (-\infty, \infty) \)</td>
    <td style="text-align: center;"><code>oo(-inf,inf)</code></td>
  </tr>
</table>
</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>
[[ if test='valx=valx_1' ]]
<br> For the function \(f(x)=\log (x), \) the domain of \(f(x)\) is \(x>0, \) thus \(x={@valx@}\) is the vertical asymptote. \(f({@valx@})=\log ({@valx@})\) which is undefined. 

[[ else ]]
<br> For the function \(f(x)=\log (x), \) the domain of \(f(x)\) is \({@ta_2@}, \) thus, the value of the function \(f({@valx@})={@val_valx_2@} = {@val_2@}\) which  is well-defined. 
[[/ if]]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[valx_1:0;
valx_2:rand(8)+1;
valx:rand([valx_1,valx_2]);
ta1: if valx=valx_2 then true else false;
val_valx_2:log(valx_2);
val_2:decimalplaces(float(val_valx_2),2);
ta_2:oo(0,inf);
ta2: if valx=valx_2 then ta_2 else x>0;

ta3_1:random_permutation([["undefined ", true],["well-defined", false]]);
ta3_2:random_permutation([["well-defined", true],["undefined", false]]);

ta3: if valx=valx_1 then ta3_1  else ta3_2;

]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(x={@valx@}\) </p>
<p>\( Dom(f) = {@ta2@}\). Therefore, the domain is \({@mcq_correct(ta3)[1]@}\) </p>
]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">https://docs.google.com/document/d/1htPOpEhhfKurf0bd3W3GARCAYb73KaHbx12kqPUxY9U/edit#heading=h.eddp73mos3v3</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer"><</p>
<p id="description">In this question, students will explore the domain of the function \(f(x)=log(x) \) to determine if \(x=0\) is included. If so, ascertain the value of \f(x) \) at \(x=0\) and validate the outcome.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>boolean</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mcq_correct(ta3)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1056848638</deployedseed>
    <deployedseed>62423388</deployedseed>
    <deployedseed>1524464892</deployedseed>
    <deployedseed>1994762540</deployedseed>
    <deployedseed>2012701328</deployedseed>
    <deployedseed>1439066584</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(ta3))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
