<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding growth of compounded interest</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>An account is opened with an initial deposit of \(\text{N\$ }{@principal@} \) and earns 
\({@rate@}\%\) interest compounded semi-annually.</p>
<p>What will be the account worth in \({@years@}\) years.</p>
<p>Amount \(= \text{N\$ } \)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Amount is calculated using the formula;</p>
<p>\(A=P(1+\displaystyle \frac{r}{n})^{\text{nt}}\)</p>
<p>where,</p>
<p> A= amount=?</p>
<p> P = principal= \({@principal@}\)</p>
<p>r = rate of interest=\({@rate_float@}\)</p>
<p>n = number of times interest is compounded per year=semi-annaully</p>
<p>t = time (in years)\(={@years@}\)</p>

<p>Substituting values to the formula we have</p>
<p>\(\begin{align}
A&=P(1+\displaystyle \frac{r}{n})^{\text{nt}}\\
&={@principal@}\left (1+\displaystyle \frac{{@rate_float@}}{{@2@}}\right )^{2({@years@})}\\
&={@principal@}\left (1+{@rate_float/2@}\right )^{2({@years@})}\\
&={@principal@}({@sol1@})\\
&={@ta1@}\\
\end{align}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*compound interest question*/
principal:rand_with_step(6000,10000,500);
rate:rand(5)+5;
rate_float:float(rate/100);
years:rand(20)+20;
/*worked solution variables*/
sol1:(1+(rate_float/2))^(2*years);
sol2:principal*sol1;
round_dp_n(float_number, decimal_places) :=
  block(
    [rounded_number: floor(float_number * 10^decimal_places + 0.5)],
    float(rounded_number / 10^decimal_places)
  );
ta1:round_dp_n(sol2,2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Principal \(={@principal@}\)</p>
<p>Interest rate \(={@rate_float@}\)</p>
<p>Time \(={@years@}\)</p>
<p>\(A=P(1+\displaystyle \frac{r}{n})^{\text{nt}}={@ta1@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p><p>
</p><p id="concept">https://docs.google.com/document/d/1eFPMYtaSAq0t7Qg7M5bE3VcF1dzmHtQ5uyU7IOH4N_M/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer"></p>
<p id="reviewer">2</p>
<p id="description">In this question, students will explore how investments grow over time through compound interest, showcasing the effects of regular compounding on the final value of an account. </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check for the correct worth after 20yrs</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>sol2</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Monetary values are typically represented with two decimal places to reflect cents. Values with more than two decimal places are not standard and should be rounded to the nearest cent.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1420465409</deployedseed>
    <deployedseed>156946340</deployedseed>
    <deployedseed>311987134</deployedseed>
    <deployedseed>687528945</deployedseed>
    <deployedseed>1798639956</deployedseed>
    <deployedseed>744726744</deployedseed>
    <deployedseed>910088848</deployedseed>
    <deployedseed>1293694071</deployedseed>
    <deployedseed>2076021600</deployedseed>
    <deployedseed>1215999568</deployedseed>
    <deployedseed>1112999184</deployedseed>
    <deployedseed>139585827</deployedseed>
    <deployedseed>149384755</deployedseed>
    <deployedseed>1536282696</deployedseed>
    <deployedseed>1205806282</deployedseed>
    <deployedseed>1953118921</deployedseed>
    <deployedseed>180228234</deployedseed>
    <deployedseed>820292903</deployedseed>
    <deployedseed>41091336</deployedseed>
    <deployedseed>374054480</deployedseed>
    <deployedseed>1451166518</deployedseed>
    <deployedseed>1507778671</deployedseed>
    <deployedseed>623699117</deployedseed>
    <deployedseed>2132015631</deployedseed>
    <deployedseed>1447005067</deployedseed>
    <deployedseed>381669323</deployedseed>
    <deployedseed>811463894</deployedseed>
    <deployedseed>1244871211</deployedseed>
    <deployedseed>1194462388</deployedseed>
    <deployedseed>493088399</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
