<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Find the value of 𝑑𝑦𝑑𝑥 in log-base-10</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>If \(\displaystyle y={@fx@}\), find the value of \(\displaystyle \frac{dy}{dx}\) at \(\displaystyle x={@bb@}\). Give your answer in exact form in terms of ln. Recall that \(\displaystyle \ln(x)\) is entered as log(\(x\)).</p><p>Answer = [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>We will first find the derivative of \(\displaystyle={@ins@}\):</p><p>\(\displaystyle \frac{dy}{dx} = \frac{({@bot@})\cdot {@diff(top,x)@} - {@top@}\cdot {@diff(bot,x)@}}{({@bot@})^2} = \frac{{@expand(bot*diff(top,x)-top*diff(bot,x))@}}{({@bot@})^2}\).</p><p>Now suppose that \(\displaystyle y = {@logbase(x,10)@}\). Then \(10^y = x\). Differentiating both sides implicitly gives us \(\displaystyle 10^y \cdot \ln(10) \cdot y' = 1\). Therefore, \(\displaystyle x\cdot \ln(10) \cdot y' = 1\) and \(\displaystyle y' = \frac{1}{x\cdot \ln(10)}\).</p><p>Combining both results we get:</p><p>\(\displaystyle \frac{dy}{dx} = \frac{1}{{@ins@}\cdot \ln(10)} \cdot \frac{{@expand(bot*diff(top,x)-top*diff(bot,x))@}}{({@bot@})^2} = \frac{{@bot*expand(bot*diff(top,x)-top*diff(bot,x))@}}{{@top*log(10)@}\cdot ({@bot@})^2} = \frac{{@expand(bot*diff(top,x)-top*diff(bot,x))@}}{{@top*log(10)@}\cdot ({@bot@})}\).</p><p>Substituting \(\displaystyle x = {@bb@}\) we get \(\displaystyle \frac{dy}{dx} = {@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand([2,3]);
top:x^aa;
bot:1+x^2;
ins:top/bot;
fx:logbase(ins,10);
bb:(rand(8)+1)/10;
dfx:(expand(bot*diff(top,x)-top*diff(bot,x)))/(top*log(10)*bot);
ta:ev(dfx,x=bb);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\displaystyle y={@fx@}\)</p>
<p>\(\displaystyle \frac{dy}{dx}\) at \(\displaystyle x={@bb@}={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey  Wabwire</p>
<p id="reviewer"> Juma Zevick </p>
<p id="description"> Find the value of 𝑑𝑦𝑑𝑥 in log-base-10 implicitly</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1186651281</deployedseed>
    <deployedseed>159331083</deployedseed>
    <deployedseed>2047733058</deployedseed>
    <deployedseed>1556712099</deployedseed>
    <deployedseed>965541709</deployedseed>
    <deployedseed>1659883530</deployedseed>
    <deployedseed>1627234313</deployedseed>
    <deployedseed>1144291054</deployedseed>
    <deployedseed>1171109645</deployedseed>
    <deployedseed>91328034</deployedseed>
    <deployedseed>1079720745</deployedseed>
    <deployedseed>1747828852</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
