<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding Inverses of Functions_9(i)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find  the inverse of the function \( \displaystyle f(x)={@func@}\)</p>
<p><b>Note:</b> \(\ln\) can be typed as <code> log</code></p>
<p>\(f^{-1}(x)=\)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To find the inverse of the function \( \displaystyle f(x)={@func@}\), express the function in terms of \(x\) and  \(y\) and proceed as follows</p>
<p> \(\begin{align}
\displaystyle f(x)&={@func@}\\
&x={@const*%e^(y^power)@}\\
&\displaystyle \frac{x}{{@const@}}={@%e^(y^power)@}\\
&{@sol3@}\\
\end{align}\)</p>
<p>Solving for \(y\) by finding  root  \({@power@}\) on both sides.</p>
<p>\(x={@sol4@}\)</p>
<p></p>
<p></p>
<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*finding inverses*/
const:rand(6)+4;
power:rand(3)+6;
func:const*%e^(x^power);
/*worked solution variables*/
sol1:x=const*%e^(y^power);
sol2:x/const=%e^(y^power);
sol3:log(x/const)=(y^power);
sol4:(log(x/const))^(1/power);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(f(x)={@func@}\) </p>
<p>\(f^{-1}(x)=\displaystyle{@sol4@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p><p>
</p><p id="concept">https://docs.google.com/document/d/1SHRuavwMgYyW58e67mna9MuZFOHpXMe7aPVGGcDpRYo/edit</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer">Santiago Borio</p>
<p id="description">This question requires the  student to find the inverse of the  given  function. </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>sol4</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct inverse</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>sol4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for the correct inverse</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>-(sol4)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1765165637</deployedseed>
    <deployedseed>496522306</deployedseed>
    <deployedseed>360421082</deployedseed>
    <deployedseed>322874811</deployedseed>
    <deployedseed>1504895083</deployedseed>
    <deployedseed>199600339</deployedseed>
    <deployedseed>839507601</deployedseed>
    <deployedseed>1461833035</deployedseed>
    <deployedseed>529160054</deployedseed>
    <deployedseed>66490238</deployedseed>
    <deployedseed>1510955902</deployedseed>
    <deployedseed>491629558</deployedseed>
    <deployedseed>1379940533</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>sol4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
