<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Find the derivative of a log function implicitly</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Find the derivative of \(y={@fx@}\). Recall that \(\ln(ax)\) is entered as log(\(a\)*\(x\)) and \(\log_a(bx)\) is entered as logbase(\(b\)*\(x\),\(a\)).</p><p>\(y' = \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>\(y = {@fx@}\), therefore, \({@aa@}^y = {@ins@}\)</p><p>Differentiating implicitly on both sides we get \({@aa^y@}\cdot \ln {@aa@} \cdot y' = {@bb*cc*x^(cc-1)@}\) and \(y' = \frac{{@bb*cc*x^(cc-1)@}}{{@aa^y@}\cdot \ln {@aa@}}\).</p><p>Since&nbsp;\({@aa@}^y = {@ins@}\), \(y' = \frac{{@bb*cc*x^(cc-1)@}}{{@ins@}\cdot \ln {@aa@}} = \frac{{@cc@}}{x\cdot \ln {@aa@}}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(8)+2;
bb:rand(6)+2;
cc:rand(2)+2;
ins:bb*x^cc;
fx:logbase(ins,aa);
ta:cc/(x*log(aa));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \(y={@fx@}\)</p>
<p>\(y' = {@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">Find the derivative of a log function given by \(𝑦=loga(b⋅𝑥^2)\) implicitly</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1503842498</deployedseed>
    <deployedseed>492656918</deployedseed>
    <deployedseed>1335266770</deployedseed>
    <deployedseed>1873231554</deployedseed>
    <deployedseed>609774995</deployedseed>
    <deployedseed>690488901</deployedseed>
    <deployedseed>1229404004</deployedseed>
    <deployedseed>754021228</deployedseed>
    <deployedseed>1670817660</deployedseed>
    <deployedseed>2084702454</deployedseed>
    <deployedseed>1153366534</deployedseed>
    <deployedseed>751356615</deployedseed>
    <deployedseed>1711383103</deployedseed>
    <deployedseed>1979908360</deployedseed>
    <deployedseed>2017968711</deployedseed>
    <deployedseed>1654964415</deployedseed>
    <deployedseed>143708963</deployedseed>
    <deployedseed>1852100605</deployedseed>
    <deployedseed>671864124</deployedseed>
    <deployedseed>1636261541</deployedseed>
    <deployedseed>1106632582</deployedseed>
    <deployedseed>1983659845</deployedseed>
    <deployedseed>278691114</deployedseed>
    <deployedseed>1863659314</deployedseed>
    <deployedseed>1696642078</deployedseed>
    <deployedseed>289741362</deployedseed>
    <deployedseed>636453370</deployedseed>
    <deployedseed>327413135</deployedseed>
    <deployedseed>24921086</deployedseed>
    <deployedseed>1831664901</deployedseed>
    <deployedseed>267222739</deployedseed>
    <deployedseed>2054118400</deployedseed>
    <deployedseed>726071344</deployedseed>
    <deployedseed>1433686490</deployedseed>
    <deployedseed>1060138937</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
