<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Find the x-coordinate of the point on the graph of fx where the tangent is horizontal</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the \(x-\)coordinate of the point on the graph of \( x^4 + y^4 + {@aa@} = {@bb@}\cdot x\cdot y^3 \) where the tangent is horizontal. </p>
<p><em>Note: Give your answer correct to 2 decimal places.</em></p>
<p> [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>\( x^4 + y^4 + {@aa@} = {@bb@}\cdot x\cdot y^3 \).</p><p>Differentiating implicitly on both sides we get</p><p>\( {@4*x^3@} + {@4*y^3@}\cdot y' = {@bb@}\cdot y^3 + {@bb*3*x*y^2@} \cdot y'\)<br></p><p>\( {@4*y^3@}\cdot y' - {@bb*3*x*y^2@} \cdot y' = {@bb@}\cdot y^3 - {@4*x^3@} \)<br></p><p>\( y' \cdot( {@4*y^3@} - {@bb*3*x*y^2@}) = {@bb*y^3 - 4*x^3@} \)<br></p><p>\( y' = \frac{{@bb*y^3 - 4*x^3@}}{ {@4*y^3@} - {@bb*3*x*y^2@} } \)<br></p><p>The tangent is horizontal when \(y' = 0\), therefore,</p><p>\(&nbsp;\frac{{@bb*y^3 - 4*x^3@}}{ {@4*y^3@} - {@bb*3*x*y^2@} } = 0\)</p><p>\({@bb*y^3 - 4*x^3@} = 0\)</p><p>\({@bb*y^3@} = {@4*x^3@}\)</p><p>\(y^3 = {@4*x^3/bb@}\)</p><p>\(y = \sqrt[3]{{@4/bb@}}\cdot x\)</p><p>Substituting this back into the original equation gives us&nbsp;</p><p>\( x^4 + \big(\sqrt[3]{{@4/bb@}}\cdot x \big)^4 + {@aa@} = {@bb@}\cdot x\cdot \big(\sqrt[3]{{@4/bb@}}\cdot x\big)^3 \)<br></p><p>\( x^4 + \sqrt[3]{{@(4/bb)^4@}}\cdot x^4 + {@aa@} = {@bb@}\cdot {@(4/bb)@}\cdot x^4 \)</p><p>\( x^4 + \sqrt[3]{{@(4/bb)^4@}}\cdot x^4 + {@aa@} = 4 \cdot x^4 \)<br></p><p>\( {@aa@} = 4 \cdot x^4 - x^4 - \sqrt[3]{{@(4/bb)^4@}}\cdot x^4\)<br></p><p>\( {@aa@} = x^4\cdot (3 - \sqrt[3]{{@(4/bb)^4@}})\)<br></p><p>\( x^4 = \frac{{@aa@}}{3 - \sqrt[3]{{@(4/bb)^4@}}}\)<br></p><p>\( x = \sqrt[4]{\frac{{@aa@}}{3 - \sqrt[3]{{@(4/bb)^4@}}}} = {@dispdp(ta,2)@}\).<br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(8)+1;
bb:rand(5)+2;
ta:(aa/(3-(4/bb)^(4/3)))^0.25;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>The \(x-\)coordinate of the point on the graph of \( x^4 + y^4 + {@aa@} = {@bb@}\cdot x\cdot y^3 \) where the tangent is horizontal is \({@dispdp(ta,2)@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">Find the x-coordinate of the point on the graph of fx where the tangent is horizontal-implicit-differentiation</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>dispdp(ta,2)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>391580083</deployedseed>
    <deployedseed>196541622</deployedseed>
    <deployedseed>793675648</deployedseed>
    <deployedseed>56198438</deployedseed>
    <deployedseed>217744758</deployedseed>
    <deployedseed>980784314</deployedseed>
    <deployedseed>723408657</deployedseed>
    <deployedseed>2080488985</deployedseed>
    <deployedseed>1989172654</deployedseed>
    <deployedseed>1583907188</deployedseed>
    <deployedseed>497613547</deployedseed>
    <deployedseed>526665285</deployedseed>
    <deployedseed>38925645</deployedseed>
    <deployedseed>1426626526</deployedseed>
    <deployedseed>1988840026</deployedseed>
    <deployedseed>2143751897</deployedseed>
    <deployedseed>668436482</deployedseed>
    <deployedseed>1953165327</deployedseed>
    <deployedseed>1415933158</deployedseed>
    <deployedseed>1606702374</deployedseed>
    <deployedseed>1094567076</deployedseed>
    <deployedseed>294578421</deployedseed>
    <deployedseed>344242177</deployedseed>
    <deployedseed>1094161911</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ta,2)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
