<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Find the rate at which the population is growing when 𝑡=a</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>A population of \({@a_1@}\) bacteria is introduced into a culture and grows according to the equation&nbsp;</p><p>\( \displaystyle p(t) = {@fx@}\)</p>
<p>where \(t\) is measured in hours. Find the rate at which the population is growing when \(t={@dd@}\).</p><p>Give your correct answer to \(2\) decimal places.</p><p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To find the rate at which the population is growing when \(t = {@dd@} \), we need to calculate the derivative of the population function \(p(t)\) with respect to time \(t\) and then evaluate it at \(t = {@dd@}\).</p>
<p>The given population growth equation is:</p>
\[p(t) =  {@fx@} \]

<p>To find \(  \displaystyle \frac{dp}{dt} \), we will differentiate this equation using the quotient rule, which states that for a function \( \displaystyle u(t) =\frac{ f(t)}{g(t)} \), the derivative \(  \displaystyle \frac{du}{dt} \) is given by:</p>

\[  \displaystyle \frac{du}{dt} = \frac{f'(t)g(t) - f(t)g'(t)}{[g(t)]^2}\]

<p>In our case, \(  \displaystyle f(t) = {@numer@} \) and \(g(t) = {@denom@} \).</p>
<p>Now, let's calculate the derivatives:</p>
<p>\(1. f'(t) = {@diff((numer),t)@} \) (the derivative of \({@numer@}\) with respect to \(t\)).</p>
<p>\(2. g'(t) = {@diff(denom,t)@} \) (the derivative of \({@denom@}\) with respect to \(t\)).</p>
<p>Now, apply the quotient rule to \(p(t)\):</p>

\[  \displaystyle p'(t) = {@a_1@} \left( \frac{({@diff((numer),t)@})({@denom@}) - ({@numer@})({@diff(denom,t)@})}{({@denom@})^2} \right) \]

<p>Simplify this expression:</p>

\[  \displaystyle p'(t) = {@a_1@} \left( \frac{ {@diff((numer),t)*denom@} - {@numer*diff(denom,t)@}}{({@denom@})^2} \right) \]

\[p'(t) = {@dfx@} \]

<p>Now, we want to find \(p'({@dd@})\), which is the rate at which the population is growing when \(t = {@dd@} \):</p>
<p>calculating this expression:</p>

\[  \displaystyle p'({@dd@}) \approx {@ta@}\]

<p>So, the rate at which the population is growing when \(t = {@dd@}\) is approximately \({@decimalplaces(ta,2)@}\), rounded to two decimal places.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(10)+1;
bb:rand(60)+40;
cc:rand_with_prohib(-5,5,[0]);
dd:rand(5)+1;
numer:aa*t;
denom:bb+t^2;
a_1:rand_with_step(300,15000,100);
fx:a_1*(cc+numer/denom);
dfx:ratsimp(diff(fx,t));
ta:ev(dfx,t=dd);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>A population of \({@a_1@}\) bacteria is introduced into a culture and grows according to the equation</p>
<p>\(p(t) = {@fx@}\)</p>
<p>where \(t\) is measured in hours. The rate at which the population is growing when \(t={@dd@}\) is \({@dispdp(ta,2)@}\).</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">A population of N bacteria is introduced into a culture and grows according to the equation  𝑝(𝑡) Finding the rate at which the population is growing when 𝑡=a</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>dispdp(ta,2)</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Check corret answer with error margin</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.7</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>You got a partial mark, because your answer is correct with \( \pm 0.01 \) margin.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>422141329</deployedseed>
    <deployedseed>1518989551</deployedseed>
    <deployedseed>224276655</deployedseed>
    <deployedseed>1337908995</deployedseed>
    <deployedseed>402001072</deployedseed>
    <deployedseed>697940619</deployedseed>
    <deployedseed>1606144449</deployedseed>
    <deployedseed>151946140</deployedseed>
    <deployedseed>1261623684</deployedseed>
    <deployedseed>1156254434</deployedseed>
    <deployedseed>1647487278</deployedseed>
    <deployedseed>1894608257</deployedseed>
    <deployedseed>311305270</deployedseed>
    <deployedseed>1137985313</deployedseed>
    <deployedseed>1085597406</deployedseed>
    <deployedseed>1231169242</deployedseed>
    <deployedseed>1385303199</deployedseed>
    <deployedseed>663896924</deployedseed>
    <deployedseed>1492242899</deployedseed>
    <deployedseed>1108574018</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>dispdp(ta,2)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
