<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Find the Volume of the  Cylinder-q3</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Use the method of cylinders to determine the volume of the solid obtained by rotating the region bounded by \({@eqn_1@}\) and \({@eqn_2@}\) about the line \({@line@}\) which is given by the sketch below.</p>
<p>
[[jsxgraph]] 
 var board = JXG.JSXGraph.initBoard(divid, {
            boundingbox: [-12, 3, 24, -12],
            keepaspectratio: true,
            showCopyright: false,
            axis: true,
            showNavigation: true
        });
 var function1 = board.create('functiongraph', ['{#line_x#}',  {#line_x1#},{#line_x2#}], { name: '{#line#}', withlabel: true });
  var function3 = board.create('functiongraph', [function(y){return y*y-{#const_1#};}, -10, 10],{name:'y^2 - {#const_1#}', withLabel:true, strokeColor:'black'});
    var function2 = board.create('functiongraph', [function(y){return {#const_2-coeff_y*y#};}, -10, 10],{name:'{#const_2-coeff_y *y#}', withLabel:true, strokeColor:'black'});
make_filled_curve(board, [function1, function3, function2], { strokeWidth: 3, fillColor: 'yellow', fillOpacity: 0.3 });
 
       
[[/jsxgraph]]
</p>
<p>\(V= \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><h3>Worked solution</h3></p>
<p>Recall that, by using the method of cylinders, the volume of the solid generated by a region bounded by the lines \(\displaystyle y=a\) and \(\displaystyle y=b\), which is revolved about the axis \(\displaystyle y=c\) is given by</p>

<p>\(\displaystyle V=2 \pi \int_{a}^{b}(y-c)f(y) dy,\) and radius \(\displaystyle r=y-c\)</p>
<p>The values \(a\) and \(b\) are obtained from the intersection points as shown in the sketch. So we set the two equations \(x={@eqn_1@}\) and \(x={@eqn_2@}\) equal and solve it.</p>
<p> 
    \( \displaystyle
    \begin{align}
    {@eqn_1@}={@eqn_2@} &amp; \Longleftrightarrow {@eqn_3@}=0 \\
    &amp;\Longleftrightarrow {@solve@} \\
    \end{align}
    \)
</p>
<p>For this problem, \(\displaystyle a={@min@},b={@max@},r={@Radius@}\) and \(f(y)=({@eqn_1@})-({@eqn_2@})={@Width@}\).</p>
<p>Recall that the value of the volume must be positive. When \(V\) is negative, you pick its absolute value.</p>
<p>Therefore,</p>
<p>
    \(\displaystyle
    \begin{align}
    V &amp;=2 \pi \int_{{@min@}}^{{@max@}}({@Radius@}) \cdot ({@Width@})\,dy \\
    &amp;= 2 \pi \int_{{@min@}}^{{@max@}}({@Area_y_1@})\,dy \\
    &amp;=\left [ {@Vol_1@} \right ]_{{@min@}}^{{@max@}} \\
    &amp;={@ta@}
    \end{align}
    \)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>const_1:rand(6)+1;/*constant in the first equation*/
eqn_1:y^2-const_1;/*the first equation*/
const_2:rand(6)+1;/*constant in the second equation*/
coeff_y:rand(4)+1;/*coefficient of y*/
eqn_2:const_2-coeff_y*y;/*the second equation*/
eqn:eqn_1=eqn_2;/*equating the two equations*/
eqn_3: ratsimp(eqn_1-eqn_2);
eqn_4: factor(eqn_3);
solve: solve(eqn_3);/*solving for the equation*/
line_x1: 0;
line_x2: 21;
min_1:first(solve);/*the minimum value of the solution*/
min:rhs(min_1);/*the real part of the minimum value of the solution*/
max_1:second(solve);/*the maximum value of the solution*/
max:rhs(max_1);/*the real part of the minimum value of the solution*/
line_x:rand(8)-9;/*the value of the equation of the line in x-axis*/
line:y=line_x;/*equation of the line in x-axis*/
Radius:-line_x+y;/*the radius of the cylinder*/
Width: ratsimp(eqn_1-eqn_2);/*the width of the cylinder*/
Area_y:2*%pi*Radius*Width;/*the area of the cylinder*/
Area_y_1:expand(Radius*Width);
Area_y_2:2*%pi*Area_y_1;
Vol_1:int(Area_y_2,y);
Vol:int(Area_y_2,y,min,max);/*the volume of the cylinder*/
ta: abs(ratsimp(Vol));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\({@eqn_1@}\) and \({@eqn_2@}\) about the line \(y=-8\)</p>
<p>Volume \(={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo </p>
<p id="concept">N/A</p>
<p id="reviewer">Mary </p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">The question focuses on finding the volume of a cylinder using a respective formula of finding it's volume</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*possible mistakes*/
/*node 1 checks for correct answer*/
/*node 2 mistake:integrating the indefinite integral without the limits */
mistake_2:int(Area_y_2,y);
/*node 3 mistake: interchanging the limits of integration*/
mistake_3:int(Area_y_2,y,max,min);
/*node 4 mistake:Adding a constant of integration after evaluation*/
mistake_4:int(Area_y_2,y,min,max)+c;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>check for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Intergrating the indefinite integral without limits</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have integrated the indefinite integral without the limits</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>changing the limits of intergration</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have interchanged the limits of integration</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Adding a constant of intergration</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have added a constant of integration after evaluation</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1067098821</deployedseed>
    <deployedseed>938262168</deployedseed>
    <deployedseed>1271915337</deployedseed>
    <deployedseed>1395646486</deployedseed>
    <deployedseed>385997030</deployedseed>
    <deployedseed>1484639036</deployedseed>
    <deployedseed>1010370627</deployedseed>
    <deployedseed>1304590230</deployedseed>
    <deployedseed>1712198462</deployedseed>
    <deployedseed>1262877917</deployedseed>
    <deployedseed>1314401911</deployedseed>
    <deployedseed>422430984</deployedseed>
    <deployedseed>2125246361</deployedseed>
    <deployedseed>1205619303</deployedseed>
    <deployedseed>673541891</deployedseed>
    <deployedseed>935173918</deployedseed>
    <deployedseed>999309928</deployedseed>
    <deployedseed>1164098411</deployedseed>
    <deployedseed>596016428</deployedseed>
    <deployedseed>823820453</deployedseed>
    <deployedseed>1734713040</deployedseed>
    <deployedseed>1446902550</deployedseed>
    <deployedseed>802929779</deployedseed>
    <deployedseed>1825128137</deployedseed>
    <deployedseed>879189482</deployedseed>
    <deployedseed>1511232968</deployedseed>
    <deployedseed>1267719755</deployedseed>
    <deployedseed>655249752</deployedseed>
    <deployedseed>2063980067</deployedseed>
    <deployedseed>1992892787</deployedseed>
    <deployedseed>376501340</deployedseed>
    <deployedseed>1870002044</deployedseed>
    <deployedseed>1143407467</deployedseed>
    <deployedseed>1436829114</deployedseed>
    <deployedseed>1475148395</deployedseed>
    <deployedseed>1362880986</deployedseed>
    <deployedseed>2000908973</deployedseed>
    <deployedseed>1868378630</deployedseed>
    <deployedseed>18569457</deployedseed>
    <deployedseed>1384117936</deployedseed>
    <deployedseed>1437450732</deployedseed>
    <deployedseed>1208279046</deployedseed>
    <deployedseed>114512421</deployedseed>
    <deployedseed>2043074637</deployedseed>
    <deployedseed>1023315696</deployedseed>
    <deployedseed>1379842156</deployedseed>
    <deployedseed>984623027</deployedseed>
    <deployedseed>677415787</deployedseed>
    <deployedseed>1916496670</deployedseed>
    <deployedseed>884979142</deployedseed>
    <deployedseed>1338276108</deployedseed>
    <deployedseed>1453978617</deployedseed>
    <deployedseed>1028017525</deployedseed>
    <deployedseed>9365694</deployedseed>
    <deployedseed>39213509</deployedseed>
    <deployedseed>1421458032</deployedseed>
    <deployedseed>1729050503</deployedseed>
    <deployedseed>1590826955</deployedseed>
    <deployedseed>1308361566</deployedseed>
    <deployedseed>783392011</deployedseed>
    <deployedseed>878457405</deployedseed>
    <deployedseed>785213342</deployedseed>
    <deployedseed>578013408</deployedseed>
    <deployedseed>1507954054</deployedseed>
    <deployedseed>1147175508</deployedseed>
    <deployedseed>1225614406</deployedseed>
    <deployedseed>1836898537</deployedseed>
    <deployedseed>1680848186</deployedseed>
    <deployedseed>1022157550</deployedseed>
    <deployedseed>1985886137</deployedseed>
    <deployedseed>796093477</deployedseed>
    <deployedseed>690402150</deployedseed>
    <deployedseed>1020923499</deployedseed>
    <deployedseed>347406943</deployedseed>
    <deployedseed>1123577052</deployedseed>
    <deployedseed>750693641</deployedseed>
    <deployedseed>735149206</deployedseed>
    <deployedseed>1466528065</deployedseed>
    <deployedseed>1414733879</deployedseed>
    <deployedseed>1206696676</deployedseed>
    <deployedseed>1574935861</deployedseed>
    <deployedseed>639639758</deployedseed>
    <deployedseed>1090219315</deployedseed>
    <deployedseed>1733695912</deployedseed>
    <deployedseed>786042668</deployedseed>
    <deployedseed>920794038</deployedseed>
    <deployedseed>164725406</deployedseed>
    <deployedseed>318650442</deployedseed>
    <deployedseed>1965302590</deployedseed>
    <deployedseed>639264871</deployedseed>
    <deployedseed>867577647</deployedseed>
    <deployedseed>1861830161</deployedseed>
    <deployedseed>593585148</deployedseed>
    <deployedseed>726310924</deployedseed>
    <deployedseed>1267320910</deployedseed>
    <deployedseed>130269365</deployedseed>
    <deployedseed>1642038070</deployedseed>
    <deployedseed>1454344779</deployedseed>
    <deployedseed>698604725</deployedseed>
    <deployedseed>2146952087</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
