<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Implicit Differentiation: Finding the Equation of Tangent Line</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the equation of the tangent line to \(\displaystyle {@func_x@}+{@func_y@} ={@lhs_question_func@}\) at ({@x_point@},{@y_point@})</p>
<p>Follow the outlined steps below. Fill in the missing blanks:</p>

<p><b>Step 1:</b> Differentiate \(\displaystyle {@func_x@}+{@func_y@} ={@lhs_question_func@}\) implicitly with respect to \(x\). </p>
<p><b>Hint:</b> \(\displaystyle a \cdot x^2+ b \cdot y^4 \cdot{@tta@}=c\) is typed as \(a*x^4  +b*y^4*diff(y,x)=c\) </p>
<br>  [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]

<p><b>Step 2:</b> Find \(y'\) using the result obtained above (Make  \(\displaystyle y'\) the subject).</p>
<br> \(y'\)= [[input:ans1]] [[validation:ans1]][[feedback:prt1]] 

<p><b>Step 3:</b>  Evaluate the derivative \(y'\) at ({@x_point@},{@y_point@}) to get the slope \(m\) of the tangent line. </p>
<br> \(m=\)[[input:ans2]] [[validation:ans2]][[feedback:prt2]] </p>

<p><b>Step 4:</b> Write down the equation of the tangent line. </p> 
<br> [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>Worked solution</p>
<p><b>Step 1:</b> Differentiate \(\displaystyle {@func_x@}+{@func_y@} ={@lhs_question_func@}\) implicitly with respect to \(x\). </p>
<p>{@ta4@}</p>
<p><b>Step 2:</b> Find \(y'\) using the result obtained above (Make  \(\displaystyle y'\) the subject).</p>
<br> \(y'={@ta1@}\)
<p><b>Step 3:</b>  Evaluate the derivative \(y'\) at ({@x_point@},{@y_point@}) to get the slope \(m\) of the tangent line. </p>
<br> \(m={@ta2@}\) </p>
<p><b>Step 4:</b> Write down the equation of the tangent line.</p> 
<p>\(y-y_1=m(x-x_1) \implies y=m(x-x_1)+y_1  \)</p>
</p>From you substitute for \(x\) and \(y\) in the above to get<p>
<br><p>\( {@ta3@}\implies  {@expand(ta3)@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Question func*/
coef_x:2+rand_with_step(-1,3,1);
coef_y:2+rand_with_step(-1,3,1);

func_x:(coef_x)*x^(2+rand(3));
func_y:(coef_y)*y^(2+rand(3));

question_func: func_x+func_y;

lhs_question_func:2+rand(4);
x_point:2+rand(4);
y_point:-sqrt(2+rand(4));

/*Worked solution*/
diff_func1:diff(question_func,x);
diff_func2:diff(question_func,y);
depends(y,x);
tta:diff(y,x); 
ta4:  diff_func1 + diff_func2*tta=0; /*input ans 4*/

ta1:-diff_func1/diff_func2; /*input ans1: Gradient function*/
ta21:subst(x_point,x,ta1);
ta2:subst(y_point,y,ta21); /*input ans2: Gradient*/

ta3: y=ta2*(x- x_point)+y_point;  /*input ans3: Equation of the tangent line*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\( {@func_x@}+{@func_y@} ={@lhs_question_func@}\) at ({@x_point@},{@y_point@})</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="concept">N/Ap>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Sal Mercuri</p>
<p id="description">Compute \( \displaystyle \frac{dz}{dx} \) using the Chain Rule for the given functions \( z(x, y) \) and \( y(x) \), Utilize the Chain Rule formula: \[ \displaystyle \frac{dz}{dx} = \frac{\partial z}{\partial x} + \frac{\partial z}{\partial y} \frac{dy}{dx} \]</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1: Checking for the correct answer*/
/*node2: Making y the subject of the formula instead of y'*/
simp:false;
node2:solve(question_func=lhs_question_func,y)[1];
node2_prt1:rhs(node2);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checking for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Making y the subject of the formula instead of y'</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>node2_prt1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You just made \( y \) the subject of expression {@question_func@}, you need to make {@tta@} the subject of the expression. </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1: checking for correct answer*/
/*node2: interchanging the points during substitution*/
var1:subst(x_point,y,ta1);
mistake1a:subst(y_point,x,var1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checking for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>interchanging the points during substitution</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake1a</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you have interchanged the points during substitution, you need to substitute \(x\) with {@x_point@} and \(y\) with {@y_point@} in the equation to find \( y' \)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1:checking for correct answer*/
/*node2: using the wrong gradient in ans2 to compute the equation of the tangent*/
mistake1b:ans2*(x- x_point)-ans2;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checking for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>using the wrong gradient in ans2 to compute the equation of the tangent</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mistake1b</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you have correctly computed the equation of the tangent using the incorrect answer in the first input answer box above. However, this is not the model answer which was expected here.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1: Checking for correct answer*/

/*node2: Checking if the student misses out the dy/dx in their input*/
ans_node2:diff_func1 + diff_func2=0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checking for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checking if the student misses out the dy/dx</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ans_node2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Note that you need to include {@tta@} as part of the derivative of {@func_y@}. Your input would be marked correct if you had {@diff_func2*tta@} as part of your answer.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>428134434</deployedseed>
    <deployedseed>2004365868</deployedseed>
    <deployedseed>1785077093</deployedseed>
    <deployedseed>276788484</deployedseed>
    <deployedseed>1129039571</deployedseed>
    <deployedseed>429894957</deployedseed>
    <deployedseed>1839832863</deployedseed>
    <deployedseed>1391071399</deployedseed>
    <deployedseed>303631562</deployedseed>
    <deployedseed>654183920</deployedseed>
    <deployedseed>2102949324</deployedseed>
    <deployedseed>2042698252</deployedseed>
    <deployedseed>1354670044</deployedseed>
    <deployedseed>571682048</deployedseed>
    <deployedseed>1462604429</deployedseed>
    <deployedseed>401979455</deployedseed>
    <deployedseed>1407331173</deployedseed>
    <deployedseed>1779230900</deployedseed>
    <deployedseed>887228270</deployedseed>
    <deployedseed>1146091139</deployedseed>
    <deployedseed>457188627</deployedseed>
    <deployedseed>1180701787</deployedseed>
    <deployedseed>884118928</deployedseed>
    <deployedseed>781378713</deployedseed>
    <deployedseed>1379818398</deployedseed>
    <deployedseed>629976316</deployedseed>
    <deployedseed>314952428</deployedseed>
    <deployedseed>197910489</deployedseed>
    <deployedseed>1575276018</deployedseed>
    <deployedseed>2120316299</deployedseed>
    <deployedseed>228404122</deployedseed>
    <deployedseed>930032868</deployedseed>
    <deployedseed>589512072</deployedseed>
    <deployedseed>176073047</deployedseed>
    <deployedseed>1893302934</deployedseed>
    <deployedseed>1766093772</deployedseed>
    <deployedseed>1178671030</deployedseed>
    <deployedseed>1457696938</deployedseed>
    <deployedseed>427426116</deployedseed>
    <deployedseed>1936602426</deployedseed>
    <deployedseed>1556646426</deployedseed>
    <deployedseed>727405227</deployedseed>
    <deployedseed>914429450</deployedseed>
    <deployedseed>1170916232</deployedseed>
    <deployedseed>1963971777</deployedseed>
    <deployedseed>17972936</deployedseed>
    <deployedseed>1416190520</deployedseed>
    <deployedseed>1096068778</deployedseed>
    <deployedseed>1490470748</deployedseed>
    <deployedseed>1099279355</deployedseed>
    <deployedseed>1546206244</deployedseed>
    <deployedseed>603730142</deployedseed>
    <deployedseed>52893283</deployedseed>
    <deployedseed>716991076</deployedseed>
    <deployedseed>1948286297</deployedseed>
    <deployedseed>145577223</deployedseed>
    <deployedseed>597259486</deployedseed>
    <deployedseed>728476885</deployedseed>
    <deployedseed>1164834594</deployedseed>
    <deployedseed>1664584760</deployedseed>
    <deployedseed>915766158</deployedseed>
    <deployedseed>1802484666</deployedseed>
    <deployedseed>462753041</deployedseed>
    <deployedseed>2134486950</deployedseed>
    <deployedseed>672124999</deployedseed>
    <deployedseed>776887173</deployedseed>
    <deployedseed>1958327571</deployedseed>
    <deployedseed>1969701594</deployedseed>
    <deployedseed>935886531</deployedseed>
    <deployedseed>867065093</deployedseed>
    <deployedseed>841481981</deployedseed>
    <deployedseed>2008509364</deployedseed>
    <deployedseed>2000567874</deployedseed>
    <deployedseed>1302278375</deployedseed>
    <deployedseed>1370372600</deployedseed>
    <deployedseed>1814825129</deployedseed>
    <deployedseed>1738965653</deployedseed>
    <deployedseed>996023868</deployedseed>
    <deployedseed>1152901592</deployedseed>
    <deployedseed>357037870</deployedseed>
    <deployedseed>416941018</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
