<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Improper Integrals that involves determining if the Integral converges or diverges</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Evaluate the integral below and determine if it converges or diverges. 
</p>
<p>Note: \( \infty\) is entered as <b>inf</b> and \( -\infty\) is entered as <b>minf</b>.</p>
<p>\(\displaystyle\int_{{@l_1@}}^{{@l_2@}}{@expr_udv@}\,dx=\)[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>Therefore, the integral  [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b>Worked solution</b></p>
<p>First, we need to recall that we can’t do the integral as long as there is an infinity in one of the limits (It is an improper integral!). Therefore, we’ll need to replace the infinity first with a parameter \(t\) as follows,
</p>
<p>\[\displaystyle\int_{{@l_1@}}^{{@l_2@}}{@expr_udv@}\,dx =\lim_{t \to \infty } \displaystyle\int_{{@l_1@}}^{t}{@expr_udv@}\,dx\]</p>
<p> </p>
<p>We now need to use the integration by parts technique to evaluate the integral. The formula for integration by parts is given by \[\displaystyle \int udv=u\cdot v-\int vdu\] </p>
<p> We now need to carefully consider the integrand and make a choice for \(u\) and \(dv\) that will simplify the integral.</p>
<p>Let \(u={@expr_u@}\) and \(dv={@expr_dv@}\,dx\)</p>
<p>\(u={@expr_u@} \, \Rightarrow du={@expr_du@}\,dx\)</p>
<p>\(dv={@expr_dv@}\,dx \, \Leftarrow v={@expr_v@}\)</p>
<p>\[
  \begin{align}
  \displaystyle\int_{{@l_1@}}^{t}{@expr_udv@}\,dx &=\left[{@expr_uv@} \right]_{{@l_1@}}^{t}+ \int_{{@l_1@}}^{t} {@expr_vdu@}\,dx \\
  & ={@expr_inta@} -{@expr_intb@} \\
&= {@expr_int_t@}
  \end{align}
  \]</p>
<p>Recall the following rule of exponential limits:</p> 
<p>\[\lim_{X \to - \infty} X \cdot e^{X}=0 \;\;\text{ and }\;\; \lim_{X \to - \infty} e^{X}=0 \]</p>
<p>Applying this rule into the integral, we get:
</p>
<p>    \[    \begin{align}
  \displaystyle\int_{{@l_1@}}^{{@l_2@}}{@expr_udv@}\,dx &=\lim_{t \to \infty}{@expr_int_t@}\\
       &={@ta1@}\\
    \end{align}    \]</p>
<p>Therefore, the integral converges </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*Defining the question*/
const_1: rand(5)+1;
const_2: rand(9);
l_1:rand(5);
l_2:inf;
coef: rand(5)+2;
func:(const_1+const_2*x)*%e^(-coef*x);
	
/*worked out solution variables*/
expr_u:(const_1+const_2*x);
expr_du:diff(expr_u)/del(x);
expr_dv:%e^(-coef*x);
expr_v:int(expr_dv,x);
expr_udv: expr_u*expr_dv;
expr_uv: expr_u*expr_v;
expr_vdu: expr_du*expr_v;
expr_inta: ev(expr_uv, x=t)-ev(expr_uv, x=l_1), simp;
expr_intb: int(expr_vdu,x,l_1,t);
expr_int_t: expand(expr_inta-expr_intb), simp;

/*Model answers*/
ta: limit(expr_int_t, t, inf);
ta1: expand(ta);
ta2: random_permutation([["converges", true],["diverges", false],["converges and diverges", false]]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Evaluate the integral below and determine if it converges or diverges.</p>
<p>\(\displaystyle\int_{{@l_1@}}^{{@l_2@}}{@expr_udv@}\,dx={@ta1@}\)</p>
<p>Therefore, the integral {@mcq_correct(ta2)[1]@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Mike</p>
<p id="description">The question focuses on evaluating an improper Integral  by parts</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1:checks for correct answers*/

/*node 2:Calculating the integral as F(a)-F(b) instead of F(b)-F(a)*/
mistake_1:F(a)-F(b) ;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for the first mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You calculated the value of the integral in a wrong format \(F(a)-F(b)\),instead of \(F(b)-F(a)\)
Hint: check the worked out solution
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta2)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>520111421</deployedseed>
    <deployedseed>1730022432</deployedseed>
    <deployedseed>392992338</deployedseed>
    <deployedseed>2106026170</deployedseed>
    <deployedseed>1655802434</deployedseed>
    <deployedseed>1173700416</deployedseed>
    <deployedseed>1929933520</deployedseed>
    <deployedseed>1749024360</deployedseed>
    <deployedseed>1959743812</deployedseed>
    <deployedseed>1257970512</deployedseed>
    <deployedseed>864306730</deployedseed>
    <deployedseed>1628128377</deployedseed>
    <deployedseed>1209968803</deployedseed>
    <deployedseed>236173587</deployedseed>
    <deployedseed>1081772027</deployedseed>
    <deployedseed>128996711</deployedseed>
    <deployedseed>1504728500</deployedseed>
    <deployedseed>60521317</deployedseed>
    <deployedseed>356992224</deployedseed>
    <deployedseed>1473971307</deployedseed>
    <deployedseed>837978684</deployedseed>
    <deployedseed>1238410704</deployedseed>
    <deployedseed>1542058202</deployedseed>
    <deployedseed>2128827445</deployedseed>
    <deployedseed>1953803961</deployedseed>
    <deployedseed>1864475807</deployedseed>
    <deployedseed>615085227</deployedseed>
    <deployedseed>461197240</deployedseed>
    <deployedseed>81616282</deployedseed>
    <deployedseed>648157353</deployedseed>
    <deployedseed>754080688</deployedseed>
    <deployedseed>1117727109</deployedseed>
    <deployedseed>1802842760</deployedseed>
    <deployedseed>2015138390</deployedseed>
    <deployedseed>1921324635</deployedseed>
    <deployedseed>497940549</deployedseed>
    <deployedseed>689429745</deployedseed>
    <deployedseed>1070950475</deployedseed>
    <deployedseed>1822872537</deployedseed>
    <deployedseed>1060081553</deployedseed>
    <deployedseed>1747072416</deployedseed>
    <deployedseed>1576936051</deployedseed>
    <deployedseed>1380428879</deployedseed>
    <deployedseed>1026038802</deployedseed>
    <deployedseed>817266571</deployedseed>
    <deployedseed>2085823376</deployedseed>
    <deployedseed>272122292</deployedseed>
    <deployedseed>1309718354</deployedseed>
    <deployedseed>2066817853</deployedseed>
    <deployedseed>1403987388</deployedseed>
    <deployedseed>757967920</deployedseed>
    <deployedseed>1176219232</deployedseed>
    <deployedseed>593919790</deployedseed>
    <deployedseed>2025411806</deployedseed>
    <deployedseed>977055744</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta2))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
