<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Integration by substitution</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find  \( \displaystyle \int {@f_t@}\, dt\)  by use of substitution method.<br></p><p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><strong>Worked solution<br></strong></p>
<p>Let us solve \( \displaystyle \int {@f_t@}\, dt\) by substitution method.<br>
    Let \(u={@constant_denom_1@}-{@coeff_denom_2@}t\)<br>
    Differentiating \(u\) with respect to \(t\) we have;<br>
    \(\frac{du}{dt}=-{@coeff_denom_2@}\), making \(dt\) the subject we have that \(dt=\frac{du}{-{@coeff_denom_2@}}\)<br>
    Replacing the obtained expressions \(u\) and \(dt\) back to the given integral we have;<br>
    \( \displaystyle \int {@f_t@}\, dt = \displaystyle \int {@g_t@}\, du=&nbsp;{@tas@} \).</p>
<p>Then,</p>
<p>\( \displaystyle \int {@f_t@}\, dt={@tas@}\)</p>
<p>But \(u={@constant_denom_1@}-{@coeff_denom_2@}t\)</p>
<p>Therefore,</p>
<p>\( \displaystyle \int {@f_t@}\, dt\)={@ta_1@}, where \(c\) is the constant of integration. </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Defining the question*/
numerator:rand(9)+2;/*random number for the numerator*/
constant_denom_1:rand(7)+1;/*random number for the constant of the denominator*/
coeff_denom_2:rand(5)+1;/*random number for the coefficient of t in the denominator*/
power:rand(9)+2;/*random value of n outside the root sign*/
f_t:numerator/((constant_denom_1-coeff_denom_2*t)^(1/power));/*the function*/
g_t:((-1/coeff_denom_2)*numerator)/((u)^(1/power));/*the function depending of the variable*/


/*Model answers*/
ta_1:int(f_t,t)+c;

/*Worked solution*/
tas:int(g_t,u)+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\int{@f_t@} dt={@ta_1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1:checking for the correct answer*/

/*node 2: mistake: missing the constant of integration*/
mistake_2:int(f_t,t);

/*node 3: mistake: not replacing the value of u back*/
mistake_3:((1/(power+1))/(power+1))*(u^(power+1))+c;

/*node 4: mistake: differentiating instead of integrating*/
mistake_4:diff(f_t,t);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions>t</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have missed the constant of integration. This is an indefinite integral. Recall, \(\int x^n\, dx=\frac{x^{n+1}}{n+1}+c\), where \( c\) is the constant of integration.<br></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You provided your answer in terms of u. Note that u was just a temporal value.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You differentiated instead of integrating. To find antiderivatives you need to integrate.  Recall,\(\int x^n\, dx=\frac{x^{n+1}}{n+1}+c\), where \(\ c\) is the constant of integration</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2057035038</deployedseed>
    <deployedseed>1712146402</deployedseed>
    <deployedseed>493763224</deployedseed>
    <deployedseed>1599781844</deployedseed>
    <deployedseed>986515359</deployedseed>
    <deployedseed>41910611</deployedseed>
    <deployedseed>388877763</deployedseed>
    <deployedseed>1302113632</deployedseed>
    <deployedseed>529462164</deployedseed>
    <deployedseed>591075878</deployedseed>
    <deployedseed>1741044716</deployedseed>
    <deployedseed>239914896</deployedseed>
    <deployedseed>1127187233</deployedseed>
    <deployedseed>1697238373</deployedseed>
    <deployedseed>104788498</deployedseed>
    <deployedseed>376604735</deployedseed>
    <deployedseed>518322255</deployedseed>
    <deployedseed>625007405</deployedseed>
    <deployedseed>1721225555</deployedseed>
    <deployedseed>1110196712</deployedseed>
    <deployedseed>1734975620</deployedseed>
    <deployedseed>598899009</deployedseed>
    <deployedseed>201237738</deployedseed>
    <deployedseed>528056831</deployedseed>
    <deployedseed>1980006515</deployedseed>
    <deployedseed>1929909294</deployedseed>
    <deployedseed>1388437109</deployedseed>
    <deployedseed>600230998</deployedseed>
    <deployedseed>1621741121</deployedseed>
    <deployedseed>1536183539</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
