<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Integration by substitution</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the antiderivatives or evaluate the definite integral</p>
<p>\(\displaystyle\int{@func@}\,d{@tt@}\)</p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Worked solution</b></p>
<p>Let us solve \(\displaystyle \int {@func@}\, d{@tt@}\) by substitution method.</p>
<p>let \(u= \displaystyle {@us@} \)</p>
<p>Differentiating \(u\) with respect to \({@tt@}\) yields to \(\frac{du}{d{@tt@}}={@bb@}\). Then \(d{@tt@}=\frac{du}{@bb@}\)</p>
<p>Replacing \(u\) and \(d{@tt@}\) expressions into the original integral yields;</p>
<p>\[&nbsp; \qquad \displaystyle \int {@func@}\, d{@tt@} = \displaystyle \int {@func_u@}\, du \] </p>
<p>\[\iff \displaystyle \int {@func@}\, d{@tt@} = \displaystyle \int {@func_u@}\, du \] </p>
<p> \[\iff \displaystyle \int {@func@}\, d{@tt@} = {@tas@} \] </p>
<p></p>
<p>Substituting the expression corresponding to \(u\) into \({@tas@}\) yields to:
</p>
<p><span style="font-size: 0.9375rem;">\[ \displaystyle \int {@func@}\, d{@tt@} = {@ta_1@} \]</span><br></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Random generating and calculating variables*/
aa:rand(5)+1;
bb:rand_with_prohib(-8,8,[0,1]);
tt:rand([t,x,z]);
power:rand(8)+9;
func:(aa+bb*tt)^power;
func_u:(1/bb)*((u)^power);

/*substitution expressions*/
us: aa+bb*tt;
ss: int(u^power,du)
ss:(u^(power+1))/(power+1)+c;

ta_1:int(func,tt)+c;

/*Worked solution*/
tas:int(func_u,u)+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Find the antiderivatives or evaluate the definite integral</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1:checking for the correct answer*/
/*node 2: mistake: Differentiating instead of integrating*/
mistake_1: diff(func,tt);
/*node 3: mistake: Giving your answer in U form instead of substituting back the integral*/
mistake_2: (u^power+1)/(power+1);
/*node 4: mistake: Forgetting the negative sign*/
mistake_3:Forgetting the negative sign;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have differentiated instead of integrating.
Recall \(\displaystyle\intx^n\,d(tt)= \(\displaystyle\int^power+1/(power+ 1)+c \)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have given your answer in u form without substituting the expression value of u
Here is a hint:\((u^10)/(10)+c = (- (1-t))^10/(10)+c\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
