<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Integration by substitution</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Evaluate the integral by use of substitution method \(\displaystyle \int {@f_x@} dx\).</p>
<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p> Worked out solution.</p>
<p>Let us solve \(\displaystyle \int {@f_x@} dx\) by substitution method.</p>
<p>Let  \( u = {@u_1@}\),&nbsp; \(\frac{du}{dx}={@diff_1@}\),&nbsp;&nbsp;&nbsp;&nbsp; \(dx=\frac{du}{{@diff_1@}}\)</p>

<p>Substituting for \(u\) in the original integral  we have:</p>
 
<p>\(\displaystyle\int {@f_x@} dx=\int x^{{@pwr-1@}}(u)^{{@power@}} \cdot \frac{du}{{@diff_1@}}\)</p>

 
<p>\(=\frac{1}{{@pwr@}} \displaystyle \int u^{@power@} du\) <br></p><p>\(=\frac{1}{{@pwr@}} [{@int(u^(power),u)@}]+c \)<br></p><p>\(=\frac{1}{{@pwr@}} [{@int_u@}]+c \)<br></p><p>But, \( u = {@u_1@}\), therefore we need to replace it then simplify the answer as shown:<br></p><p>\( \frac{1}{{@pwr@}} [{@int_u@}]+c = {@ta_1@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Static question variables*/
coeff_contant:rand(9)+1; /* random number for the coefficient of x*/
power :rand(100)+40; /* random number for the power of x*/
pwr:rand(5)+1;/* random number for the power of x*/
u_1:x^pwr+coeff_contant; /*the chose value of u*/
f_x:x^(pwr-1)*(x^pwr +coeff_contant)^power; /* The function */

/*model answer*/
ta_1:int(f_x,x)+c;

/*Worked out solution question variables*/
diff_1:diff(u_1,x);
int_u:int(u^(power),u);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\displaystyle \int {@f_x@} dx\)= {@ta_1@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1: check for correct answer*/

/*node2: Differentiating instead of integrating*/
mistake_1:diff(f_x,x);

/*node3: Failing to add the constant of integration*/
mistake_2:int(f_x,x);

/*node4: Failing to replace the u-substitution back*/
mistake_3:(1/pwr)*int_u+c;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions>x</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<span id="docs-internal-guid-89e395bd-7fff-092e-c5cd-d8db0e217a26">You seem to have integrated correctly but failed to replace the u to its original value.</span>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have missed the constant of integration.
    Recall, \(\displaystyle\int f(x)dx=F(x)+c\) , where c is the constant of integration and f is the derivative of F.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>mistake_3</tans>
        <testoptions>u</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<span id="docs-internal-guid-89e395bd-7fff-092e-c5cd-d8db0e217a26">You seem to have integrated correctly but failed to replace the u to its original value.</span>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
