<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Integration by substitution of definite Integral</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Evaluate by the use of substitution method:
<br> \(\displaystyle\int_{{@lower@}}^{{@upper@}} {@f_x@} dx =\) [[input:ans1]] [[validation:ans1]]&nbsp;[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><h3>Worked solution</h3></p>
<p>Find \(\displaystyle\int_{{@lower@}}^{{@upper@}} {@f_x@} dx\) by use of substitution method.</p>
<p>We recall that the antiderivative of \(\displaystyle u^{n}\) is \(\displaystyle \frac{u^{n+1}}{n+1}\) where \(\displaystyle n \neq -1\).</p>
<p>Let \(u= {@us@}\). Then \(du={@diff_u@} \, dx \). So \( \displaystyle dx = \frac{du}{{@diff_u@}} \).</p>
<p> When \( x \rightarrow {@lower@}, u \rightarrow  {@lower_u@} \).</p>
<p> When \( x \rightarrow {@upper@}, u \rightarrow  {@upper_u@} \).</p>
 <p>By substituting the values of \(x\) by their corresponding values of \(u\), and also \( dx\) by \( \frac{du}{{@diff_u@}}   \) in the original integral, we get;</p>
<p>\(
\begin{align}
\displaystyle \int_{{@lower@}}^{{@upper@}} {@f_x@} dx & = \int_{{@lower_u@}}^{{@upper_u@}} \frac{@constant@}{{@diff_u*u^power@}} \, du \\
& = \left[ {@expr2@} \right]_{{@lower_u@}}^{{@upper_u@}} \\
& = {@ta_1@} 
\end{align}
\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Defining the question*/
lower:rand(9)+1;/*random number for the lower limit*/
upper:(lower +1);/*random number for the upper limit*/
coeff_x:rand_with_prohib(-8,8,[0]);/*random number for the coefficient of x in the numerator*/
constant:rand(9)+1;/*random number for the constant of the denominator*/
power:rand(2)+2;/*random value of n outside the root sign*/
f_x:ev((constant/(coeff_x*x - constant)^(power)), simp);/*the function*/
part_frac_f_x:partfrac(f_x,x);
num_part:num(part_frac_f_x);
den_part:denom(part_frac_f_x);
func_1:coeff_x*x -constant;
fx_1:factor(func_1);



/*Defining the question*/
lower:rand(9)+1;/*random number for the lower limit*/
upper:(lower +1);/*random number for the upper limit*/
coeff_x:rand_with_prohib(-8,8,[0]);/*random number for the coefficient of x in the numerator*/
constant:rand(9)+1;/*random number for the constant of the denominator*/
power:rand(2)+2;/*random value of n outside the root sign*/
f_x:ev((constant/(coeff_x*x - constant)^(power)), simp);/*the function*/
part_frac_f_x:partfrac(f_x,x);
num_part:num(part_frac_f_x);
den_part:denom(part_frac_f_x);
func_1:coeff_x*x -constant;
fx_1:factor(func_1);






/*Substitution and other tasks*/
us:ev((coeff_x*x - constant), simp);

diff_u:diff(us,x);
func_2:diff_u*(u^(-power+1)/(-power+1))
subst_u:subst(us,u,func_2);
den_subst:denom(subst_u);

u_lower: ev(us, x=lower);
u_upper:ev(us, x=upper);
g_u:ev((1/diff_u)*(constant/(u^power)), simp);/*the function*/

/*Evaluating lower and upper bounds as follow*/
lower_u: at(us,[x=lower]);
upper_u: at(us,[x=upper]);

/*Model answers*/
ta_g:integrate(g_u,u);
upper_value:subst(u_upper,x,ta_g);
lower_value:subst(u_lower,x,ta_g);
eval_ta_g: ev(ta_g,u=us);
ta_1:integrate(f_x,x,lower,upper);

/*Worked solution*/
eval_expr1: constant/(diff_u*u^power);
expr2: integrate(eval_expr1,u);
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\displaystyle\int_{{@lower@}}^{{@upper@}} {@f_x@} dx={@ta_1@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey  Wabwire</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">The question focuses on evaluating a definite Integral by u substitution</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1:checking correct answer*/

/*node2:Interchanging the limits of integration*/
mistake_1:int(f_x,x,upper,lower);

/*node3:Differentiating instead of integrating*/
mistake_2:diff(f_x,x);

/*node4:Failing to replace the u substitution back*/
/*mistake_3:int(1/constant *u^{power+1},u,lower,upper);*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checking correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Interchanging the limits of integration</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have interchanged the limits of integration. Note that given \( \int_a^b{f(x)}dx = F(b) - F(a)\)0</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Differentiating instead of integrating</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have differentiated the function instead of integrating.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Failing to replace the u substitution back</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have integrated correctly but failed to substitute u in the final answer place.
Recall, if you let u=z  then the value of u needs to be replaced with the original value that is z</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>446693261</deployedseed>
    <deployedseed>1273218880</deployedseed>
    <deployedseed>1528538617</deployedseed>
    <deployedseed>1871835231</deployedseed>
    <deployedseed>1248633040</deployedseed>
    <deployedseed>467263349</deployedseed>
    <deployedseed>1512878727</deployedseed>
    <deployedseed>1073121424</deployedseed>
    <deployedseed>2111495287</deployedseed>
    <deployedseed>1041881011</deployedseed>
    <deployedseed>964693181</deployedseed>
    <deployedseed>1331216042</deployedseed>
    <deployedseed>1451711669</deployedseed>
    <deployedseed>976958022</deployedseed>
    <deployedseed>939367090</deployedseed>
    <deployedseed>1248509344</deployedseed>
    <deployedseed>436024234</deployedseed>
    <deployedseed>346228962</deployedseed>
    <deployedseed>612147426</deployedseed>
    <deployedseed>1226710943</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
