<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Inverse Function Computation</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>Find the inverse of \(R(x)={@fx@}\).</p><p>\(R^{-1}(x)=\) [[input:ans1]] [[validation:ans1]]<br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2022071300</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand(3)+2;
pm1:rand_with_prohib(-1,1,[0]);
bb:pm1*(rand(15)+1);
cc:rand(3)+2;
pm2:rand_with_prohib(-1,1,[0]);
dd:pm2*(rand(15)+1);
fx:(aa*x+bb)/(cc*x+dd);
ta:(bb-dd*x)/(cc*x-aa);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>aa,pm1,bb</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p><p>When calculating the inverse of a function \(R(x)\) we are finding a new&nbsp;function that will tell us which \(x\) value will give as a particular&nbsp;\(y\) value when we input it in \(R\). Therefore, in the case of \(R(x)={@(aa*x+bb)/(cc*x+dd)@}\), we write it as \(y={@(aa*x+bb)/(cc*x+dd)@}\) and invert the variables \(x\) and \(y\) to get \(x={@(aa*y+bb)/(cc*y+dd)@}\) and rearrange as follows:</p><p>\[x={@(aa*y+bb)/(cc*y+dd)@}\]<br></p><p></p><p>\[x\cdot ({@cc*y+dd@})={@aa*y+bb@}\]<br></p><p></p><p>\[{@cc*y*x+dd*x@}={@aa*y+bb@}\]<br></p>\[{@cc*y*x-aa*y@}={@bb-dd*x@}\]<br>\[{@y*(cc*x-aa)@}={@bb-dd*x@}\]<br>\[{@y@}={@(bb-dd*x)/(cc*x-aa)@}\]<br>The inverse of \(R(x)\) is therefore \(R^{-1}(x) = {@(bb-dd*x)/(cc*x-aa)@}\).<br><p></p><p></p><p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
