<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Limit of a Function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find \(\displaystyle \lim_{x \rightarrow {@aa@}} {@fx@}\). If the limit doesn't exist enter <code>NO</code> as your answer.</p><p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b><h3>Worked Solution</h3></b></p>
<p>Given the expression:</p>
<p style="margin-left: 30px">\(\displaystyle \lim_{x \rightarrow {@aa@}} {@fx@}\).</p>
<p>To find the limit of the above expression, we first need to simplify it;</p>

<p style="margin-left: 30px">\(\displaystyle \lim_{x \rightarrow {@aa@}} {@fx@} = \lim_{{x \rightarrow {@aa@}}} ({@x - aa@})\). <p>

<p>Now, we can directly evaluate the limit by plugging in the value of \(x\), which is \({@aa@}\).</p>

<p style="margin-left: 30px">\(\lim_{{x \rightarrow {@aa@}}} ({@x - aa@}) = {@aa@}{@sign_ev@}{@ev(abs(aa),simp)@} = 0\) </p>

<p>Therefore,</p>

<p style="margin-left: 30px">\(\displaystyle \lim_{{x \rightarrow {@aa@}}} \frac{{({@x - aa@})^2}}{{{@x - aa@}}} = 0\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[pm1:rand_with_prohib(-1,1,[0]);
aa:pm1*(rand(9)+1);
simp:false
fx:(x-aa)^2/(x-aa);
fsimp:ev(fx,simp);
tad:ev(fsimp,x=aa);
ta:ev(tad,simp);
sign_ev:if aa>0 then "-" else "+";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \(\displaystyle \lim_{x \rightarrow {@aa@}} {@fx@}={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="concept">N/A</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description"> This STACK question asks students to find the limit of a given function as \(x\) approaches a specified value. They need to compute the limit and indicate whether it exists or not.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You could simplify \(f(x)\) as follows \(f(x) = {@fx@} = {@fsimp@} \)</p>
<p> Therefore \(\displaystyle \lim_{x\rightarrow {@aa@}} f(x) = \lim_{x\rightarrow {@aa@}} ({@fsimp@})=f({@aa@})={@aa-aa@}={@ta@}\)</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>185883987</deployedseed>
    <deployedseed>72013422</deployedseed>
    <deployedseed>10371126</deployedseed>
    <deployedseed>53721626</deployedseed>
    <deployedseed>938933859</deployedseed>
    <deployedseed>1625466034</deployedseed>
    <deployedseed>1861125925</deployedseed>
    <deployedseed>437554440</deployedseed>
    <deployedseed>5211273</deployedseed>
    <deployedseed>1461854746</deployedseed>
    <deployedseed>1375007853</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
