<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Limit Evaluation of a Function Approaching the Origin Along Different Paths</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>  Fill in the missing computational steps. </p>
<p>Evaluate \(\lim_{(x,y) \to (0,0)} {@question_func@}\)</p>
<p>The first step will be to plug in \(y=0\) into the function and then take the limit as \( x\) approaches zero. 
</p>
<p>\(\lim_{(x,0) \to (0,0)} {@question_func@} = \)  [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<br>
<p> And along the \( y\)-axis we have \(x=0\) and so the limit is,</p>
<p>\(\lim_{(0,y) \to (0,0)} {@question_func@} = \)  [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<br>
<p> At this point, let’s take a look at a third fairly common path. To do this we’ll replace all the \(y’s \) with \(x’s\) and then let \( x\) approach zero.</p>
<p>\(\lim_{(x,y) \to (0,0)} {@question_func@} = \lim_{(x,x) \to (0,0)} \frac{({@x^aa@}){@x^aa@}}{{@x^(2*aa)@}+{@bb*x^(2*aa)@}} = \)  [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>
<br>
<p>Therefore we can say that the limit of this function [[input:ans4]] [[validation:ans4]][[feedback:prt4]] </p>
<p></p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>Worked solution</p>
<p>If we are going to approach \( (0,0) \) along the \(x\)-axis we can take advantage of the fact that along the \( x \)-axis we know that \( y=0\). This means that, along the \( x\) -axis, we will plug in \( y=0\) into the function and then take the limit as \( x \) approaches zero as follows;</p>
<br>
<p>\(\lim_{(x,0) \to (0,0)} {@question_func@} = \lim_{(x,0) \to (0,0)} \frac{({@x^aa@}){@0^aa@}}{{@x^(2*aa)@}+{@bb*0^(2*aa)@}} = {@ta1@}\)</p>
<br>
<p>So, along the \( x\)-axis the function will approach zero as we move in toward the origin. Now, let’s try the \(y\)-axis. Along this axis we have \(x=0\) and so the limit becomes,</p>
<p>\(\lim_{(x,y) \to (0,0)} {@question_func@} = \lim_{(0,y) \to (0,0)} \frac{({@0^aa@}){@y^aa@}}{{@0^(2*aa)@}+{@bb*y^(2*aa)@}} = {@ta2@}\)</p>
<br>
<p>This does NOT say that the limit exists and has a value of zero. This only means that the limit happens to have the same value along two paths. Let’s take a look at a third fairly common path. To do this we will replace all the \(y’s\) with \(x’s\) and then let \( x\) approach zero.</p>

<p>\(\lim_{(x,y) \to (0,0)} {@question_func@} = \lim_{(x,x) \to (0,0)} \frac{({@x^aa@}){@x^aa@}}{{@x^(2*aa)@}+{@bb*x^(2*aa)@}} = {@ta3@}\)</p>

<p>So, a different value from the previous two paths and this means that the limit can’t possibly exist.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[aa:rand(2)+1;
a1:aa^2;
bb:rand(5)+1;
num:x^aa*y^aa;
denom:x^(2*aa)+bb*y^(2*aa);
question_func:num/denom;

/*worked solution*/
ta1_a:subst(0,y,question_func); ta1:limit (ta1_a, y,0);
ta2_a:subst(0,x,question_func);  ta2:limit (ta2_a, x,0);
ta3_a:subst(x,y,question_func);  ta3:limit (ta3_a, x,0);
ta4:[["Exists", true],["Doesn't exist", false]];]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>({@var_x@},{@var_y@}) {@question_func@} = {@ta@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>ta4</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>If we plug in the values \(x={@var_x@}\) and \(y={@var_y@}\), you’ll realize that we’ll have \( \frac{0}{0}\), to show that this limit doesn't exist. Notice that we can factor both the numerator and denominator of the function;</p>
<p>\({@question_func@}=\frac{({@(cc*x+y)@})({@(aa*x-bb*y)@})}{({@(aa*x-bb*y)@})({@(aa*x+bb*y)@})}={@(cc*x+y)/(aa*x+bb*y)@}\)<p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mcq_correct(ta4)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
