<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Limits of functions with several variables- Example 3 (Direct Question)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Assuming the limit of the function below exists, compute its value. </p>
<p> \( \displaystyle \lim_{(x,y) \to ({@var_x@},{@var_y@})} {@question_func@} = \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p></p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>If we plug in the values \(\displaystyle x={@var_x@}\) and \(y={@var_y@}\), you’ll realize that we’ll have \(\displaystyle \frac{0}{0}\), and this might not be the case.</p>
<p> Notice that we can factor both the numerator and denominator of the function as follows,</p>
<p>\(\displaystyle
\begin{align} \lim_{(x,y) \to ({@var_x@},{@var_y@})} {@question_func@} &= \lim_{(x,y) \to ({@var_x@},{@var_y@})} \frac{({@(cc*x+y)@})({@(aa*x-bb*y)@})}{({@(aa*x-bb*y)@})({@(aa*x+bb*y)@})}\\
& = \lim_{(x,y) \to ({@var_x@},{@var_y@})} {@(cc*x+y)/(aa*x+bb*y)@} 
\end{align}\)</p>
<p>From here we can evaluate the limit as,</p>
<p>\( \qquad \qquad \displaystyle \lim_{(x,y) \to ({@var_x@},{@var_y@})} {@(cc*x+y)/(aa*x+bb*y)@} = {@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand_with_step(1,4,1);
bb:rand_with_step(1,4,1);
cc: rand_with_step(1,4,1);
var_x:bb;
var_y:aa;
num_func:expand((cc*x+y)*(aa*x-bb*y));
denom_func:expand((aa*x-bb*y)*(aa*x+bb*y));
question_func:num_func/denom_func;

question_func1:ratsimp(question_func);
lim_x:limit (question_func1, x, var_x); /*substituting for x*/
lim_y:limit (lim_x, y, var_y); /*substituting for y*/
ta:lim_y;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\[\displaystyle \lim_{(x,y) \to ({@var_x@},{@var_y@})} {@question_func@} = \lim_{(x,y) \to ({@var_x@},{@var_y@})} {@(cc*x+y)/(aa*x+bb*y)@} ={@ta@} \]</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="concept">https://docs.google.com/document/d/18XomBwrJn_JflBGU6drEJ7lj0ixz1FaQcKzYeeAdCoU/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer"></p>
<p id="description">The problem involves finding \(\lim_{{(x, y) \to (a, b)}}\) for the expression \(\frac{px+qy}{px^2-qy^2}\). Substituting \(x = a\) and \(y = b\) results in an indeterminate form (\(0/0\)), requiring further algebraic simplification or advanced techniques for evaluation, and method of cancellation.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for correct answer as a float</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>float(ta)</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks for correct answer to 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p style="color:blue">Well done for giving your answer correct to \(2\) decimal places.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>732049907</deployedseed>
    <deployedseed>852028347</deployedseed>
    <deployedseed>2120048754</deployedseed>
    <deployedseed>1639542001</deployedseed>
    <deployedseed>888872197</deployedseed>
    <deployedseed>1002320026</deployedseed>
    <deployedseed>1912610537</deployedseed>
    <deployedseed>2001248442</deployedseed>
    <deployedseed>1617519709</deployedseed>
    <deployedseed>486622751</deployedseed>
    <deployedseed>1169964031</deployedseed>
    <deployedseed>376713261</deployedseed>
    <deployedseed>1910383335</deployedseed>
    <deployedseed>1166231257</deployedseed>
    <deployedseed>374275090</deployedseed>
    <deployedseed>1196533641</deployedseed>
    <deployedseed>1617076286</deployedseed>
    <deployedseed>1753045173</deployedseed>
    <deployedseed>424416437</deployedseed>
    <deployedseed>1871455819</deployedseed>
    <deployedseed>1456917137</deployedseed>
    <deployedseed>1085157800</deployedseed>
    <deployedseed>260470749</deployedseed>
    <deployedseed>1868102635</deployedseed>
    <deployedseed>2014642889</deployedseed>
    <deployedseed>237166595</deployedseed>
    <deployedseed>2125612009</deployedseed>
    <deployedseed>1755895682</deployedseed>
    <deployedseed>1976013330</deployedseed>
    <deployedseed>1929414133</deployedseed>
    <deployedseed>585085864</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
