<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Limit of a Function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find \(\displaystyle \lim_{x \rightarrow {@aa@}} {@fx@}\)</p>
<p><em>Note: If no limit exists, give <code>NO</code> as your answer.</em></p>
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>First, we can simplify the function to \(\displaystyle f(x) = {@fx@}\)</p>

<p>We should consider both sides of the limit to determine the directional limits of the function.
<br>To do this, we could consider values on either side of the limit, however it is to easier to simplify the function and observe the limit from this.</p>

<br>

<p>Taking the limit from the <b>left</b>, \(x \rightarrow {@aa@}^{-}\), we have that \(x < {@aa@}\), therefore,</p>
<p style="margin-left: 30px">\({@abs(x-aa)@} = -({@x-aa@})\)</p>
<p style="margin-left: 30px">\(\displaystyle \Rightarrow {@fx@} = \frac{-({@x-aa@})}{{@x-aa@}} = -1\)</p>
<p>Hence, since the limit of a constant function is the constant itself,</p>
<p style="margin-left: 30px">\(\displaystyle \lim_{x \rightarrow {@aa@}^{-}} {@fx@} = -1\)</p>

<br>

<p>Now, taking the limit from the <b>right</b>, \(x \rightarrow {@aa@}^{+}\), we have that \(x > {@aa@}\), therefore,</p>
<p style="margin-left: 30px">\({@abs(x-aa)@} = {@x-aa@}\)</p>
<p style="margin-left: 30px">\(\displaystyle \Rightarrow {@fx@} = \frac{{@x-aa@}}{{@x-aa@}} = 1\)</p>
<p>Hence, as argued above,</p>
<p style="margin-left: 30px">\(\displaystyle \lim_{x \rightarrow {@aa@}^{+}} {@fx@} = 1\)</p>

<br>

<p>We conclude that, because the directional limits of the function are <b>not equal</b>, the limit of the function <b>does not exist</b>.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[pm1: rand_with_prohib(-1, 1, [0]); 
aa: pm1 * (rand(9) + 1); 
fx: abs(x - aa) / (x - aa); 
ta: "NO";  

]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\displaystyle \lim_{x \rightarrow {@aa@}} {@fx@}=\) {@ta@} limit.</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">NA</p>
<p id="reviewer">Nixson Kiplagat</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">This STACK question involves finding the limit of a given function as \(x\) approaches a specified value. Students are asked to compute the limit or determine if it does not exist./p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>string</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>StringSloppy</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>439408515</deployedseed>
    <deployedseed>695110142</deployedseed>
    <deployedseed>1073873018</deployedseed>
    <deployedseed>361347526</deployedseed>
    <deployedseed>657733006</deployedseed>
    <deployedseed>216024613</deployedseed>
    <deployedseed>702377097</deployedseed>
    <deployedseed>947226480</deployedseed>
    <deployedseed>1724320549</deployedseed>
    <deployedseed>1042212251</deployedseed>
    <deployedseed>177158422</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
