<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Nepero Eulero limits</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p> Calcola il limite della seguente funzione: </p>
\[ \lim_{x \mapsto \infty}\left(1 + \frac{{@fx@}}{{@gx@}}\right)^{{@hx@}} \]
<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text>Soluzioni:

Semplifica numeratore e denominatore fino ad ottenere:

\[ \lim_{x \mapsto \infty}\left(1 + \frac{1}{{@gg1@}x+{@gg2@}}\right)^{{@hx@}} \]

Poni \( y(x) := {@gg1@}x+{@gg2@} \) e verifica che  \( y(x) \mapsto + \infty \) se  \( x \mapsto \infty \). Osserva che:

\[ \left(1 + \frac{1}{{@gg1@}x+{@gg2@}}\right)^{{@hx@}} =\left(\left(1 + \frac{1}{y}\right)^y\right)^{({@hx@})/y} \]

Utilizza:

\[ \lim_{y \mapsto \infty}\left(1 + \frac{1}{y}\right)^{y} = e \],

dove \(e \) è il numero di Nepero, e che è possibile scambiare l' ordine di limite e esponenziale. Inoltre abbiamo che:

\[ \lim_{x \mapsto \infty} \frac{{@hx@}}{{@gg1@}x + {@gg2@}} = \frac{{@hh1@}}{{@gg1@}} \],

Quindi abbiamo:

\[ \lim_{x \mapsto \infty}\left(1 + \frac{{@fx@}}{{@gx@}}\right)^{{@hx@}} = \lim_{x \mapsto \infty}\left(1 + \frac{1}{{@gg1@}x+{@gg2@}}\right)^{{@hx@}} = e^{{@hh1@}/{@gg1@}}  \]</text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>ff1:rand(5)+2;
ff2:rand(4)+1;
gg1:rand(3)+3;
gg2:rand(3)+2;
hh1:rand(6)+2;
hh2:rand(3)+1;
fx:ff1*x+ff2;
gx:fx*(gg1*x+gg2)
hx:hh1*x+hh2
ta1:e^{hh1/gg1};</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>ta1</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
