<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Proving Non-Existence of a Limit</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find \(\displaystyle \lim_{x \rightarrow {@aa@}} {@fx@}\). If the limit doesn't exist enter <code>NO</code> as your answer.</p>
<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Remember that we need to test values from both sides in this case. You will see that from the left, values tend towards \(-1\) and from the right they tend to \(1\), and therefore there is no limit.<br></p>
<p>We can also look at the simplification of the function. In order to do so we need to consider the following two cases:</p>
<ol>
    <li>Values to the left of {@aa@}. In this case, the value of {@x-aa@} is negative and the function becomes \(\displaystyle f(x) = \frac{{@x-aa@}}{-({@x-aa@})} = -1\) so we can see that from the left the function tends to \(-1\).</li>
    <li>Values to the right of {@aa@}. In this case, the value of {@x-aa@} is positive and the function becomes \(\displaystyle f(x) = \frac{{@x-aa@}}{{@x-aa@}} = 1\) so we can see that from the right the function tends to \(1\).</li>
</ol>
<p>Therefore, the <strong>limit does not exist</strong>.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[pm1:rand_with_prohib(-1,1,[0]);
aa:pm1*(rand(9)+1);
fx:(x-aa)/abs(x-aa);
ta:"NO";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<br>Limit point \({@pm1@},{@aa@}\)
<br> \(\displaystyle f(x)=\lim_{x \rightarrow {@aa@}} {@fx@}\)]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="concept">https://docs.google.com/document/d/1Wfjz_C2qdEwFf_Ga4CHZGvkepwz7vncbskoV6gJtKAQ/edit#heading=h.vy8x1tbmoha9</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">2</p>
<p id="description">This question requires students to demonstrate through mathematical reasoning and analysis why the limit of the absolute value function divided by x does not converge to a single value as x approaches the given point.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>string</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>StringSloppy</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1657614897</deployedseed>
    <deployedseed>1345379203</deployedseed>
    <deployedseed>1751365047</deployedseed>
    <deployedseed>285234364</deployedseed>
    <deployedseed>2116871532</deployedseed>
    <deployedseed>2055225461</deployedseed>
    <deployedseed>782446712</deployedseed>
    <deployedseed>1848083067</deployedseed>
    <deployedseed>1389673367</deployedseed>
    <deployedseed>561969250</deployedseed>
    <deployedseed>1976970049</deployedseed>
    <deployedseed>1180723711</deployedseed>
    <deployedseed>706688105</deployedseed>
    <deployedseed>1409460369</deployedseed>
    <deployedseed>365895305</deployedseed>
    <deployedseed>1627725703</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
