<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Question 5 - Area Between Curves [Broken wrong extremes of integration]</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Determine the area of the region bounded by \({@expr1@}, \, {@expr2@}, \, y={@y_1@}\) and \(y={@y_2@}\)</p>
<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Worked Solution</p>
<p>Here is a sketch of the bounded region we want to find the area.</p>

[[jsxgraph]]

const board = JXG.JSXGraph.initBoard(divid, {
            boundingbox: [-3, 5, 10, -3], axis: true
        });
        
        // Plot functions only up to the mutual intersections
        var function1 = board.create('functiongraph', ['sqrt(x-3)', 4, 3], {name: 'x=3+y^2', withlabel: true});
        var function2 = board.create('functiongraph', ['-sqrt(x-3)', 3, 7]);
        var function3 = board.create('functiongraph', ['-2', 7, -2]);
        var function5 = board.create('functiongraph', ['-sqrt(-x+2)', -2, 2], {name: 'x=2-y^2', withlabel: true});
        var function6 = board.create('functiongraph', ['sqrt(-x+2)', 2, 1]);
        var function4 = board.create('functiongraph', ['1', 1, 4]);

         // The elements can be FunctionGraphs or Pairs of coordinates (encoding points)
        // properties is a dictionary of properties for the generated function.
        function make_filled_curve(board, element_list, properties) {
            var curve = board.create('curve', [[], []], properties);
            curve.updateDataArray = function () {
                this.dataX = [];
                this.dataY = [];
                for (const element of element_list) {
                    if (Array.isArray(element)) {
                        // Add the point to the curve
                        this.dataX.push(element[0]);
                        this.dataY.push(element[1]);
                    } else {
                        // Copy all points from function
                        this.dataX = this.dataX.concat(element.points.map(p =&gt; p.usrCoords[1]));
                        this.dataY = this.dataY.concat(element.points.map(p =&gt; p.usrCoords[2]));
                    }
                }
                // Close the loop
                this.dataX.push(this.dataX[0]);
                this.dataY.push(this.dataY[0]);
            };
            board.update();
        }

        // Plot different filled functions
        make_filled_curve(board, [function1, function2, function3, function5, function6, function4], { strokeWidth: 3, fillColor: 'green', fillOpacity: 0.3 });
      

[[/jsxgraph]]


<p>It is clear from the graph that the curves \({@expr1@}\) and \({@expr1@}\) with lines \(y={@y_1@}\) and \(y={@y_2@}\) are really bounded by the region indicated in green colour, (in between  the blue and red curves).
</p>
<p><b>Recall:</b> Area between two curves with respect to the y-axis.</p>
<p>If \(f(y)\) and \(g(y)\) are continuous on \([a, b]\) and \(f(y) &lt; g(y)\) for all \(y\) in \([a, b]\), then</p>
<p>\(Area = \displaystyle \int_{a}^{b} [g(y) -f(y)] dy\)</p>
<p>In other words if the area of the region bounded by \(x=f(y), \, x= g(y), \, y=a \, \text{and} \,y=b\) exists then the area if given by
</p>
<p>\(Area = \displaystyle \int_{a}^{b} | f(y) -g(y) | dy\)</p>
<p>Therefore,</p>
<p>
\(
\begin{align}
A &amp;= \displaystyle \int_{{@y_1@}}^{{@y_2@}} \left( {@rhs(expr1)@}-({@rhs(expr2)@}) \right) dy \\
   &amp;= \displaystyle \int_{{@y_1@}}^{{@y_2@}} \left(&nbsp;{@(rhs(expr1)-rhs(expr2))@}\right) dy \\
   &amp;= \displaystyle \left[{@area_func@}\right]_{{@y_1@}}^{{@y_2@}} \\
   &amp;= {@ta@}

\end{align}
\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Defining the question variables*/ 
const1:rand(5)+8; /*constant in the first function*/
const2:rand(5)-8; /*constant in the second function*/
expr1:x=const1+y^2; /*first function*/
expr2:x=const2-y^2; /*second function*/
y_1:rand(5)-8; /*lower limit*/
y_2:rand(5)+8; /*upper limit*/
x_1:rhs(ev(expr1,x,y=y_1)); /*min value of x*/
x_2:rhs(ev(expr2,x,y=y_2)); /*max value of x*/
area:rhs(expr1)-rhs(expr2); 
area_func:int(area,y);

/*Worked out solution*/
graph:plot([rhs(expr1),rhs(expr2)],[y,y_1,y_2]);

/*Model answer*/
ta:int(area,y,y_1,y_2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p><p>Determine the area of the region bounded by \({@expr1@}, \qquad {@expr2@}, \, y={@y_1@} \, \text{and} \,  y={@y_2@}\)</p>
<p>\({@ta@} \)</p>
<p>{@graph@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1 checks for correct answer*/
/*Node 2 mistake:Integrating the indefinite integral without the limits*/
mistake_2:int(area,y);
/*Node 3 mistake:Interchanging the limits of integration, however, with the correct definite integral.*/
mistake_3:int(area,y,y_2,y_1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have integrated the function correctly but did not input the limits. Take a look at the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your computation looks okay for the definite integral, however, it looks like you have interchanged the limits of integration.
Take a look at the worked solution below.
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>677693235</deployedseed>
    <deployedseed>1708962587</deployedseed>
    <deployedseed>372342529</deployedseed>
    <deployedseed>1928540007</deployedseed>
    <deployedseed>1611114569</deployedseed>
    <deployedseed>1672485668</deployedseed>
    <deployedseed>428365685</deployedseed>
    <deployedseed>1779230301</deployedseed>
    <deployedseed>932060850</deployedseed>
    <deployedseed>16712909</deployedseed>
    <deployedseed>1243057385</deployedseed>
    <deployedseed>1649965615</deployedseed>
    <deployedseed>353165177</deployedseed>
    <deployedseed>1504680450</deployedseed>
    <deployedseed>1665962385</deployedseed>
    <deployedseed>1087072534</deployedseed>
    <deployedseed>634946301</deployedseed>
    <deployedseed>1191986205</deployedseed>
    <deployedseed>1164745065</deployedseed>
    <deployedseed>1223800991</deployedseed>
  </question>
</quiz>
