<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Representation of exponential growth and decay of function</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Determine whether the function below represents exponential growth, exponential decay, or neither.</p>

	<ol type ="a">
    	<li>
<p>\( \displaystyle f(x)={@initial_x@}({@base_x@})^{x}\)</p>

                <p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
    	</li>

    	<li>
<p>\( \displaystyle f(t)={@initial_t@}({@base_t@})^{t}\)</p>
                <p>[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>
      	</li>
	</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To determine whether the given functions represent exponential growth, exponential decay, or neither, we need to examine the base of the exponential term in each function.</p><p> The general form of an exponential function is \( f(x) = a \cdot b^x \), where \(a\) is a constant and \(b\) is the base of the exponent.</p>
<ul>
<li>If \(b > 1\), the function represents exponential growth.<br></li>
<li> If \(0 < b < 1\), the function represents exponential decay.<br></li>
<li>If \(b = 1\) or \(b \leq 0\), it is neither exponential growth nor decay.</p></li>
</ul>
<p>Let's analyze each function:</p>
	<ol type ="a">
    	<li>
<p>\( \displaystyle f(x)={@initial_x@}({@base_x@})^{x}\)</p>
<p>In this function, the base \(b = {@base_x@}\)</p>

<p><strong>{@is_growth@}</strong>.</p>
 	</li>

    	<li>
<p>\( \displaystyle f(t)={@initial_t@}({@base_t@})^{t}\)</p>

<p>In this function, the base \(b =  {@base_t@}\).</p>

<p><strong>{@is_decay@}</strong>.</p>

    	</li>
	</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Part (a) */
/* Define a function to generate random numbers with a specified number of decimal places */
decimalplaces(num, places) := float(round(num * 10^places) / 10^places);

/* Function to randomize base_x and ensure a 50/50 split for >1 and <1 */
random_base_x() := block(
    if rand(2) = 0 then
        /* Generate base_x < 1 */
        decimalplaces(0.1 + rand(0.9), 2)
    else
        /* Generate base_x > 1 */
        decimalplaces(1.01 + rand(0.99), 2)
);

/* Function to randomize base_t and ensure a 50/50 split for >1 and <1 */
random_base_t() := block(
    if rand(2) = 0 then
        /* Generate base_t < 1 */
        decimalplaces(0.1 + rand(0.9), 2)
    else
        /* Generate base_t > 1 */
        decimalplaces(1.01 + rand(0.99), 2)
);

/* Generate random base_x and base_t */
base_x: random_base_x();
base_t: random_base_t();
/* Define the functions */
initial_x: rand_with_step(50,300,5);
initial_t: rand_with_prohib(50,300,[initial_x]);
/*base_x:decimalplaces(rand(1.06)+0.1,2);
base_t:decimalplaces(rand(0.97)+0.8,2);*/
fx : initial_x * (base_x)^x;
ft: initial_t * (base_t)^t;
is_growth: if base_x > 1 then  sconcat("\\( \\text{Exponential growth, the growth rate is}, \\, ",base_x," \\, \\text{is greater than} \\, ",1,"\\)")
elseif base_x < 1 and base_x > 0 then
sconcat("\\( \\text{Exponential decay, the decay factor is}\\, ",base_x,"  \\, \\text{is between} \\, ",0,"  \\, \\text{and } \\, ",1,"\\)");
/*Destructors*/
base_x1:decimalplaces(base_x+(rand(0.051)+0.01),2);
base_t1:decimalplaces(base_t+(rand(0.061)+0.1),2);
/* correct answer */
corr_ans:[[1, true, is_growth]];
/* wrong answers */
incorr_ans:[
[2, false, sconcat("\\( \\text{Exponential growth, the growth rate is}, \\, ",base_x1," \\, \\text{is greater than} \\, ",1,"\\)")],
[3, false, sconcat("\\( \\text{Exponential decay, the decay factor is}\\, ",base_t1,"  \\, \\text{is between} \\, ",0,"  \\, \\text{and } \\, ",1,"\\)")]];
ta1: random_permutation(append(corr_ans, incorr_ans));
/*Worked Soln*/
ta_index1: sort(makelist(x[1], x, ta1));

/* Part (b) */
is_decay: if base_t > 1 then  sconcat("\\( \\text{Exponential growth, the growth rate is}, \\, ",base_t," \\, \\text{is greater than} \\, ",1,"\\)")
elseif base_t < 1 and base_t > 0 then
sconcat("\\( \\text{Exponential decay, the decay factor is}\\, ",base_t,"  \\, \\text{is between} \\, ",0,"  \\, \\text{and } \\, ",1,"\\)");
/*Destructors*/
base_x2:base_x+0.010;
base_t2:base_t+0.101;
/* correct answer */
corr_anss:[[1, true, is_decay]];
/* wrong answers */
incorr_anss:[
[2, false, sconcat("\\( \\text{Exponential growth, the growth rate is}, \\, ",base_x2," \\, \\text{is greater than} \\, ",1,"\\)")],
[3, false, sconcat("\\( \\text{Exponential decay, the decay factor is}\\, ",base_t2,"  \\, \\text{is between} \\, ",0,"  \\, \\text{and } \\, ",1,"\\)")]];
ta2: random_permutation(append(corr_anss, incorr_anss));
/*Worked Soln*/
ta_index2: sort(makelist(x[1], x, ta2));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\( \displaystyle f(x)={@initial_x@}({@base_x@})^{x}\)</p>

<p><strong>{@is_growth@}</strong>.</p>

<p>\( \displaystyle f(t)={@initial_t@}({@base_t@})^{t}\)</p>

<p><strong>{@is_decay@}</strong>.</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p>
<p id="concept">https://docs.google.com/document/d/1JX3hBVYGCg3EQ7oPHRboEHHIALYeRERkPJZhfWwXp9Q/edit</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">2</p>
<p id="description">Given a function, students are required to determine when it represents an exponetiial growth or decay  </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans2</name>
      <type>radio</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta1)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta2)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>851100800</deployedseed>
    <deployedseed>1282692556</deployedseed>
    <deployedseed>744668170</deployedseed>
    <deployedseed>792516150</deployedseed>
    <deployedseed>151970586</deployedseed>
    <deployedseed>1273430042</deployedseed>
    <deployedseed>1588273397</deployedseed>
    <deployedseed>776550063</deployedseed>
    <deployedseed>1471749913</deployedseed>
    <deployedseed>2011769320</deployedseed>
    <deployedseed>548353478</deployedseed>
    <deployedseed>148616483</deployedseed>
    <deployedseed>1011698018</deployedseed>
    <deployedseed>1277551163</deployedseed>
    <deployedseed>1112907207</deployedseed>
    <deployedseed>1426139585</deployedseed>
    <deployedseed>1764024543</deployedseed>
    <deployedseed>874652993</deployedseed>
    <deployedseed>241725608</deployedseed>
    <deployedseed>1619536918</deployedseed>
    <deployedseed>618408848</deployedseed>
    <deployedseed>847393525</deployedseed>
    <deployedseed>1783302123</deployedseed>
    <deployedseed>1098420989</deployedseed>
    <deployedseed>1987100039</deployedseed>
    <deployedseed>1374024164</deployedseed>
    <deployedseed>2016010930</deployedseed>
    <deployedseed>1873172112</deployedseed>
    <deployedseed>584189174</deployedseed>
    <deployedseed>156745104</deployedseed>
    <deployedseed>171966077</deployedseed>
    <deployedseed>1836008977</deployedseed>
    <deployedseed>485553653</deployedseed>
    <deployedseed>259243278</deployedseed>
    <deployedseed>1573176634</deployedseed>
    <deployedseed>325924629</deployedseed>
    <deployedseed>2061645369</deployedseed>
    <deployedseed>1527755768</deployedseed>
    <deployedseed>210555701</deployedseed>
    <deployedseed>659808239</deployedseed>
    <deployedseed>1578912941</deployedseed>
    <deployedseed>1126957274</deployedseed>
    <deployedseed>1489489075</deployedseed>
    <deployedseed>1342575769</deployedseed>
    <deployedseed>1596250464</deployedseed>
    <deployedseed>145110</deployedseed>
    <deployedseed>1103980654</deployedseed>
    <deployedseed>968035510</deployedseed>
    <deployedseed>856953064</deployedseed>
    <deployedseed>1423527856</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta2))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
