<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Rolle's Theorem</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Rolle's Theorem is a fundamental theorem in calculus that applies to differentiable functions. It states that if a function \( f(x) \) is continuous on a closed interval \([a, b]\), differentiable on the open interval \((a, b)\), and \( f(a) = f(b) \), then there exists at least one point \( c \) in the open interval \((a, b)\) such that \( f'(c) = 0 \).</p>
<p></p><p></p>
<p>Let \( f(x) ={@fx@} \) on the interval \( \displaystyle [{@a@},{@b@}]\). Verify whether the conditions of Rolle's Theorem are satisfied, and if so, find the value of \( c \) where \( f'(c) = 0\).</p>

<ol>
    <li><p>Verify the conditions of Rolle's Theorem:</p>
        <ul>
            <li><p>\( f(x) = {@fx@} \) is continuous and differentiable on the interval \([{@a@},{@b@}]\).</p></li>
            <li><p>\(\displaystyle  f({@a@}) = \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p></li>
            <li><p>\( \displaystyle f({@b@}) = \) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p></li>
            <li><p>Since \(\displaystyle  f({@a@})\)[[input:ans3]] \(f\displaystyle ({@b@})\), the function [[input:ans4]] the conditions of Rolle's Theorem [[validation:ans3]][[feedback:prt3]] [[validation:ans4]][[feedback:prt4]]</p></li>
        </ul>
    </li>
    <li><p>Apply Rolle's Theorem:</p>
        <ul>
            <li><p>Since the conditions are met, there exists at least one value \( c \) in the interval \(({@a@}, {@b@})\) such that \( f'(c) = 0 \).</p></li>
            <li><p>Find the derivative of \( f(x) \): </p></li>
             <li><p>\( f'(x) =  \) [[input:ans5]] [[validation:ans5]][[feedback:prt5]]</p></li>
            <li><p>Set \( f'(c) = 0 \) and solve for \( c \):</p></li>
            <li>
           <p><b>Give your answer in exact form where necessary Hint: \(\displaystyle \frac{\sqrt{x}}{y} \) can be typed as sqrt(x)/y</b></p>
          <p>So, \( c = \) [[input:ans6]] or \( c = \) [[input:ans7]]</p>
           </li>
             <p>\(\qquad \qquad \) [[validation:ans6]] \(\quad \qquad \) [[validation:ans7]] [[feedback:prt6]] [[feedback:prt7]]</p>
        </ul>
    </li>
</ol>

<p>Therefore, according to Rolle's Theorem, there exists a point  \( c = \) [[input:ans8]] [[validation:ans8]] in the interval \( \displaystyle ({@a@}, {@(b)@})\) where \( f'(c) = 0 \).[[feedback:prt8]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Rolle's Theorem is a fundamental theorem in calculus that applies to differentiable functions. It states that if a function \( f(x) \) is continuous on a closed interval \([a, b]\), differentiable on the open interval \((a, b)\), and \( f(a) = f(b) \), then there exists at least one point \( c \) in the open interval \((a, b)\) such that \( f'(c) = 0 \).</p>
<p></p><p></p>
<p>Let \( f(x) ={@fx@} \) on the interval \( \displaystyle [{@a@},{@b@}]\). Verify whether the conditions of Rolle's Theorem are satisfied, and if so, find the value of \( c \) where \( f'(c) = 0\).</p>

<ol>
    <li><p>Verify the conditions of Rolle's Theorem:</p>
        <ul>
            <li><p>\( f(x) = {@fx@} \) is continuous and differentiable on the interval \([{@a@},{@b@}]\).</p></li>
            <li><p>\(\displaystyle  f({@a@}) = {@ta1@}\)</p></li>
            <li><p>\( \displaystyle f({@b@}) = {@ta2@}\)</p></li>
            <li><p>Since \(\displaystyle  f({@a@})\) {@mcq_correct(ta3)[1]@} \(f \displaystyle ({@b@})\), the function {@mcq_correct(ta4)[1]@} the conditions of Rolle's Theorem</p></li>
        </ul>
    </li>
    <li><p>Apply Rolle's Theorem:</p>
        <ul>
            <li><p>Since the conditions are met, there exists at least one value \( c \) in the interval \(({@a@}, {@b@})\) such that \( f'(c) = 0 \).</p></li>
            <li><p>Find the derivative of \( f(x) \): </p></li>
             <li><p>\( f'(x) =  {@ta5@}\)</p></li>
            <li><p>Set \( f'(c) = 0 \) and solve for \( c \):</p></li>
            <li><p>\( {@f_pime_c@} = 0 \).</p></li>
            <li><p>Factorize the above equation: \( {@factor(f_pime_c)@}= 0 \).</p></li>
            <li><p>Solve the factorized quadratic equation: \( {@factor(f_pime_c)@}= 0 \)</p></li>
            <li><p>So, \(\displaystyle c = {@ta6@}\) or \(\displaystyle c = {@ta7@}\)</p></li>
        </ul>
    </li>
</ol>
<p>Therefore, according to Rolle's Theorem, there exists a point \(\displaystyle c = {@ta7@} \) in the interval \( \displaystyle ({@a@}, {@(b)@})\) where \( f'(c) = 0 \).</p>




<p></p><p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[a: rand_with_step(-3,3,1);
b: a+rand(4)+1;
coef1: rand(5)+2;
coef2: rand_with_prohib(-9,9,[coef1]);
coef3: rand_with_step(-5,5,1);

fxx: coef1*x^3+coef2*x^2+t*x+coef3;
expr1: at(fxx,[x=a]);
expr2: at(fxx,[x=b]);
expr3: expr1-expr2;
solv: solve(expr3);
remov_br: rhs(solv[1]);
fx: subst(remov_br,t,fxx);
interva:[a,b];
fa:at(fx,[x=a]);
fb:at(fx,[x=b]);

satisfies: if fa=fb then "satisfies" else "does not satisfy";
sat:if satisfies="satisfies" then "does not satisfy" else "satisfies"; 

ss: if fa=fb then "=" else "≠";
st: if ss="=" then "≠" else "=";

f_prime:diff(fx,x);
f_pime_c:subst(c,x,f_prime);
sol_c:solve(f_pime_c=0);
c1:sol_c[1];
c2:sol_c[2];

/*Model answers*/
ta1:fa;
ta2:fb;
ta3: random_permutation([[ss,true],[st,false], ["<>",false]]);
ta4: random_permutation([[satisfies,true],[sat,false], [" is not clear with",false]]);
ta5:f_prime;
ta6:rhs(c1);
ta7:rhs(c2);
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[  <ul>
            <li><p>\( f(x) = {@fx@} \) is continuous and differentiable on the interval \([{@a@},{@b@}]\).</p></li>
            <li><p>\(\displaystyle  f({@a@}) = {@ta1@}\)</p></li>
            <li><p>\( \displaystyle f({@b@}) = {@ta2@}\)</p></li>
            <li><p>Since \(\displaystyle  f({@a@})\) {@mcq_correct(ta3)[1]@} \(f \displaystyle ({@b@})\), the function {@mcq_correct(ta4)[1]@} the conditions of Rolle's Theorem</p></li>
        </ul>
    </li>
   
 <li><p>\( f'(x) =  {@ta5@}\)</p></li>
            <li><p>\( {@f_pime_c@} = 0 \).</p></li>
            <li><p>Factorize the above equation: \( {@factor(f_pime_c)@}= 0 \).</p></li>
            <li><p>So, \(\displaystyle c = {@ta6@} \approx {@decimalplaces(ta6,2)@}\) or \(\displaystyle c = {@ta7@} \approx {@decimalplaces(ta7,2)@}\) within the interval \(({@a@}, {@b@})\).</p></li>
        </ul>
<p>Point \(\displaystyle c = {@ta7@} \approx {@decimalplaces(ta7,2)@} \in \displaystyle ({@a@}, {@(b)@})\) where \( f'(c) = 0 \).</p>
    </li>
</ol>
]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey Wabwirep>
<p id="concept">https://docs.google.com/document/d/1bgiDq1Yru63dqCi0IxRv7E9Y321DdM5r/edit</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">Verify the conditions of Rolle's Theorem.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>ta7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>ta7</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mcq_correct(ta3)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mcq_correct(ta4)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sans:{ans6,ans7};
tans:{ta6,ta7};

sans_d:{decimalplaces(ans6,2),decimalplaces(ans7,2)};
tans_d:{decimalplaces(ans6,2),decimalplaces(ans7,2)};</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>interchanging correct answers </description>
        <answertest>AlgEquiv</answertest>
        <sans>sans</sans>
        <tans>tans</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt6-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks for correct answer 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt6-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>checks for correct answer 2dp interchangeably</description>
        <answertest>AlgEquiv</answertest>
        <sans>sans_d</sans>
        <tans>tans_d</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sans:{ans6,ans7};
tans:{ta6,ta7};

sans_d:{decimalplaces(ans6,2),decimalplaces(ans7,2)};
tans_d:{decimalplaces(ans6,2),decimalplaces(ans7,2)};</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>interchanging correct answers </description>
        <answertest>AlgEquiv</answertest>
        <sans>sans</sans>
        <tans>tans</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt7-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks for correct answer 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans7</sans>
        <tans>ta7</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt7-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>checks for correct answer 2dp interchangeably</description>
        <answertest>AlgEquiv</answertest>
        <sans>sans_d</sans>
        <tans>tans_d</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>ta7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for correct answer 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans8</sans>
        <tans>ta7</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>826515442</deployedseed>
    <deployedseed>803249421</deployedseed>
    <deployedseed>1480437573</deployedseed>
    <deployedseed>423278955</deployedseed>
    <deployedseed>1338524981</deployedseed>
    <deployedseed>352509841</deployedseed>
    <deployedseed>1230870796</deployedseed>
    <deployedseed>236931952</deployedseed>
    <deployedseed>17171297</deployedseed>
    <deployedseed>1794502613</deployedseed>
    <deployedseed>590422579</deployedseed>
    <deployedseed>1093794173</deployedseed>
    <deployedseed>8752643</deployedseed>
    <deployedseed>1671981225</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(ta3))</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>first(mcq_correct(ta4))</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta6</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>ta7</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>ta7</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
