<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Solving Exponential and Logarithmic Equations</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Solve for \(x\) in the equations below:</p>

<ol class="HELM_parts">
        <li>
<p> \(\displaystyle {@eqn_1@}\)</p>
<p>\(x= \) [[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>
    </li>

        <li>
<p> \(\displaystyle {@eqn_2@}\)</p>
<p>\(x= \)  [[input:ans2]] [[validation:ans2]]</p>
<p>[[feedback:prt2]]</p>
  </li>

        <li>
<p> \(\displaystyle {@eqn_3@}\)
</p>
<p>\(x= \) [[input:ans3]] [[validation:ans3]]</p>
<p>[[feedback:prt3]]</p>
    </li>
   
        </ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>var_1:rand(5)+2;
var_2:rand_with_prohib(1,7,[var_1]);
var_3:rand_with_prohib(1,7,[var_1,var_2]);
const_1:rand(8)+2;
const_2:rand_with_prohib(1,7,[var_1,var_2,var_3]);
var_5:rand_with_prohib(1,7,[var_1,var_2,var_3,const_2]);
const_3:rand_with_prohib(1,7,[var_1,var_2,var_3,const_2,var_5]);
eqn_1:e^(var_1*x+var_2)=const_1;
eqn_2:e^(var_3*x)=const_2;
eqn_3:e^(e^(var_5*x))=const_3;
ta_1:float(rhs(solve(eqn_1,x)[1]));
ta_2:float(rhs(solve(eqn_2,x)[1]));
ta_3:float((solve(eqn_3,x)[1]));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\({@eqn_1@}={@ta_1@}\)</p>
<p>\({@eqn_2@}={@ta_2@}\)</p>
<p>\({@eqn_3@}={@ta_3@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">https://docs.google.com/document/d/176hYh0jAetw-rifazqmuye5TC2UG4ztVJR0E7RUl0kw/edit#heading=h.eddp73mos3v3</p>
<p id="reviewer"></p>
<p id="reviewer">[Reviewer 2 Name]</p>
<p id="description">This question involves solving equations involving exponentials for the variable \(x\). </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta_3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2107214213</deployedseed>
    <deployedseed>24085954</deployedseed>
    <deployedseed>81694635</deployedseed>
    <deployedseed>697146984</deployedseed>
    <deployedseed>901237995</deployedseed>
    <deployedseed>744051283</deployedseed>
    <deployedseed>567400220</deployedseed>
    <deployedseed>1643968953</deployedseed>
    <deployedseed>11789328</deployedseed>
    <deployedseed>1453359232</deployedseed>
    <deployedseed>1054117113</deployedseed>
    <deployedseed>1342114548</deployedseed>
    <deployedseed>373645459</deployedseed>
    <deployedseed>883050872</deployedseed>
    <deployedseed>1598494136</deployedseed>
    <deployedseed>972192453</deployedseed>
    <deployedseed>1629372789</deployedseed>
    <deployedseed>465108509</deployedseed>
    <deployedseed>1233515882</deployedseed>
    <deployedseed>1187630521</deployedseed>
    <deployedseed>1968773041</deployedseed>
    <deployedseed>441996190</deployedseed>
    <deployedseed>1175599297</deployedseed>
    <deployedseed>2145769695</deployedseed>
    <deployedseed>1501125318</deployedseed>
    <deployedseed>36384853</deployedseed>
    <deployedseed>1180558934</deployedseed>
    <deployedseed>858152791</deployedseed>
    <deployedseed>736063</deployedseed>
    <deployedseed>1662076851</deployedseed>
  </question>
</quiz>
