<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Stating whether the change is exponential growth or decay and the relevant factor</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p><p>The relationship between the elapsed time \(t\), in minutes, since a medicine is ingested and its concentration in the bloodstream, \(C(t)\) in mg/l is modeled by the following function:</p><p>\[C(t) = {@aa@}\cdot ({@dispdp(bb,2)@})^t\]</p><p>Every minute, the concentration of the medicine changes by a factor. State whether the change is exponential growth or decay and the relevant factor:</p><p>Change: [[input:ans1]]&nbsp;[[validation:ans1]]&nbsp;[[feedback:prt1]]</p><p>Factor: [[input:ans2]]&nbsp;[[validation:ans2]]&nbsp;[[feedback:prt2]]<br></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>An exponential function has the form \(f(x) = k\cdot a^x\). </p>
<p>If \(a&gt;1\), the function represents exponential growth. </p><p>If \(0&lt;a&lt;1\), the function represents exponential decay. </p><p>The growth or decay factor is \(a\). In this case, the function represents exponential decay with a decay factor of \({@dispdp(bb,2)@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[aa:rand(40)+80;
bb:(rand(40)+60)/100;
ta:[[1,true,"Decay"], [2,false, "Growth"]];
ta:random_permutation(ta);
tac:1;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(C(t) = {@aa@}\cdot ({@dispdp(bb,2)@})^t\)</p>
<p>\({@(ta)@}\)</p>
<p>\({@bb@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo </p>
<p id="concept">N/A</p>
<p id="reviewer"> Godfrey Wabwire</p>
<p id="reviewer">Juma Zevick</p>
<p id="description">Stating whether the change is exponential growth or decay and the relevant factor every minute, the concentration of the medicine changes by a factor.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>dispdp(bb,2)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>tac</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>bb</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>bb</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1350152978</deployedseed>
    <deployedseed>1702216432</deployedseed>
    <deployedseed>997946754</deployedseed>
    <deployedseed>298072509</deployedseed>
    <deployedseed>908115626</deployedseed>
    <deployedseed>77498397</deployedseed>
    <deployedseed>2008416940</deployedseed>
    <deployedseed>1037369795</deployedseed>
    <deployedseed>379227687</deployedseed>
    <deployedseed>624717887</deployedseed>
    <deployedseed>2112719627</deployedseed>
    <deployedseed>615159002</deployedseed>
    <deployedseed>1721712389</deployedseed>
    <deployedseed>1085925604</deployedseed>
    <deployedseed>746560585</deployedseed>
    <deployedseed>683937670</deployedseed>
    <deployedseed>296617664</deployedseed>
    <deployedseed>662068335</deployedseed>
    <deployedseed>1498942789</deployedseed>
    <deployedseed>887750481</deployedseed>
    <deployedseed>1513549224</deployedseed>
    <deployedseed>1850273996</deployedseed>
    <deployedseed>1910052868</deployedseed>
    <deployedseed>324183023</deployedseed>
    <deployedseed>1682761024</deployedseed>
    <deployedseed>1175072147</deployedseed>
    <deployedseed>1304356163</deployedseed>
    <deployedseed>278284520</deployedseed>
    <deployedseed>1613495035</deployedseed>
    <deployedseed>1033796627</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dispdp(bb,2)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
