<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>The Work Required to Stretch or Compress a Spring</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Suppose it takes a force of \({@force@} \textbf{N}\) (in the negative direction) to compress a spring \({@float(spring_compress)@} \textbf{m}\) from equilibrium position. How much work is done to stretch the spring from \({@float(spring_stretch)@}\textbf{m}\) from the equilibrium</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>First, find the spring constant, \(\textbf{k}\). When \(x=-{@spring_compress@}\) we know \(F(x)=-{@force@}\) so,</p>
<p>\[F(x)=k(x)\]</p>
<p>\[-{@force@}=k(-{@spring_compress@})\]</p>
<p>\[k={@sol2@}\]</p>
<p>and \(F(x)={@sol2@}x\). Then to calculate work, we integrate the force function, obtaining</p>
<p>\(\begin{align}
W & =\displaystyle \int_a^b F(x)dx\\
&=\displaystyle \int_0^{@spring_stretch@}{@sol2@}x \,dx\\
&=\displaystyle \big[ {@sol4@} \big]_{0}^{@spring_stretch@}\\
&={@sol5@}\\
& \approx {@decimalplaces(sol5,2)@}\\
\end{align}
\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*workdone in calculus*/
force:rand(20)+10;
spring_compress:(rand_with_step(1,8,2)/10);
spring_stretch:(rand_with_step(1,8,1)/10);
/*variables*/
sol1:-force=k(-spring_compress);
sol2:-force/(-spring_compress);
sol3:F(x)=sol2*(x);
sol4:int(sol2*(x),x);
sol5:int(sol2*(x),x,0,spring_stretch);
ta:sol5;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Suppose it takes a force of \({@force@}\) (in the negative direction) to compress a spring \({@spring_compress@} \textbf{m}\) from equilibrium position.How much work is done to stretch the spring from \({@spring_stretch@}\textbf{m}\) from the equilibrium</p>
<p>{@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Wastalas</p>
<p id="description">The question focuses on finding  how much work is done to stretch the spring of given  m from the equilibrium position? Suppose it takes a given force of 10 N (in the negative direction) to compress a spring of given 0.2 m from the equilibrium position </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p style="color: blue;">Well done for giving your answer correct to \(2\) decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>291500895</deployedseed>
    <deployedseed>1797623290</deployedseed>
    <deployedseed>1776988627</deployedseed>
    <deployedseed>1057634442</deployedseed>
    <deployedseed>1674502778</deployedseed>
    <deployedseed>1697241671</deployedseed>
    <deployedseed>1833347901</deployedseed>
    <deployedseed>2012981451</deployedseed>
    <deployedseed>893222553</deployedseed>
    <deployedseed>78813960</deployedseed>
    <deployedseed>1363513247</deployedseed>
    <deployedseed>774369728</deployedseed>
    <deployedseed>933456953</deployedseed>
    <deployedseed>297574785</deployedseed>
    <deployedseed>446089149</deployedseed>
    <deployedseed>931398054</deployedseed>
    <deployedseed>1750016273</deployedseed>
    <deployedseed>1922613565</deployedseed>
    <deployedseed>928667140</deployedseed>
    <deployedseed>1258709831</deployedseed>
    <deployedseed>920365095</deployedseed>
    <deployedseed>1327349869</deployedseed>
    <deployedseed>928246072</deployedseed>
    <deployedseed>581474435</deployedseed>
    <deployedseed>1506912550</deployedseed>
    <deployedseed>1653117557</deployedseed>
    <deployedseed>1206735815</deployedseed>
    <deployedseed>1806936672</deployedseed>
    <deployedseed>1227425934</deployedseed>
    <deployedseed>2035337063</deployedseed>
    <deployedseed>1265290054</deployedseed>
    <deployedseed>723672977</deployedseed>
    <deployedseed>149468633</deployedseed>
    <deployedseed>2080603454</deployedseed>
    <deployedseed>1810820571</deployedseed>
    <deployedseed>1738853367</deployedseed>
    <deployedseed>1848582738</deployedseed>
    <deployedseed>779264575</deployedseed>
    <deployedseed>85279119</deployedseed>
    <deployedseed>463580211</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
