<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Trigonometric Integration</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the antiderivative of the function \( \displaystyle \int {@trig_func@}d{@xx@}\) by use of the sine reduction formula.</p>
<p>  [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b>Worked solution</b></p>
<p>We recall a function \( \displaystyle \int \sin^{n}xdx \) whose antiderivative can be computed using sine reduction formula as follows:
</p>
<p>\[ \displaystyle \int \sin^{n}xdx = \frac{n-1}{n} \int \sin^{n-2}xdx-\frac{\cos(x)\sin^{n-1}(x)}{n} \]</p>
<p>Then for \(n={@nn@}\);</p>
<p>\[ \displaystyle \int \sin^{{@nn@}}{@xx@}d{@xx@} = \frac{{@nn@}-1}{{@nn@}}\int \sin^{{@nn@}-2}{@xx@}d{@xx@}-\frac{\cos({@xx@})\sin^{{@nn@}-1}({@xx@})}{{@nn@}} \]</p>

<p>\[\displaystyle \int \sin^{{@nn@}}{@xx@}d{@xx@} = \frac{{@nn-1@}}{{@nn@}} \int \sin^{{@nn-2@}}{@xx@}d{@xx@}-\frac{\cos({@xx@})\sin^{{@nn-1@}}({@xx@})}{{@nn@}} \qquad \qquad (*) \]</p>
<p>But \( \displaystyle  \frac{{@nn-1@}}{{@nn@}} \int \sin^{{@nn-2@}}{@xx@}d{@xx@}= \frac{{@nn-1@}}{{@nn@}}\left({@int_trig@}\right)\),&nbsp; \(\sin^{2}x=1-\cos^{2}x\) and \(\cos^{2}x= \frac{1-\cos(2x)}{2}\).</p><p><span style="font-size: 0.9375rem;">Therefore, the expression \((*)\) becomes;</span></p><p><span style="font-size: 0.9375rem;">\[ \displaystyle \quad \int \sin^{{@nn@}}{@xx@}d{@xx@}  = {@ta@} \]</span><br></p>

<p>Where \(c\) is a constant of integration.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Defining the question*/
xx:rand([x,y,z,t]);/*variable of integration*/
nn:rand(5)+2;/*the power - lowest value 2*/
trig_func:sin(xx)^nn;/*the function*/

trig:sin(xx)^(nn-2); /*the first part of the reduction formula*/
int_trig:int(trig,xx)+c;/*the integral of the first part of the reduction formula*/

/*Model answer*/
ta:int(trig_func,xx)+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@trig_func@}= {@ta@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1 checking for the correct answer*/

/*node 2 checking if the student is missing the constant of integration*/
mistake_2:int(trig_func,xx);

/*node 3: mistake:Dividing with the wrong power*/
mistake_3:int_trig:int(trig,xx)+c;

/*node 4: mistake: Differentiating instead of integrating*/
mistake_4:diff(trig_func);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>xx</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have forgotten the constant of integration.</p>
<p>Recall, \( \displaystyle \int x^n dx=\frac{x^{n+1}}{n+1}+c\), where \(c\) is the constant of integration.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You calculated your integral using the wrong power. The static power is \({@nn@}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have differentiated the function instead of integrating.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>731982094</deployedseed>
    <deployedseed>1687566889</deployedseed>
    <deployedseed>1654204275</deployedseed>
    <deployedseed>1966196131</deployedseed>
    <deployedseed>1657280105</deployedseed>
    <deployedseed>471859488</deployedseed>
    <deployedseed>1169707456</deployedseed>
    <deployedseed>524804392</deployedseed>
    <deployedseed>1609103486</deployedseed>
    <deployedseed>2082411839</deployedseed>
    <deployedseed>1053870859</deployedseed>
    <deployedseed>140248220</deployedseed>
    <deployedseed>1111804459</deployedseed>
    <deployedseed>638200866</deployedseed>
    <deployedseed>1018025540</deployedseed>
    <deployedseed>195870431</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
