<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Understanding Logarithmic Functions with Base 1 and Graphs part a</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>How is the Logarithmic function \(y = \log_a x\) defined?</p>
        <p>[[input:ans1]] [[validation:ans1]]</p>
        <p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>How is the Logarithmic function \(y = \log_a x\) defined?</p>
<p>The logarithmic function \(y = \log_a x\) is defined as the inverse function of the exponential function \(y = a^x\), where \(a\) is the base of the logarithm. In other words, the logarithmic function undoes what the exponential function does.
</p>
<p>It is defined as for \(x>0 ,\,a>0\) and \(a \neq 1,\) </p>
<p>\(y = \log_a x \Leftrightarrow  a^y=x\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*part a*/
correct_answer:[
[1, true, "The logarithmic function is an inverse function of the exponential function"]];

incorrect_answers:[
[2, false, "The equation \\(y = \\log_a x\\) is equivalent to \\(y = x^a\\)."],
[3, false, "The base \\(a\\) of the logarithm must be a negative real number greater than \\(1\\)."],
[4, false, "The argument \\(x\\) of the logarithm must be a negative real number."]
    ];
ta:random_permutation(append(correct_answer,incorrect_answers));

ta_index: sort(makelist(x[1], x, ta));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<br>{@(ta)@}]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey Wabwire</p><p>
</p><p id="concept">https://docs.google.com/document/d/1gopE2Y3AnKuL517nuCiGuIHFFTIWX4t7Je88qIPcyow/edit</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">2</p>
<p id="description">This question delves into the peculiar nature of logarithmic functions with a base of 1, examining their definition and characteristics. It explores the domain and range of the function \( y = \log_{1}x \) and provides insights into its graph. Additionally, it sketches the general shape of the graph of logarithmic functions \( y = \log_{a}x \) when \( a > 1 \). </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: checks for all the correct answer from the list*/

/*Node 2: Checks if option 2 (incorrect) was provided and selected*/
true_opt_2:if member(2,ta_index) then ( if member(2,ans1) then 1 else 0 ) else 0;

/*Node 3: Checks if option 3 (incorrect) was provided and selected*/
true_opt_3:if member(3,ta_index) then ( if member(3,ans1) then 1 else 0 ) else 0;

/*Node 4: Checks if option 4 (incorrect) was provided and selected*/
true_opt_4:if member(4,ta_index) then ( if member(4,ans1) then 1 else 0 ) else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the corrrect answer from the list</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks if option 2 (incorrect) was provided and selected</description>
        <answertest>AlgEquiv</answertest>
        <sans>true_opt_2</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The equation \(y = \log_a x\) is equivalent to \(x = a^y\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks if option 3 (incorrect) was provided and selected</description>
        <answertest>AlgEquiv</answertest>
        <sans>true_opt_3</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The base \(a\) of the logarithm must be a positive real number greater than \(0\) and \(a \neq 1\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Checks if option 4 (incorrect) was provided and selected</description>
        <answertest>AlgEquiv</answertest>
        <sans>true_opt_4</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The argument \(x\) of the logarithm must be a positive real number.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>930029979</deployedseed>
    <deployedseed>2137984753</deployedseed>
    <deployedseed>1432229719</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
