<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Use Trigonometric identity to evaluate the Integral</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find \( \int {@p@} d{@v@}. \)</p>  
<p>Hint: Use a trigonometric identity and note that the derivative of \(\tan(t)\) is \(\sec^2(t)\)</p> 
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>For this question we need to use the identity \[tan^2(A) = sec^2(A) - 1\] as follows: \[{@p@} = {@q@}\] 
<p>So,
 \(\begin{align}
\displaystyle \int {@p@} d{@v@}&=\displaystyle \int ({@q@}) d{@v@} \\ 
 &=\displaystyle \int {@q1@} d{@v@} + \int {@q2@} d{@v@} \\ 
 &=\displaystyle {@ta@}\\
\end{align}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>v:t;
rdm:2+rand(5);
p:(tan(rdm*v))^2;
q1:(sec(rdm*v))^2;
q2:-1;
q:q1+q2;
ta:int(q,v)+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\int {@p@}d{@v@} ={@ta@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Godfrey</p>
<p id="concept">N/A</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">The question focuses on evaluating an Integral that Involves  a tan function by use of trigonometric identity</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>Int</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>v</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1418325152</deployedseed>
    <deployedseed>30222064</deployedseed>
    <deployedseed>1399734088</deployedseed>
    <deployedseed>240162612</deployedseed>
    <deployedseed>1758399097</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>int(q,v)+c</value>
      </testinput>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>int(q,v)</value>
      </testinput>
    </qtest>
  </question>
</quiz>
