<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Use Trig Substitutions to evaluate the Integral</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p> Use trig substitution to evaluate
 \(\displaystyle\int{@f_w@}dw\)</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>\( \text{Worked Solution}\)</p>
<p></p>
<p>
</p>
<p> \( \text{In order to carry out trig substitution integrals, let's take a look at the information in the table below.} \)<br></p>
<div>
    <table width="750" cellspacing="1" cellpadding="0" border="1">
        <colgroup>
            <col>
            <col>
            <col>
            <col>
        </colgroup>
        <tbody>
            <tr>
                <td>
                    <p dir="ltr" style="text-align: left;">\( \text{If you see }\)<br></p>
                </td>
                <td>
                    <p dir="ltr" style="text-align: left;">\( \text{Use the following substitution}\)<br></p>
                </td>
                <td>
                    <p dir="ltr" style="text-align: left;">\(\text{Where}\)</p>
                </td>
                <td>
                    <p dir="ltr" style="text-align: left;">\(\text{So that,}\)</p>
                </td>
            </tr>
            <tr>
                <td>
                    <p dir="ltr" style="text-align: left;">\(\sqrt{a^2-x^2}\)<br></p>
                </td>
                <td>
                    <p dir="ltr" style="text-align: left;"> \(x=a\sin(\theta)\)</p>
                </td>
                <td>
                    <p dir="ltr" style="text-align: left;">\(dx=a\cos(\theta)d\theta\)<br></p>
                </td>
                <td>
                    <p dir="ltr" style="text-align: left;">\(\sqrt{a^2-x^2} = a\cos(\theta)\)<br></p>
                </td>
            </tr>
            <tr>
                <td>
                    <p dir="ltr" style="text-align: left;">\(\sqrt{a^2+x^2}\)<br></p>
                </td>
                <td>
                    <p dir="ltr" style="text-align: left;">\(x=a\tan(\theta)\)</p>
                </td>
                <td>
                    <p dir="ltr" style="text-align: left;">\(dx=a\sec^{2}(\theta)d\theta\)<br></p>
                </td>
                <td>
                    <p dir="ltr" style="text-align: left;">\(\sqrt{a^2+x^2} = a\sec(\theta)\)<br></p>
                </td>
            </tr>
            <tr>
                <td>
                    <p dir="ltr" style="text-align: left;">\(\sqrt{x^2-a^2}\)<br></p>
                </td>
                <td>
                    <p dir="ltr" style="text-align: left;">\(x=a\sec(\theta)\)</p>
                </td>
                <td>
                    <p dir="ltr" style="text-align: left;">\(dx=a\sec(\theta)\tan(\theta)d\theta\)<br></p>
                </td>
                <td>
                    <p dir="ltr" style="text-align: left;">\(\sqrt{x^2-a^2} = a\tan(\theta)\)<br></p>
                </td>
            </tr>
        </tbody>
    </table>
</div>
<p></p>
<p>Let \(w=\frac{1}{\sqrt{{@coeff_w@}}} \sin{\theta}\). Then, \(w^2 = \frac{1}{{@coeff_w@}}\sin^{2}{\theta}\), and so,</p>

<p>
    \(
    \begin{align}
 \displaystyle   \sqrt{1-{@coeff_w@} w^2} &amp;= \sqrt{1-\sin^2{\theta}} \\
     \displaystyle  &amp;= \sqrt{\cos^2{\theta}} \\
  \displaystyle     &amp;= |\cos{\theta}|.
    \end{align}
    \)
</p>

<p>As we are calculating an indefinite integral, we will take the positive square root, and so \( \displaystyle  \sqrt{1-{@coeff_w@} w^2} = \cos{\theta}\).</p>

<p>Also, \( \displaystyle  dw = \frac{1}{\sqrt{{@coeff_w@}}} \cos{\theta} d{\theta}\).</p>

<p>Therefore, substituting into the integral,</p>

<p>
    \(
    \begin{align}
 \displaystyle     \int \sqrt{1-{@coeff_w@} w^2} dw &amp;= \int \cos{\theta} \cdot \frac{1}{\sqrt{{@coeff_w@}}} \cos{\theta} d\theta \\
 \displaystyle     &amp;= \frac{1}{\sqrt{{@coeff_w@}}} \int \cos^2{\theta} d\theta \\
   \displaystyle   &amp;= \frac{1}{\sqrt{{@coeff_w@}}} \int \frac{1}{2} (1+\cos{2\theta}) d\theta &amp;&amp; \text{as } \cos{2\theta} = 2\cos^2{\theta} - 1 \\
 \displaystyle     &amp;= \frac{1}{\sqrt{{@coeff_w@}}} \cdot \frac{1}{2} \left( \theta + \frac{1}{2} \sin{2\theta} \right) +c \\
 \displaystyle     &amp;= \frac{1}{2\sqrt{{@coeff_w@}}} ( \theta + \sin{\theta} \cos{\theta} ) + c &amp;&amp; \text{as }  \sin{2\theta}=2\sin{\theta} \cos{\theta}.
    \end{align} \\
    \)
</p>

<p>Since \( \displaystyle  w=\frac{1}{\sqrt{{@coeff_w@}}} \sin{\theta}\) and \(\cos{\theta} = \sqrt{1-{@coeff_w@} w^2}\), \(\sin{\theta} = \sqrt{{@coeff_w@}}w\) and \(\theta = \sin^{-1}(\sqrt{{@coeff_w@}}w)\).</p>

<p>Substituting these back,</p>

<p>
    \(
    \begin{align}
 \displaystyle     \int \sqrt{1-{@coeff_w@} w^2} dw &amp;= \frac{1}{2\sqrt{{@coeff_w@}}} (\theta + \sin{\theta} \cos{\theta} ) + c \\
 \displaystyle     &amp;= \frac{1}{2\sqrt{{@coeff_w@}}} \left( \sin^{-1}(\sqrt{{@coeff_w@}}w) + \sqrt{{@coeff_w@}}w \sqrt{1-{@coeff_w@} w^2} \right) +c
    \end{align}
    \)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>coeff_w:rand([2,3,5,6,7,8,10,11,12,13,14]);
f_w:sqrt(1-coeff_w*w^(2));

ta:integrate(f_w,w)+c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text> \(\displaystyle\int{@f_w@}dw={@ta@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">The question focuses on evaluating a integral Trig Functions by Trig substitution</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node1:checking correct answer*/

/*node 2: missing constant of integration*/
mistake1:integrate(f_w,w);

/*node3:Differentiating instead of integrating*/
mistake2:diff(f_w,w);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have missed the constant of integration. This is an indefinite integral. </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have differentiated instead of integrating it.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>843570352</deployedseed>
    <deployedseed>489747789</deployedseed>
    <deployedseed>18884920</deployedseed>
    <deployedseed>1466509550</deployedseed>
    <deployedseed>222413603</deployedseed>
    <deployedseed>1601001285</deployedseed>
    <deployedseed>1759116074</deployedseed>
    <deployedseed>56609122</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
