<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Using Product rule to find the derivatives</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the following derivatives of \(\displaystyle f(x) ={@f(x)@}\).</p>
<p>\(\quad \displaystyle f'(x)=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>\(\quad \displaystyle f''(x)=\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p> Inorder to find the derivative of \( \displaystyle f(x) = {@f(x)@}\) we will proceed as follows:</p>

  <p>Let \(u(x) =  {@func1@}\) and \(v(x) = {@func2@}\).</p>
  <p>Differentiating \(u\) and \(v\) with respect to \(x\) respectively we get;</p>


<p>
\(u'(x) = {@diff_func1@}\) and \(
v'(x) =  {@diff_func2@}
\)
</p>

 <p>Applying  product rule formula:</p>
<p> \((uv)' = u'v + uv'\)<p>
\(
({@f(x)@})' = ( {@diff_func1@}) \cdot ({@func2@}) + ({@func1@}) \cdot ( {@diff_func2@})
\)
</p>


  <p>Simplifying the expression we get;</p>

<p>
\(
( {@f(x)@})' = {@ta@}
\)
</p>

<p>Therefore the first derivative of \( f(x)={@f(x)@}\) is:</p>
<p>\( f'(x)= {@ta@}\).</p>
<p>For the second derivative \( f''(x)\) we will still apply the product rule formula and repeat the same steps</p>
<p>Therefore the second derivative is \( f''(x)= {@ta2@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/* Define the function */
var_x:rand(8)+2;
power:rand(3)+1;
trig:rand([sin,cos,tan]);
func1:var_x*x^power;
diff_func1:diff(func1,x);
func2:rand([trig(x),e^x,(x-1)^2,(e^(-2*x)),(x+3)^2]);
diff_func2:diff(func2,x);

f(x) := func1 * func2;

/* Find the derivative */ 
ta : diff(f(x), x);  
ta2:diff(ta,x);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p> \(\displaystyle f(x) ={@f(x)@}\).</p>
<p> \(f'(x)={@ta@}\)</p>
<p> \(f''(x)={@ta2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">Using method of differentiation (Product Rule)</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>Diff</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1145545241</deployedseed>
    <deployedseed>487824712</deployedseed>
    <deployedseed>1116513571</deployedseed>
    <deployedseed>2060189887</deployedseed>
    <deployedseed>2145545095</deployedseed>
    <deployedseed>463448715</deployedseed>
    <deployedseed>1861602491</deployedseed>
    <deployedseed>873869702</deployedseed>
    <deployedseed>691356043</deployedseed>
    <deployedseed>295072351</deployedseed>
    <deployedseed>1157565181</deployedseed>
    <deployedseed>1188676397</deployedseed>
    <deployedseed>1326322556</deployedseed>
    <deployedseed>1955414411</deployedseed>
    <deployedseed>1526030615</deployedseed>
    <deployedseed>546667644</deployedseed>
    <deployedseed>1336889223</deployedseed>
    <deployedseed>600967924</deployedseed>
    <deployedseed>1886231772</deployedseed>
    <deployedseed>1796115263</deployedseed>
    <deployedseed>476514700</deployedseed>
    <deployedseed>200424157</deployedseed>
    <deployedseed>427188748</deployedseed>
    <deployedseed>1846749532</deployedseed>
    <deployedseed>1184753563</deployedseed>
    <deployedseed>1460146329</deployedseed>
    <deployedseed>1021015460</deployedseed>
    <deployedseed>1897796864</deployedseed>
    <deployedseed>2018684811</deployedseed>
    <deployedseed>2086934107</deployedseed>
    <deployedseed>1410568769</deployedseed>
    <deployedseed>1805569136</deployedseed>
    <deployedseed>961990739</deployedseed>
    <deployedseed>648360725</deployedseed>
    <deployedseed>707704754</deployedseed>
    <deployedseed>158071509</deployedseed>
    <deployedseed>1368649081</deployedseed>
    <deployedseed>1020920861</deployedseed>
    <deployedseed>2009779390</deployedseed>
    <deployedseed>530980150</deployedseed>
    <deployedseed>363834759</deployedseed>
    <deployedseed>1233819997</deployedseed>
    <deployedseed>2033982284</deployedseed>
    <deployedseed>1635867976</deployedseed>
    <deployedseed>400155930</deployedseed>
    <deployedseed>461247473</deployedseed>
    <deployedseed>1853263200</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
