<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Volumes of Solids of Revolution-Washer method-q1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the volume of the solid generated by revolving the region bounded by \(y = {@f1(x)@}\) and \(y = {@f2(x)@}\) about the \(x\)‐axis.</p>
<p>Here is a sketch of the bounded region with the axis of rotation shown,</p>
<p>
[[jsxgraph]]
        var board = JXG.JSXGraph.initBoard(divid, { 
            boundingbox: [{#minx#}, {#maxy#}, {#maxx#}, {#miny#}],
            keepaspectratio: true,
            showCopyright: false,
            axis: true,
            showNavigation: false
        });
    // Define the functions
    var f1 = function(x) { return x*x + {#const1#}; };
    var f2 = function(x) { return x + {#const2#}; };

    // Plot the curves
    var curve1 = board.create('functiongraph', [f1, -3, 3], { strokeColor: 'blue' });
    var curve2 = board.create('functiongraph', [f2, -3, 3], { strokeColor: 'red' });

// Find the intersection points
    var interPoint1 = board.create('intersection', [curve1, curve2, 0], { visible: true, name:'\\(({#a#},{#x1#})\\)'});
    var interPoint2 = board.create('intersection', [curve1, curve2, 1], { visible: true, name:'\\(({#b#},{#y1#})\\)'});


    // Find the minimum point of the curve
    var minPoint = board.create('point', [0, f1(0)], { style: 5, size: 4, fillColor: 'green', strokeColor: 'green', visible: false});
    var minPointY = minPoint.Y();

    // Create the first curly bracket
    var height1 = 0.5; // height of the first curly bracket
    var crl1 = board.create('curve', [[0],[0]], {strokeWidth:1, strokeColor:'black' }, {visible: false});
    crl1.bezierDegree = 3;

    crl1.updateDataArray = function() {
        var A = minPoint, B = board.create('point', [minPoint.X(), 0], {visible: false});
        var d = [B.X() - A.X(), B.Y() - A.Y()];
        var dl = Math.sqrt(d[0] * d[0] + d[1] * d[1]);
        var mid = [(A.X() + B.X()) * 0.5, (A.Y() + B.Y()) * 0.5];

        d[0] *= height1 / dl;
        d[1] *= height1 / dl;

        this.dataX = [ A.X(), A.X() - d[1], mid[0], mid[0] - d[1], mid[0], B.X() - d[1], B.X() ];
        this.dataY = [ A.Y(), A.Y() + d[0], mid[1], mid[1] + d[0], mid[1], B.Y() + d[0], B.Y() ];
    };

    // Text for the first curly bracket
    var txt1 = board.create('text', [
                        function() { 
                            var A = minPoint, B = board.create('point', [minPoint.X(), 0], {visible: false});
                            var d = [B.X() - A.X(), B.Y() - A.Y()];
                            var dl = Math.sqrt(d[0] * d[0] + d[1] * d[1]);
                            var mid = (A.X() + B.X()) * 0.5;

                            d[1] *= height1 / dl;
                            return mid - d[1] + 0.1; 
                        },
                        function() { 
                            var A = minPoint, B = board.create('point', [minPoint.X(), 0], {visible: false});
                            var d = [B.X() - A.X(), B.Y() - A.Y()];
                            var dl = Math.sqrt(d[0] * d[0] + d[1] * d[1]);
                            var mid = (A.Y() + B.Y()) * 0.5;

                            d[0] *= height1 / dl;
                            return mid + d[0] + 0.1; 
                        },
                        () => 'g(x)={#f1(x)#}'
                    ]);

    // Create the second curly bracket
    var height2 = 0.5; // height of the second curly bracket
    var crl2 = board.create('curve', [[0],[0]], {strokeWidth:1, strokeColor:'black'});
    crl2.bezierDegree = 3;

    crl2.updateDataArray = function() {
        var A = board.create('point', [2, f1(2)], {visible: false});
        var B = board.create('point', [2, 0], {visible: false});
        var d = [B.X() - A.X(), B.Y() - A.Y()];
        var dl = Math.sqrt(d[0] * d[0] + d[1] * d[1]);
        var mid = [(A.X() + B.X()) * 0.5, (A.Y() + B.Y()) * 0.5];

        d[0] *= height2 / dl;
        d[1] *= height2 / dl;

        this.dataX = [ A.X(), A.X() - d[1], mid[0], mid[0] - d[1], mid[0], B.X() - d[1], B.X() ];
        this.dataY = [ A.Y(), A.Y() + d[0], mid[1], mid[1] + d[0], mid[1], B.Y() + d[0], B.Y() ];
    };

    // Text for the second curly bracket
    var txt2 = board.create('text', [
                        function() { 
                            var A = board.create('point', [2, f1(2)], {visible: false});
                            var B = board.create('point', [2, 0], {visible: false});
                            var d = [B.X() - A.X(), B.Y() - A.Y()];
                            var dl = Math.sqrt(d[0] * d[0] + d[1] * d[1]);
                            var mid = (A.X() + B.X()) * 0.5;

                            d[1] *= height2 / dl;
                            return mid - d[1] + 0.1; 
                        },
                        function() { 
                            var A = board.create('point', [2, f1(2)], {visible: false});
                            var B = board.create('point', [2, 0], {visible: false});
                            var d = [B.X() - A.X(), B.Y() - A.Y()];
                            var dl = Math.sqrt(d[0] * d[0] + d[1] * d[1]);
                            var mid = (A.Y() + B.Y()) * 0.5;

                            d[0] *= height2 / dl;
                            return mid + d[0] + 0.1; 
                        },
                        () => '\\(f(x)={#f2(x)#}\\)'
                    ]);

    board.update(); // This update is necessary to call updateDataArray for the brackets.
    [[/jsxgraph]] </p>
<p style="color:blue;">Note: Give your answer in exact form or correct to \(2\) decimal places where necessary</p>
<p>\( V=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>We will use the Washer method formula to find the volume of the solid generated by revolving the region bounded \(y = {@f1(x)@}\) and \(y = {@f2(x)@}\) about the \(x\)‐axis, 
<p>Let's first find the points of intersection:</p>
<p>\({@f1(x)@} = {@f2(x)@}\)  </p>
<p>\({@f1(x)-f2(x)@} = 0\)  </p>
<p>\({@factor(f1(x)-f2(x))@} = 0\)  </p>
<p>\(x = {@b@}\) or \(x = {@a@}\)</p>

</p>The limits of integration will be from \( {@a@}\) to \({@b@}\) (the \(x\)-values where the curves intersect).</p>

<p>The formula for this case would be:</p>

  <p>\( \displaystyle V = \pi \int_{a}^{b} \big((f(x))^2 - (g(x))^2\big) \, dx\)</p>

  <p>Where \(f(x)\) is the outer function and \(g(x)\) is the inner function.</p>

  <p>Given \(f(x)= {@f2(x)@}\) and \(g(x)= {@f1(x)@}\) about the \(x\)‐axis, we'll find the volume using the washer method formula:</p>

  <p>\( \displaystyle V = \pi \int_{{@a@}}^{{@b@}} \big (\big(  {@f2(x)@} \big)^2 - \big(  {@f1(x)@} \big)^2 \big) \, dx\)</p>

  <p>Let's solve this integral:</p>

  <p>\( \displaystyle V = \pi \int_{{@a@}}^{{@b@}} \big (\big(  {@f2(x)@} \big)^2 - \big(  {@f1(x)@} \big)^2 \big) \, dx\)</p>
  <p>\( \displaystyle V = \pi \int_{{@a@}}^{{@b@}} \big (\big(  {@eqn1@} \big) - \big(  {@eqn2@} \big) \big) \, dx\)</p>
  <p>\( \displaystyle V = \pi \int_{{@a@}}^{{@b@}} \big ( {@eqn3@} \big) \, dx\)</p>
  <p>\( \displaystyle V = \pi  \left ( {@eqn4@} \right)_{{@a@}}^{{@b@}} \)</p>
  <p>\( \displaystyle V = \pi \left(\big({@eqn5@}\big) - \big({@eqn6@}\big)\right) \)</p>
  <p>\(\displaystyle V = \pi \left({@eqn7@}\right)\)</p>
  <p>\(V = \boxed{{@volume@} \approx {@decimalplaces(volume,2)@} }\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/* Define the functions */ 
const1:rand(3)+1;
const2:const1+2;
f1(x) := x^2 + const1; 
f2(x) := x + const2;  
/* Find the intersection points */ 
intersection_points: solve(f1(x) = f2(x), x); 
b: rhs(intersection_points[1]); 
a: rhs(intersection_points[2]);  
x1:f1(a);
y1:f1(b);
/* Define the integral for volume calculation */ 
integral_expr: int((f2(x))^2 - (f1(x))^2, x, a, b);  
/* Calculate the volume */ 
volume: %pi*(integral_expr); 


/*worked soln*/
eqn1:expand((f2(x))^2);
eqn2: expand((f1(x))^2);
eqn3:eqn1-eqn2;
eqn4:int(eqn3,x);
eqn5:at(eqn4,[x=b]);
eqn6:at(eqn4,[x=a]);
eqn7:eqn5-eqn6;


/*board coord */
minx:-1*(const1+1);
miny:minx;
maxy:(-1*minx)+7;
maxx:-1*minx</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Find the volume of the solid generated by revolving the region bounded by \(y = {@f1(x)@}\) and \(y = {@f2(x)@}\) about the \(x\)‐axis.</p>
<p>\( V={@volume@}\) </p>

<p>Here is a sketch of the bounded region with the axis of rotation shown,</p>
<p>
[[jsxgraph]]
        var board = JXG.JSXGraph.initBoard(divid, { 
            boundingbox: [{#minx#}, {#maxy#}, {#maxx#}, {#miny#}],
            keepaspectratio: true,
            showCopyright: false,
            axis: true,
            showNavigation: false
        });
    // Define the functions
    var f1 = function(x) { return x*x + {#const1#}; };
    var f2 = function(x) { return x + {#const2#}; };

    // Plot the curves
    var curve1 = board.create('functiongraph', [f1, -3, 3], { strokeColor: 'blue' });
    var curve2 = board.create('functiongraph', [f2, -3, 3], { strokeColor: 'red' });

// Find the intersection points
    var interPoint1 = board.create('intersection', [curve1, curve2, 0], { visible: true, name:'\\(({#a#},{#x1#})\\)'});
    var interPoint2 = board.create('intersection', [curve1, curve2, 1], { visible: true, name:'\\(({#b#},{#y1#})\\)'});


    // Find the minimum point of the curve
    var minPoint = board.create('point', [0, f1(0)], { style: 5, size: 4, fillColor: 'green', strokeColor: 'green', visible: false});
    var minPointY = minPoint.Y();

    // Create the first curly bracket
    var height1 = 0.5; // height of the first curly bracket
    var crl1 = board.create('curve', [[0],[0]], {strokeWidth:1, strokeColor:'black' }, {visible: false});
    crl1.bezierDegree = 3;

    crl1.updateDataArray = function() {
        var A = minPoint, B = board.create('point', [minPoint.X(), 0], {visible: false});
        var d = [B.X() - A.X(), B.Y() - A.Y()];
        var dl = Math.sqrt(d[0] * d[0] + d[1] * d[1]);
        var mid = [(A.X() + B.X()) * 0.5, (A.Y() + B.Y()) * 0.5];

        d[0] *= height1 / dl;
        d[1] *= height1 / dl;

        this.dataX = [ A.X(), A.X() - d[1], mid[0], mid[0] - d[1], mid[0], B.X() - d[1], B.X() ];
        this.dataY = [ A.Y(), A.Y() + d[0], mid[1], mid[1] + d[0], mid[1], B.Y() + d[0], B.Y() ];
    };

    // Text for the first curly bracket
    var txt1 = board.create('text', [
                        function() { 
                            var A = minPoint, B = board.create('point', [minPoint.X(), 0], {visible: false});
                            var d = [B.X() - A.X(), B.Y() - A.Y()];
                            var dl = Math.sqrt(d[0] * d[0] + d[1] * d[1]);
                            var mid = (A.X() + B.X()) * 0.5;

                            d[1] *= height1 / dl;
                            return mid - d[1] + 0.1; 
                        },
                        function() { 
                            var A = minPoint, B = board.create('point', [minPoint.X(), 0], {visible: false});
                            var d = [B.X() - A.X(), B.Y() - A.Y()];
                            var dl = Math.sqrt(d[0] * d[0] + d[1] * d[1]);
                            var mid = (A.Y() + B.Y()) * 0.5;

                            d[0] *= height1 / dl;
                            return mid + d[0] + 0.1; 
                        },
                        () => 'g(x)={#f1(x)#}'
                    ]);

    // Create the second curly bracket
    var height2 = 0.5; // height of the second curly bracket
    var crl2 = board.create('curve', [[0],[0]], {strokeWidth:1, strokeColor:'black'});
    crl2.bezierDegree = 3;

    crl2.updateDataArray = function() {
        var A = board.create('point', [2, f1(2)], {visible: false});
        var B = board.create('point', [2, 0], {visible: false});
        var d = [B.X() - A.X(), B.Y() - A.Y()];
        var dl = Math.sqrt(d[0] * d[0] + d[1] * d[1]);
        var mid = [(A.X() + B.X()) * 0.5, (A.Y() + B.Y()) * 0.5];

        d[0] *= height2 / dl;
        d[1] *= height2 / dl;

        this.dataX = [ A.X(), A.X() - d[1], mid[0], mid[0] - d[1], mid[0], B.X() - d[1], B.X() ];
        this.dataY = [ A.Y(), A.Y() + d[0], mid[1], mid[1] + d[0], mid[1], B.Y() + d[0], B.Y() ];
    };

    // Text for the second curly bracket
    var txt2 = board.create('text', [
                        function() { 
                            var A = board.create('point', [2, f1(2)], {visible: false});
                            var B = board.create('point', [2, 0], {visible: false});
                            var d = [B.X() - A.X(), B.Y() - A.Y()];
                            var dl = Math.sqrt(d[0] * d[0] + d[1] * d[1]);
                            var mid = (A.X() + B.X()) * 0.5;

                            d[1] *= height2 / dl;
                            return mid - d[1] + 0.1; 
                        },
                        function() { 
                            var A = board.create('point', [2, f1(2)], {visible: false});
                            var B = board.create('point', [2, 0], {visible: false});
                            var d = [B.X() - A.X(), B.Y() - A.Y()];
                            var dl = Math.sqrt(d[0] * d[0] + d[1] * d[1]);
                            var mid = (A.Y() + B.Y()) * 0.5;

                            d[0] *= height2 / dl;
                            return mid + d[0] + 0.1; 
                        },
                        () => '\\(f(x)={#f2(x)#}\\)'
                    ]);

    board.update(); // This update is necessary to call updateDataArray for the brackets.
    [[/jsxgraph]] </p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author"> Godfrey Wabwire </p>
<p id="concept"> N/A </p>
<p id="reviewer">Mary Sayuni</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">Finding the volume of the solid generated by revolving the region bounded by \(y = {@f1(x)@}\) and \(y = {@f2(x)@}\) about the \(x\)‐axis using the Washer method formula.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>volume</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>volume</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for correct answer to 2dp</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>volume</tans>
        <testoptions>2</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p style="color:blue;">Well done for giving your answer correct to \(2\) decimal places</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>324955526</deployedseed>
    <deployedseed>93115717</deployedseed>
    <deployedseed>1951989228</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>volume</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>decimalplaces(volume,2)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
