<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Volumes of Solids of Revolution-Washer method-q3</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the volume of the solid generated by revolving the region bounded by \(y={@y1@}\) and \(y={@y2@}\) on the interval  \( \displaystyle \left[0,{@interval_2@}  \displaystyle \right ]\) about \(x-axis\) using the disk Washer method.</p>
<p> 
<p>The figure formed when revolved is as shown below</p>
<p>
[[jsxgraph]]
        var board = JXG.JSXGraph.initBoard(divid, { 
            boundingbox: [-1, 1.5, Math.PI / 2 + 1, -1.5],
            keepaspectratio: true,
            showCopyright: false,
            axis: true,
            showNavigation: false
        });

// Create functions
    var f = x => Math.sin(x / 2);
    var g = x => Math.cos(x / 2);

    // Plot the curves
     var curve1 = board.create('functiongraph', [f]);
     var curve2 = board.create('functiongraph', [g]);

  // Fill the region between the curves with green color
    var region = board.create('curve', [
        (t) => t,
        (t) => Math.max(f(t), g(t)),
        0, Math.PI/2
    ], {
        fillColor: 'green',
        fillOpacity: 2
    });
// Create a polygon to fill the region between the curves up to their intersection
        var region = board.create('polygon', [
            [0, f(0)],
            [0, g(0)],
            [1.6, g(1.6)]
        ], {
            fillColor: 'green',
            fillOpacity: 1,
            fixed:true 
        });

    // Find the intersection point
    var intersectionPoint = board.create('intersection', [curve1, curve2, 0], {
      visible: false, fixed: true
    });

    // Draw a dotted line from the intersection point to the x-axis at x = pi/2
    var dottedLine = board.create('line', [
      [intersectionPoint.X(), intersectionPoint.Y()],
      [Math.PI / 2, 0]
    ], {
      straightFirst: true,
      straightLast: true,
      dash: 2,
      strokeColor: 'black',
       fixed:true
    });

    // Label x = pi/2 on the x-axis
    var labelX = board.create('text', [Math.PI / 2, -0.2, 'x = \u03C0/{#aa#}'], {
      anchorX: 'middle',
      fontSize: 18
    });

    // Plot functions only up to the mutual intersections
       // var function1 = board.create('functiongraph', ['{#f_x#}', {#lower_limit#},{#upper_limit#}], { name: '\\({#simplify(y=f_x)#}\\)', withlabel: true, strokeColor: 'blue' });
        //var function2 = board.create('functiongraph', ['x={#lower_limit#}'],{ name: '\\({#x=lower_limit#}\\)', withlabel: true, strokeColor: 'blue' });

        // Note: This function goes backwards, i.e. from 3 to -3.
        //var function4 = board.create('functiongraph', ['x={#upper_limit#}'],{ name: '\\({#x=upper_limit#}\\)', withlabel: true, strokeColor: 'blue' });

        // You can copy this function and then use it like in the examples below.
        // element_list is a list of elements defining the outline of the function.
        // The elements can be FunctionGraphs or Pairs of coordinates (encoding points)
        // properties is a dictionary of properties for the generated function.
        function make_filled_curve(board, element_list, properties) {
            var curve = board.create('curve', [[], []], properties);
            curve.updateDataArray = function () {
                this.dataX = [];
                this.dataY = [];
                for (const element of element_list) {
                    if (Array.isArray(element)) {
                        // Add the point to the curve
                        this.dataX.push(element[0]);
                        this.dataY.push(element[1]);
                    } else {
                        // Copy all points from function
                        this.dataX = this.dataX.concat(element.points.map(p => p.usrCoords[1]));
                        this.dataY = this.dataY.concat(element.points.map(p => p.usrCoords[2]));
                    }
                }
                // Close the loop
                this.dataX.push(this.dataX[0]);
                this.dataY.push(this.dataY[0]);
            };
            board.update();
        }

        // Plot different filled functions
       // make_filled_curve(board, [curve1, curve2], { strokeWidth: 3, fillColor: 'yellow', fillOpacity: 0.3 });
  // make_filled_curve(board, [ [0,0], curve1, [0,1]], { strokeWidth: 1, fillColor: 'green', fillOpacity: 0.3 });
        //make_filled_curve(board, [function1, [0, 3], [-2, 3]], { strokeWidth: 3, fillColor: 'red', fillOpacity: 0.3 });
    [[/jsxgraph]] </p>             
</p>
           
<p>Volume \(=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Since the revolution is about the \(x-axis\), we can use the disk washer method to find the volume as follows.</p>
<p>\( \begin{align}
V&=\displaystyle \int_{a}^{b}\pi([f(x)]^2-[g(x)]^2)dx\\
&=\displaystyle \int_{0}^{{@interval_2@}}\pi([{@y2@}]^2-[{@y1@}]^2)dx\\
&=\displaystyle \int_{0}^{{@interval_2@}}\pi([1-\sin^2(\frac{x}{2})]-[\sin^2(\frac{x}{2})])dx\\
&=\displaystyle \int_{0}^{{@interval_2@}}\pi([1-2\sin^2(\frac{x}{2})])dx\\
&=\displaystyle \int_{0}^{{@interval_2@}}\pi \cos(x)dx\\
\end{align}
\)</p>
<p>use the trigonometric identity</p>
<p>\(\displaystyle \int \cos (x)\,dx=sin(x)+C\)</p>
<p>Applying the integral we have,</p>
<p>\(\begin{align}
\displaystyle \int_{0}^{{@interval_2@}}\pi \cos(x)dx &=\pi \displaystyle \left[  \sin(x) \displaystyle \right ]_0^{{@interval_2@}}
\\
&=\pi \displaystyle \left[ \sin({@interval_2@})-\sin(0) \displaystyle \right]\\
\end{align}\)
</p>
<p>Since \(\displaystyle \sin⁡ {@interval_2@}=\displaystyle {@sol3@}\) and \(\sin(0)=0\) we have,</p>
<p>\( \pi \cdot ({@sol3@}-0)=\displaystyle{@ta@}\)</p>
<p>Therefore the volume of the solid is \(\displaystyle{@ta@}\)</p>
<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Method of washers*/
y1:sin(x/2);
y2:cos(x/2);
interval_1:0;
aa:rand([2,4,6]);
interval_2:pi/aa;
/*solution variables*/
sol1:pi*((y2)^2-(y1)^2);
sol2:int(sol1,x,0,interval_2);
ta:sol2;
sol3:if interval_2=pi/2 then 1 elseif interval_2=pi/4 then sqrt(2)/2 elseif interval_2=pi/6 then 1/2 ;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Find the volume of the solid generated by revolving the region bounded by \(y={@y1@}\) and \(y={@y2@}\) on the interval  \( \displaystyle \left[0,{@interval_2@}  \displaystyle \right ]\) about \(x-axis\)</p>
<p>V={@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Wastalas</p>
<p id="description">This question focuses on finding the volume by slicing that is generated by revolving the area bounded by \(y={@y1@}\) and \(y={@y2@}\) on the interval  \( \displaystyle \left[0,{@interval_2@}  \displaystyle \right ]\) about \(x-axis\)</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer of the volume</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>470716919</deployedseed>
    <deployedseed>61575384</deployedseed>
    <deployedseed>678071426</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
