<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Volumes of Solids of Revolution-Washer method-q2</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the volume of the solid generated by revolving the region bounded by \(f(x)={@const1@}-{@var1*x^2@} , g(x)=\sqrt{{@const2@}-{@var2*x^2@} }, x ={@lower_limit@}, \) and \(x={@upper_limit@}\) about the 
\(x-\) axis using the disk Washer method.</p>
<p><em>Note: Give your answer in exact form.</em></p>
<p>\(V=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Because the revolution is about the \(x-\) axis, the volume is determined as follows</p>
<p>\[\begin{align} V&=\int_a^b \pi \left( [f(x)]^2-[g(x)]^2 \right) dx. \\
    &=\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}} \pi \left( [{@const1@}-x^2]^2-[\sqrt{{@const2@}-x^2}]^2 \right) dx.\\
 &= \displaystyle \pi \int_{{@lower_limit@}}^{{@upper_limit@}} \left( [{@expand(((const1-x^2))^2)@}]-[{@const2@}-x^2] \right) dx.\\
&=\displaystyle \pi\int_{{@lower_limit@}}^{{@upper_limit@}} \left[ {@expand(((const1-x^2))^2)-(const2-x^2)@} \right] dx.\\
&=\displaystyle \pi \bigg[ {@ta1@} \bigg]_{{@lower_limit@}}^{{@upper_limit@}}\\
&=\displaystyle {@ta@}
  \end{align}\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>const1:2;
const2:rand([1,1/4,1/9,1/16]);
var1:1;
var2:1;
lower_limit:-sqrt(const2);
upper_limit:sqrt(const2);
aa:expand(((const1-x^2))^2);
dd:int(2*x,x);
fx:const1-x^2;
gx:sqrt(const2-x^2);
ta1:int((((fx)^2-(gx)^2)),x);
ta:int((pi*((fx)^2-(gx)^2)),x,lower_limit,upper_limit);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Find the volume of the solid generated by revolving the region bounded by \(f(x)={@const1@}-{@var1*x^2@} , g(x)=\sqrt{{@const2@}-{@var2*x^2@} }, x ={@lower_limit@}, \) and \(x={@upper_limit@}\) about the 
\(x-\) axis.</p><p>\(V={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Nixson Kiplagat</p>
<p id="concept">N/A</p>
<p id="reviewer">Emmaculate Odhiambo</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">[Actual question description]</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node 1: Checks for the correct answer*/</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2075010521</deployedseed>
    <deployedseed>1260146959</deployedseed>
    <deployedseed>553004603</deployedseed>
    <deployedseed>91784313</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
