<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>CQ - Evaluating the Convergence and Value of the Improper Integral (Infinite interval of integration)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="border: 2px solid #000; background-color: #f0f0f0; padding: 10px; color: green; min-width:100px; display: inline-block">
<h2>Improper Integrals Involving Unbounded Intervals</h2> 
<p>This involves integration over an interval whose upper limit, (or lower limit) grows without bound. You can think of this as evaluating the area of an unbounded region.</p>
<p>An example of such integrals is \[\int_{-\infty}^{3}e^{2x}dx=\lim_{t\to-\infty}\int_{t}^{3}e^{2x}dx\] or \[\int_{1}^{\infty}\frac{1}{2x^{3}}dx=\lim_{t\to\infty}\int_{1}^{t}\frac{1}{2x^{3}}dx\]</p>
</div>
<p> Consider the following integral \(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx\). This is improper integral because  [[input:ans1a]] [[validation:ans1a]][[feedback:prt1a]]</p>
<div style="border: 2px solid #000; background-color: #f0f0f0; padding: 10px; color: green; min-width:100px; display: inline-block">
<p> To evaluate this type of an improper integral, we first express it as a limit. That is  \(\displaystyle\lim_{t\to\infty} \int_{{@lower_limit@}}^{t} {@fx@}\,dx\). We then use the <a href="https://activecalculus.org/single/sec-4-4-FTC.html" target="_blank">Fundamental Theorem of Calculus</a>
 to evaluate the integral \(\displaystyle \int_{{@lower_limit@}}^{t} {@fx@}\,dx\).</p>
 </div>
<p>To simplify the integral, we will need to make a \(u\) substitution. Here \(u=\) [[input:ans1b]] and \(du =\)[[input:ans1c]]\(dx\) [[validation:ans1b]][[validation:ans1c]].</p>
<p> Making this change of variable, the new integrand is given by [[input:ans1d]][[validation:ans1d]] and the limits of integration are
<ul>
  <li>lower limit =[[input:ans1e]][[validation:ans1e]] </li>
  <li>upper limit = [[input:ans1f]][[validation:ans1f]]</li>
</ul>

 So that the antiderivative is \(\displaystyle \int{@func_2@}\,du=\) [[input:ans1g]].[[validation:ans1g]]</p>
<p> Evaluating the antiderivative at the limits of integration gives </p>
<p>\(\displaystyle \int_{{@lower_limit@}}^{t}{@fx@}\,dx=\)[[input:ans1h]] [[validation:ans1h]][[feedback:prt1b]]</p> 

<div style="border: 2px solid #000; background-color: #f0f0f0; padding: 10px; color: green; min-width:100px; display: inline-block">
<h2>Convergence and divergence of an improper integral</h2>
<p>The limit \(\displaystyle \lim_{t\to\infty} \int_{{@lower_limit@}}^{t}{@fx@}\,dx\)
 is either finite or infinite (or it doesn’t exist). With these possibilities in mind, we introduce the following terminology.</p>

If \(f(x)\) is nonnegative for \(x\geq a\), then we say that the improper integral \(\displaystyle \int_{a}^{\infty}f(x)dx\) converges provided that \[\lim_{t\to\infty} \int_{a}^{t}f(x)dx\] exists and is finite. otherwise, we say that \(\displaystyle \int_{a}^{\infty}f(x)dx\) diverges.
</div>
<div style="border: 2px solid #000; background-color: #f0f0f0; padding: 10px; color: green; min-width:100px; display: inline-block">
<p> We now want to determine if the integral \(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx\) converges or diverges.  To do this, we evaluate the limit as \(t\longrightarrow \infty\) of your answer above. <em>You might want to review examples <a href="https://math.libretexts.org/@go/page/10189" target="_blank"> 2.5.5 and 2.5.6 </a> to refresh your memory on how to evaluate limits of this nature.</em>

(Note: \( \infty\) is entered as <b>inf</b> and \( -\infty\) is entered as <b>minf</b>). 
</p>
</div>
<p>\(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx=\lim_{t\to\infty} \int_{{@lower_limit@}}^{t}{@fx@}\,dx=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>Therefore, the integral is [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3> 
<p> First this is improper integral because the limits of integration include infinity.</p>
<p> So we express the integral as a limit. That is  \(\displaystyle\lim_{t\to\infty} \int_{{@lower_limit@}}^{t} {@fx@}\,dx\). </p>
<p>To simplify the integral, we will need to make a \(u\) substitution. Here \(u={@ta_1b@}\) and \(du ={@ta_1c@} dx\) .</p>
<p> Making this change of variable, the integrand is given by  \({@ta_1d@}\) and the limits of integration are
<ul>
  <li>lower limit  \(={@ta_1e@}\)</li>
  <li>upper limit  \(u={@ta_1f@}\)</li>
</ul>
 So that the antiderivative is  \({@ta_1g@}\) </p>
<p> Evaluating the antiderivative at the limits of integration gives </p>
<p>\(\displaystyle \int_{{@lower_limit@}}^{t}{@fx@}\,dx={@ta_1h@}\)</p> 
<p> Finally evaluating the limit as \(t\longrightarrow \infty\) of your answer above gives the value of the integral. (Note: \( \infty\) is entered as <b>inf</b> and \( -\infty\) is entered as <b>minf</b>). </p>
<p>\(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx=\lim_{t\to\infty} \int_{{@lower_limit@}}^{t}{@fx@}\,dx={@ta_1@}\) </>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Define the constants and variables */
num:rand(10)+1;
denom_const:rand(4)+4;
coeff: rand(10)+1;
power:rand(5)+2/(rand(4)+2);
upper_limit:inf;
lower_limit:rand(5)+1;
point_disc:t;
den:(coeff*x+denom_const)^power;
fx : num/den;
int_fx: int(fx, x, lower_limit, upper_limit);

/*worked solution*/
expr_1: coeff*x+denom_const;
diff_u: diff(expr_1,x);
func_1: fx*(1/diff_u);
func_2: subst(u,expr_1,func_1);
expr_2:subst(t,x,expr_1);
expr_3:ev(expr_2,t=lower_limit);
func_3: int(func_2,u);
int_fu: func_3;

lower_limit_u: ev(expr_1,x=lower_limit);
upper_limit_u: ev(expr_1,x=t);
func_4: ev(func_3,u=upper_limit_u)-ev(func_3,u=lower_limit_u);
func_5: ratsimp(func_4);


/*Model answers*/
ta_1a:random_permutation([["The limits of integration are unbounded (include infinity)", true],["The limits of integration contain a point of discontinuity", false],["The integrand is not bounded (not continuous at a point) in the interval of integration", false]]);

ta_1b: expr_1;
ta_1c: diff_u;
ta_1d: func_2;
ta_1e: lower_limit_u;
ta_1f: upper_limit_u;
ta_1g: func_3;
ta_1h: func_5;

ta_1: int_fx;
ta_2:if ta_1=inf then "Divergent" elseif ta_1=minf then "Divergent" else"Convergent";
ta_3:if ta_2="Divergent" then "Convergent" else "Divergent";
ta_4:random_permutation([[ta_2, true],[ta_3, false],["Divergent and convergent", false]]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Evaluate the integral below and determine if it converges or diverges.</p>
<p>\(\displaystyle\int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx={@ta_1@}\)</p>
<p>Therefore, the integral is {@mcq_correct(ta_4)[1]@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">1</p>
<p id="reviewer">2</p>
<p id="description">The question focuses on evaluating an Integral  into Inverse Trigonometric function by substitution by a defined formula</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1a</name>
      <type>checkbox</type>
      <tans>ta_1a</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1b</name>
      <type>algebraic</type>
      <tans>ta_1b</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1c</name>
      <type>algebraic</type>
      <tans>ta_1c</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1d</name>
      <type>algebraic</type>
      <tans>ta_1d</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1e</name>
      <type>algebraic</type>
      <tans>ta_1e</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1f</name>
      <type>algebraic</type>
      <tans>ta_1f</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1g</name>
      <type>algebraic</type>
      <tans>ta_1g</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1h</name>
      <type>algebraic</type>
      <tans>ta_1h</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta_4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sans1: ev(ans1,t=upper_limit);
/*node 1:checks for correct answers*/

/*node 2:Calculating the integral as F(a)-F(b) instead of F(b)-F(a)*/
mistake_1:int(fx,upper_lm, lower_lm) ;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct evaluation of the limit</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>sans1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>You correctly evaluated the limit</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>Check the evaluation of the limit. You might have made a mistake in the evaluation of the limit!</text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for the first mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You calculated the value of the integral in a wrong format \(F(a)-F(b)\),instead of \(F(b)-F(a)\)</p>
<p>Hint: Check the worked out solution</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks that the student limit is the value of the integral</description>
        <answertest>AlgEquiv</answertest>
        <sans>sans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>Check that you evaluated the antiderivative correctly.</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt1a</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct choices</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1a</sans>
        <tans>mcq_correct(ta_1a)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1a-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1a-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt1b</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check the correct u-substitution</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1b</sans>
        <tans>ta_1b</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1b-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>check the correct \(du\)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1c</sans>
        <tans>ta_1c</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1b-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>check the correct integrand</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1d</sans>
        <tans>ta_1d</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1b-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>check the correct lower limit of new integral</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1e</sans>
        <tans>ta_1e</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1b-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>check the correct upper limit of the new integral</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1f</sans>
        <tans>ta_1f</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1b-5-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description>Check the correct antiderivative</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1g</sans>
        <tans>ta_1g</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty/>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt1b-6-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-6-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description>Check the correct evaluation of limits of integration</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1h</sans>
        <tans>ta_1h</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1b-7-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-7-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_4)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1713077684</deployedseed>
    <deployedseed>1135097558</deployedseed>
    <deployedseed>1085257259</deployedseed>
    <deployedseed>1735702239</deployedseed>
    <deployedseed>1382360763</deployedseed>
    <deployedseed>1184361143</deployedseed>
    <deployedseed>1897229937</deployedseed>
    <deployedseed>84512596</deployedseed>
    <deployedseed>320197684</deployedseed>
    <deployedseed>1611767008</deployedseed>
    <deployedseed>1410145781</deployedseed>
    <deployedseed>1544667461</deployedseed>
    <deployedseed>1666898777</deployedseed>
    <deployedseed>324379575</deployedseed>
    <deployedseed>1252877097</deployedseed>
    <deployedseed>1847195690</deployedseed>
    <deployedseed>2050226279</deployedseed>
    <deployedseed>1553150849</deployedseed>
    <deployedseed>784009488</deployedseed>
    <deployedseed>148913360</deployedseed>
    <deployedseed>900855818</deployedseed>
    <deployedseed>1420873199</deployedseed>
    <deployedseed>584399493</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta_4))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <hint format="html">
      <text><![CDATA[<p>This is an improper integral because the upper limit is \(\infty \). So, you need to rewrite the integral in terms of a limit:</p>
<p>\(\displaystyle \int_{{@lower_lm@}}^{{@upper_lm@}}{@fx@}\,dx=\lim_{b \to \infty } \displaystyle \int_{{@lower_lm@}}^{{@upper_lm@}}{@fx@}\,dx\)</p>
<p>Solve the integral taking the upper limit  as finite.</p>
<p>The integrand, \(\displaystyle {@fx@}\,dx\) is a rational function of \(x\), which suggests a standard power rule for integration.</p>
<p>\(\displaystyle \int\frac{1}{(x+a)^n}\,dx=\displaystyle \frac{(x+a)^{(1-n)}}{1-n}\) for \(n\neq 1\)</p>]]></text>
    </hint>
  </question>
</quiz>
