<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>MQH - An improper integral on an infinite interval</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Evaluate the integral below and determine if it converges or diverges. 
</p>
<p>Note: \( \infty\) is entered as <b>inf</b> and \( -\infty\) is entered as <b>minf</b>.</p>
<p>\(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@func@}\,dx=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>Therefore, the integral is [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Recall that we can’t evaluate the integral as long as there is an infinity in one of the limits. Therefore, we’ll need to replace infinity with a parameter \(t\) as follows,
</p>
<p>The given improper integral is</p>

<p>
\(\displaystyle  I = \int_{@lower_limit@}^{{@upper_limit@}}{@func@} \, dx. \)
</p>

<p>We will first evaluate whether the integral converges or diverges by finding the limit of the integral as \( t \to \infty \):</p>

<p>
\( \displaystyle I =  \int_{@lower_limit@}^{{@t@}}{@func@} \, dx. \)
</p>
<p>To find antiderivative of \(\displaystyle {@func@}  \) we apply exponential rule: \( \displaystyle \int a^x \, dx = \frac{a^x}{\ln(a)} \,\) with \(a=e\)</p>
   <p>Hence  \( \displaystyle \int e^{x} \, dx = e^{x}\)</p>
   <p>\(\qquad \implies  \displaystyle  {@const@} \int {@fx@} \, dx = {@func@} \)</p>
<p>Evaluate limit points:</p>
<p>
\(
\begin{align}
\displaystyle 
\int_{{@lower_limit@}}^{{@t@}} {@func@} \, dx &= \Bigg[ {@func@} \Bigg]_{{@lower_limit@}}^{{@t@}} \\
                              &= {@indef_val@} \\
\end{align}
\)
</p>
<p>Take the limit as \( t \to \infty \)</p>

<p>
\(\displaystyle  I(t) = \lim_{t \to \infty} \left( {@indef_val@}\right) = {@ta_1@} \)
</p>

<p>The improper integral is {@mcq_correct(ta_4)[1]@}.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*Defining the question*/
power:rand(3)+2;
const:(rand(10)+3)/(rand(5)+1);
func:const*%e^(x);
fx: %e^(x);
/*worked solution*/
int_fx: int(fx,x);
upper_limit:inf;
upper_lmt:t;
lower_limit:rand(5)+2;

upper:at(func,x=t);
lower:at(func,x=lower_limit);
indef_val: upper-lower;

/*Model answers*/
ta_1: limit(indef_val, t, inf);
ta_2:if ta_1=inf then "Divergent" elseif ta_1=minf then "Divergent" else"Convergent";
ta_3:if ta_2="Divergent" then "Convergent" else "Divergent";
ta_4:random_permutation([[ta_2, true],[ta_3, false],["Divergent and convergent", false]]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\displaystyle\int_{{@lower_limit@}}^{{@upper_limit@}}{@func@}\,dx={@ta_1@}\)</p>
<p>Therefore, the integral is {@mcq_correct(ta_4)[1]@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Mike</p>
<p id="description">The question focuses on evaluating an improper Integral  by parts</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta_4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_4)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>580199743</deployedseed>
    <deployedseed>92216480</deployedseed>
    <deployedseed>998671051</deployedseed>
    <deployedseed>1232590071</deployedseed>
    <deployedseed>1195084257</deployedseed>
    <deployedseed>1308885196</deployedseed>
    <deployedseed>74934039</deployedseed>
    <deployedseed>1138953322</deployedseed>
    <deployedseed>2062482540</deployedseed>
    <deployedseed>2065205317</deployedseed>
    <deployedseed>2052446670</deployedseed>
    <deployedseed>1260386620</deployedseed>
    <deployedseed>1282441321</deployedseed>
    <deployedseed>1917639779</deployedseed>
    <deployedseed>1245349126</deployedseed>
    <deployedseed>980794931</deployedseed>
    <deployedseed>48362390</deployedseed>
    <deployedseed>2051585219</deployedseed>
    <deployedseed>1032710032</deployedseed>
    <deployedseed>1545868104</deployedseed>
    <deployedseed>1476902002</deployedseed>
    <deployedseed>1147549053</deployedseed>
    <deployedseed>2128474347</deployedseed>
    <deployedseed>535427140</deployedseed>
    <deployedseed>367018388</deployedseed>
    <deployedseed>639023529</deployedseed>
    <deployedseed>2014952589</deployedseed>
    <deployedseed>1363450042</deployedseed>
    <deployedseed>380681776</deployedseed>
    <deployedseed>100552814</deployedseed>
    <deployedseed>1565384485</deployedseed>
    <deployedseed>134931666</deployedseed>
    <deployedseed>693232943</deployedseed>
    <deployedseed>1288666056</deployedseed>
    <deployedseed>926518868</deployedseed>
    <deployedseed>289805965</deployedseed>
    <deployedseed>347470223</deployedseed>
    <deployedseed>2007111449</deployedseed>
    <deployedseed>575596891</deployedseed>
    <deployedseed>1566144808</deployedseed>
    <deployedseed>1472822535</deployedseed>
    <deployedseed>639537797</deployedseed>
    <deployedseed>2101204700</deployedseed>
    <deployedseed>1430260922</deployedseed>
    <deployedseed>483757411</deployedseed>
    <deployedseed>81358987</deployedseed>
    <deployedseed>1184931003</deployedseed>
    <deployedseed>1981910375</deployedseed>
    <deployedseed>522437371</deployedseed>
    <deployedseed>1262161485</deployedseed>
    <deployedseed>268363647</deployedseed>
    <deployedseed>1509781287</deployedseed>
    <deployedseed>2059963965</deployedseed>
    <deployedseed>1013802726</deployedseed>
    <deployedseed>1154306504</deployedseed>
    <deployedseed>37915452</deployedseed>
    <deployedseed>798999460</deployedseed>
    <deployedseed>905482329</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta_4))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <hint format="html">
      <text><![CDATA[<p>The integral involves an exponential function, \(e^x\), which grows rapidly for \(x\gt 0\)</p>
<p>To evaluate the improper integral, rewrite it as a limit</p>
<p>\(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@func@}\,dx=\lim_{b \to \infty } \int_{{@lower_limit@}}^{b}{@func@}\,dx \)</p>
<p>For an improper integral to converge, the function being integrated must decrease rapidly enough to approach zero as \(x\to \infty \)</p>]]></text>
    </hint>
  </question>
</quiz>
