<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>*** Using Geometric Formulas to Calculate Definite Integrals</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Use the formula for the area of a circle to evaluate : \(\displaystyle \int_{@-radius@}^{@radius@} \sqrt{{@radius^2@}-\left({@var1-x^2@}\right)}dx\)</p>

<p></p>
<p>Note: You need to find the area under the curve over the interval \(\left [{@-radius@},{@radius@}\right]\).</p>
<br>
<p>Therefore, area under the curve is: \(\displaystyle A =\frac{1}{2}\pi r^2 = \)[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>/p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h2>Worked solution</h2>
<p>\(\displaystyle \int_{@-radius@}^{@radius@} \sqrt{{@radius^2@}-\left({@var1-x^2@}\right)}dx\)</p>

<p>Recall the equation of a circle  is given as: \(\displaystyle x^2 + y^2 = r^2 \implies y= \sqrt {r^2-x^2} \)</p>
<p> We have the function </p>
<p>\(\displaystyle  y= \sqrt{{@radius^2@}-\left({@var1-x^2@}\right)}\),  and the interval {@-radius@}, {@radius@}</p>
<p>Let's first visualize the curve over the given interval as follows.</p>

[[jsxgraph]]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [{#-radius-1#}, {#radius+1#}, {#radius+1#}, {#-radius-1#}],
    keepaspectratio: true,
    showCopyright: false,
    axis: true,
    showNavigation: true
});

// Creating points p and q.
var point1 = board.create('point', [0, 0], { name: '' });
var point2 = board.create('point', [{#radius#}, 0], { name: 'q' });
var point3 = board.create('point', [{#-radius#}, 0], { name: 'p' });

// Creating the circle with center (0, 0) and radius {#radius#}
var circle = board.create('circle', [point1, point2]);
[[/jsxgraph]]
<p>The limits of integration are from {@-radius@} to {@radius@}, which means we need to find the area of the semicircle between {@-radius@} and {@radius@}.</p>

<p>\(\displaystyle A =\frac{1}{2}\pi r^2 = {@decimalplaces(ta1,3)@}\) </p>
<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Board dimensions*/
radius:rand(4)+2;
var1:rand(5)+1;
func1:sqrt(radius^2-(var1-x^2));
ta1:1/2*(22/7)*radius^2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\displaystyle \int_{@-radius@}^{@radius@} \sqrt{{@radius^2@}-\left({@var1-x^2@}\right)}dx = {@ta1@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>920968199</deployedseed>
    <deployedseed>1830767697</deployedseed>
    <deployedseed>664610122</deployedseed>
    <deployedseed>2105492502</deployedseed>
    <deployedseed>274851883</deployedseed>
    <deployedseed>1711105131</deployedseed>
    <deployedseed>636471604</deployedseed>
    <deployedseed>1238516173</deployedseed>
    <deployedseed>1510670119</deployedseed>
    <deployedseed>919817382</deployedseed>
    <deployedseed>1371278628</deployedseed>
    <deployedseed>218918913</deployedseed>
    <deployedseed>1558725803</deployedseed>
    <deployedseed>1827728403</deployedseed>
  </question>
</quiz>
