<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Approximating Area using Riemann Sums_Definite integral</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Use Riemann Sum to approximate area under the curve of \(\displaystyle f(x) = {@func@}\) over the interval \([{@left_limit@}, {@right_limit@}]\) with \(n = {@rect_num@}\) rectangles.</p>

<p>The function has been plotted as shown below:</p>

<hr>
[[jsxgraph]]
/*Values for bounding box.*/
var x_min = {#x_min#};
var x_max = {#x_max#};
var y_min = {#y_min#};
var y_max = {#y_max#};

var rect_width = {#rect_width#};
var rect_num = {#rect_num#};
var left_limit = {#left_limit#};

var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [x_min, y_max, x_max, y_min],
    keepaspectratio: false,
    showCopyright: false,
    axis: true,
    showNavigation: false
});

var f = board.jc.snippet('{#func#}', true, 'x', true);
var sketch = board.create('functiongraph', [f, x_min, x_max]);

/* Add labels*/
var label1 = board.create('text', [x_min + 2, y_max - 2, '\\(f(x) = {#func#}\\)'], {
    fixed: true,
    anchorX: 'left',
    fontSize: 15,
    useMathJax: true,
    parse: false
});
[[/jsxgraph]]
<hr>

<br> 

<p>Compute the approximate area under the curve using the left endpoint.
<br>Give your answer to 2 decimal places.</p>
<p>\(\displaystyle A_{L} = \) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>

<br>

<p>Compute the area under the curve using the right endpoint.
<br>Give your answer to 2 decimal places.</p>
<p>\(\displaystyle A_{R} = \) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>

<br>

<p>Evaluate the definite integral, giving your answer to 2 decimal places.
<br>\(\displaystyle \int_2^{10} \left( {@func@}\right) \mathrm{d}x = \) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>

<p>For the interval \([a,b] = [{@left_limit@}, {@right_limit@}]\) and  \( n={@rect_num@}\), the first step is to find the partition width \(\displaystyle \Delta x = \frac{b-a}{n} = \frac{{@right_limit@}-{@left_limit@}}{@rect_num@} = {@rect_width@}\)</p>

<br>

<p>Let \(x_i\) be the right endpoint of the \(i\)th rectangle and \(x_{i-1}\) be the left endpoint.</p>
<p>Consider the sketch below with {@rect_num@} rectangular partitions of width {@rect_width@} each. </p>

<hr>
[[jsxgraph]]
/*Values for bounding box.*/
var x_min = {#x_min#};
var x_max = {#x_max#};
var y_min = {#y_min#};
var y_max = {#y_max#};

var rect_width = {#rect_width#};
var rect_num = {#rect_num#};

var left_limit = {#left_limit#};

var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [x_min, y_max, x_max, y_min],
    keepaspectratio: false,
    showCopyright: false,
    axis: true,
    showNavigation: false
});

var f = board.jc.snippet('{#func#}', true, 'x', true);

var sketch = board.create('functiongraph', [f, x_min, x_max]);

/* Add labels*/
var label1 = board.create('text', [x_min + 2, y_max - 2, '\\(f(x)={#func#}\\)'], {
    fixed: true,
    anchorX: 'left',
    fontSize: 15,
    useMathJax: true,
    parse: false
});

/*Create rectangles*/
var y_right_rect_vals = {#rect_top_y_coords_lowersum#};
var y_left_rect_vals = {#rect_top_y_coords_uppersum#};
var position = "right";
var y_rect_vals = [];

if (position == "right") {
    y_rect_vals = y_right_rect_vals
} else {
    y_rect_vals = y_left_rect_vals
};
for (var i = 1; i <= rect_num; i++) {
    var rectangle = board.create('polygon', [
        [left_limit + (i - 1) * rect_width, 0],
        [left_limit + (i - 0) * rect_width, 0],
        [left_limit + (i - 0) * rect_width, y_rect_vals[i-1]], 
        [left_limit + (i - 1) * rect_width, y_rect_vals[i-1]] 
    ], {
        fixed: true,
        highlight: false,
        color: '#A8FE9E',
        borders: { fixed: true, highlight: false },
        vertices: { visible: false }
    })
};
[[/jsxgraph]]
<hr>

<p>To find the height of the rectangles, evaluate the function at the left endpoint of each rectangle.</p>

<p>For the left endpoint approximation \(\displaystyle A_L\), we compute the sum of the areas of rectangles touching the curve from the left side.</p>

<p style="margin-left:30px">\(\begin{align*}
    A_L &= \sum_{i=0}^{@rect_num-1@} f(x_i)\Delta x \\
    &= ({@sum_fxi_lower@})\Delta x \\
    &= ({@sum_efxi_lower@})\cdot{@rect_width@} \\
    &= {@ta1@} \hspace{10px} \text{(2dp)}
\end{align*}\)</p>

<p>\(n\) is the number of rectangles (subintervals).</p>
<p>\(f(x_i)\) is the value of the function at the \(i^{th}\) partition point.</p>
<p>\(\displaystyle \Delta x \) is the width of each rectangle (partition width).</p>

<br>
<hr>
[[jsxgraph]]
/*Values for bounding box.*/
var x_min = {#x_min#};
var x_max = {#x_max#};
var y_min = {#y_min#};
var y_max = {#y_max#};

var rect_width = {#rect_width#};
var rect_num = {#rect_num#};
var left_limit = {#left_limit#};
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [x_min, y_max, x_max, y_min],
    keepaspectratio: false,
    showCopyright: false,
    axis: true,
    showNavigation: false
});
var f = board.jc.snippet('{#func#}', true, 'x', true);
var sketch = board.create('functiongraph', [f, x_min, x_max]);

/* Add labels*/
var label1 = board.create('text', [x_min + 2, y_max - 2, '\\(f(x)={#func#}\\)'], {
    fixed: true,
    anchorX: 'left',
    fontSize: 15,
    useMathJax: true,
    parse: false
});

/*Create rectangles*/
var y_right_rect_vals = {#rect_top_y_coords_lowersum#};
var y_left_rect_vals = {#rect_top_y_coords_uppersum#};
var position = "left";
var y_rect_vals = [];
if (position == "right") {
    y_rect_vals = y_right_rect_vals
} else {
    y_rect_vals = y_left_rect_vals
};

for (var i = 1; i <= rect_num; i++) {
    var rectangle = board.create('polygon', [
        [left_limit + (i - 1) * rect_width, 0],
        [left_limit + (i - 0) * rect_width, 0],
        [left_limit + (i - 0) * rect_width, y_rect_vals[i-1]], 
        [left_limit + (i - 1) * rect_width, y_rect_vals[i-1]] 
    ], {
        fixed: true,
        highlight: false,
        color: '#A8FE9E',
        borders: { fixed: true, highlight: false },
        vertices: { visible: false }
    })
};
[[/jsxgraph]]
<hr>

<p>For the Right endpoint approximation \(\displaystyle A_L\), we use the formula</p>

<p style="margin-left:30px">\(\begin{align*} 
    A_R &= \sum_{i=1}^{@rect_num@} f(x_i)\Delta x \\
    &= ({@sum_fxi_upper@})\Delta x \\
    &= ({@sum_efxi_upper@})\cdot{@rect_width@} \\
    &= {@ta2@} \hspace{10px} \text{(2dp)} \\
\end{align*}\)</p>

<p>This time we take the sum of the areas of rectangles touching the curve at the right-hand side.</p>

<br>

<p>Finally, \(\displaystyle \int_{@left_limit@}^{{@right_limit@}} \left( {@func@}\right) \mathrm{d}x = {@ta3@} \hspace{10px} \text{(2dp)} \)]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*Variables for the function definition*/
aa:rand_with_step(20,40,5);
bb:rand(20)+2;
func1:x^2/aa + bb/10;

cc:rand(5)+1;
dd:rand(3)+2;
func2:x/dd + cc;

ee:rand_with_step(15,30,5);
ff:rand_with_step(4,10,1);
func3:x^2/ee + x/ff + 1/10;

gg:rand_with_step(115,120,1);
hh:rand_with_step(15,20,1);
jj:rand(20);
func4:x^3/gg + x/hh + 1/jj;
func:rand([func1,func2, func3,func4]);
f(y) := ev(func, x=y);

/*Variables for rectangle definition*/
rect_num:rand([3,4,5]);
rect_width:rand([1,2]);
left_limit:rand(4)+1;   /*Lower bound*/
right_limit:left_limit + rect_num * rect_width; /*Upper bound*/
rect_left_x_coords: makelist(i,i,left_limit,right_limit-rect_width,rect_width); /*x-coords of left corners*/

/*Variables for bounding box*/
x_min:-1;
x_max:right_limit+1;
y_min:-1;
y_vals:makelist(f(i),i,x_min,x_max,0.1);
y_max:ceiling(last(sort(y_vals)));

/*List of y-values for rectangles*/
rect_top_y_coords_lowersum: makelist(f(i), i, left_limit, right_limit-rect_width, rect_width);
rect_top_y_coords_uppersum: makelist(f(i+rect_width), i, left_limit, right_limit-rect_width, rect_width);

/* Computer lower/upper sum and integral*/
ta1:decimalplaces(apply("+", rect_top_y_coords_lowersum)*rect_width,2);
ta2:decimalplaces(apply("+", rect_top_y_coords_uppersum)*rect_width,2);
ta3:decimalplaces(int(func,x,left_limit,right_limit),2);

/* Strings for typesetting step by step solution */
indices: makelist(i, i, 1, rect_num);
fxi_lower: makelist(sconcat("f(x_", i, ")"), i, 0, rect_num-1);
efxi_lower: makelist(rect_top_y_coords_lowersum[i], i, 1, rect_num);
sum_fxi_lower: simplode(fxi_lower, "+");
sum_efxi_lower: simplode(efxi_lower, "+");
fxi_upper: makelist(sconcat("f(x_", i, ")"), i, 1, rect_num);
efxi_upper: makelist(rect_top_y_coords_uppersum[i], i, 1, rect_num);
sum_fxi_upper: simplode(fxi_upper, "+");
sum_efxi_upper: simplode(efxi_upper, "+");
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(\displaystyle f(x) = {@func@}\) over the interval \([{@left_limit@}, {@right_limit@}]\) with \(n = {@rect_num@}\) rectangles.
</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Juma Zevick</p>
<p id="concept">N/A</p>
<p id="reviewer">N/A</p>
<p id="reviewer"></p>
<p id="description"This question involves approximating the area under the curve f(x) over the interval [2,10] using both Left and Right endpoints of a Riemann sum with n=4 rectangles. Students will calculate these approximations and evaluate the definite integral for the given function. </p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2138456832</deployedseed>
    <deployedseed>819071794</deployedseed>
    <deployedseed>1369052303</deployedseed>
    <deployedseed>1377888873</deployedseed>
    <deployedseed>1182804454</deployedseed>
    <deployedseed>1663438218</deployedseed>
    <deployedseed>822200543</deployedseed>
    <deployedseed>194478737</deployedseed>
    <deployedseed>2058957272</deployedseed>
    <deployedseed>799135742</deployedseed>
    <deployedseed>699668082</deployedseed>
    <deployedseed>124603246</deployedseed>
    <deployedseed>1920414820</deployedseed>
    <deployedseed>1251720478</deployedseed>
    <deployedseed>334976423</deployedseed>
    <deployedseed>1537689942</deployedseed>
    <deployedseed>1887206944</deployedseed>
    <deployedseed>542615993</deployedseed>
    <deployedseed>136663545</deployedseed>
    <deployedseed>1071812084</deployedseed>
    <deployedseed>1526890471</deployedseed>
    <deployedseed>1487528370</deployedseed>
    <deployedseed>1004319102</deployedseed>
    <deployedseed>1415036359</deployedseed>
    <deployedseed>1043264303</deployedseed>
    <deployedseed>239944665</deployedseed>
    <deployedseed>2135091625</deployedseed>
    <deployedseed>2141329998</deployedseed>
    <deployedseed>712673075</deployedseed>
    <deployedseed>1585261631</deployedseed>
    <deployedseed>723011187</deployedseed>
    <deployedseed>1544958017</deployedseed>
    <deployedseed>956905891</deployedseed>
    <deployedseed>233302925</deployedseed>
    <deployedseed>274367481</deployedseed>
    <deployedseed>1175705319</deployedseed>
    <deployedseed>250114939</deployedseed>
    <deployedseed>1183407762</deployedseed>
    <deployedseed>928569270</deployedseed>
    <deployedseed>573100076</deployedseed>
    <deployedseed>1094229504</deployedseed>
    <deployedseed>705063633</deployedseed>
    <deployedseed>247502111</deployedseed>
    <deployedseed>151908276</deployedseed>
    <deployedseed>1106067815</deployedseed>
    <deployedseed>11963518</deployedseed>
    <deployedseed>1602172454</deployedseed>
    <deployedseed>201643537</deployedseed>
    <deployedseed>932362368</deployedseed>
    <deployedseed>1850163842</deployedseed>
    <deployedseed>213442806</deployedseed>
    <deployedseed>1362955009</deployedseed>
    <deployedseed>1618975161</deployedseed>
    <deployedseed>1973403201</deployedseed>
    <deployedseed>424000477</deployedseed>
    <deployedseed>610359530</deployedseed>
    <deployedseed>825053374</deployedseed>
    <deployedseed>202534156</deployedseed>
    <deployedseed>558853958</deployedseed>
    <deployedseed>1752967057</deployedseed>
    <deployedseed>805676118</deployedseed>
    <deployedseed>22149287</deployedseed>
    <deployedseed>299812502</deployedseed>
    <deployedseed>1592371094</deployedseed>
    <deployedseed>1044103775</deployedseed>
    <deployedseed>1315695226</deployedseed>
    <deployedseed>340381583</deployedseed>
    <deployedseed>259938978</deployedseed>
    <deployedseed>58621906</deployedseed>
    <deployedseed>1588391949</deployedseed>
    <deployedseed>543398773</deployedseed>
    <deployedseed>623502484</deployedseed>
    <deployedseed>1740772391</deployedseed>
    <deployedseed>1291370509</deployedseed>
    <deployedseed>1119151205</deployedseed>
    <deployedseed>986162459</deployedseed>
    <deployedseed>225031958</deployedseed>
    <deployedseed>1589603276</deployedseed>
    <deployedseed>588503632</deployedseed>
    <deployedseed>873227986</deployedseed>
    <deployedseed>1118817003</deployedseed>
    <deployedseed>368086086</deployedseed>
    <deployedseed>2065905103</deployedseed>
    <deployedseed>1224206267</deployedseed>
    <deployedseed>1300100368</deployedseed>
    <deployedseed>31285454</deployedseed>
    <deployedseed>544602877</deployedseed>
    <deployedseed>1682333678</deployedseed>
    <deployedseed>1879015065</deployedseed>
    <deployedseed>157940866</deployedseed>
    <deployedseed>1475887524</deployedseed>
    <deployedseed>1253346013</deployedseed>
    <deployedseed>1444658342</deployedseed>
    <deployedseed>1930077401</deployedseed>
    <deployedseed>1533576655</deployedseed>
    <deployedseed>255497708</deployedseed>
    <deployedseed>750897426</deployedseed>
    <deployedseed>1749798243</deployedseed>
    <deployedseed>1663218699</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
