<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Riemann Sums and Definite Integrals: Calculations and Limits.</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Given the definite integral \(\displaystyle \int_{@aa@}^{@bb@}{@func@} dx\)</p>
<br>
<p> Compute the following and fill in the blanks.</p>
<p>\(\displaystyle \Delta x = \frac{b-a}{n} \implies \Delta x = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p style="margin-bottom:12px">\(\displaystyle x_j = a+ j\Delta x \implies x_j = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>
<p>\(\displaystyle f(x_j) = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>
<br>
<p>\(\displaystyle \sum_j ^n  f(x_j) \Delta x = \) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]] </p>
<br>
<p>Evaluate the integral .</p>
<p>\(\displaystyle \lim_{{n \to \infty}} \sum_j ^n  f(x_j) \Delta x = \) [[input:ans5]] [[validation:ans5]] [[feedback:prt5]] </p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p>Given that \(\displaystyle \int_a^b f(x) = \lim_{{n \to \infty}} \sum_j ^n  f(x_j) \Delta x  \)</p>

<p>Let us go ahead and evaluate this integral \(\displaystyle \int_{@aa@}^{@bb@}{@func@}\); </p>

<p>\(\displaystyle \Delta x = \frac{b-a}{n} \implies \Delta x = {@ta1@} \) </p>
<p style="margin-bottom:12px">\(\displaystyle x_j = a+ j\Delta x \implies x_j = {@ta2@}\) </p>
<p>\(\displaystyle f(x_j) = {@ta3@}\) </p>
<p>\( \displaystyle f(x_j) \Delta x = {@f_xj@}\)</p>

<br>
<p>Now we go ahead and compute \(\displaystyle \sum_j ^n  f(x_j)\)</p>

<p>We need to substitute the following when setting up the Reimann sum:</p>
<p> \(\displaystyle \sum_j ^n j^2 = \frac{n(n+1)(2n+1)}{6}\) and \(\displaystyle  \sum_j ^n j = \frac{n(n+1)}{n}\), we can go ahead and form the Reimann sum as shown,</p>

<p>\(\displaystyle \sum_j ^n  f(x_j) \Delta x = {@coef_j2@}\left( {@sum_j2@} \right ) + {@coef_j@}\left( {@sum_j@} \right )+{@coef_j0@} = {@ta4@}\) </p>


<br>
<p>Finally,</p>
<p>\(\displaystyle \lim_{{n \to \infty}} \sum_j ^n  f(x_j) \Delta x  \implies \lim_{{n \to \infty}} {@ta4@} = {@ta5@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand_with_step(-5,0,1);
bb:rand_with_step(1,5,1);
var1:rand(2)+1; 
func1:x^(var1);
func2:expand((x-aa)*(x-bb));
func3:expand((x+aa)*(x+aa));
func:rand([func1,func2,func3]);

/*Solutions of the steps*/
ta1:(bb-aa)/n;
ta2:aa+ta1*j;
ta3:subst(ta2,x,func);
f_xj:expand(ta3*ta1);

coef_j0:coeff(f_xj,j,0);
coef_j:coeff(f_xj,j);
coef_j2:coeff(f_xj,j^2);

sum_j:(n*(n+1))/2;
sum_j2:(n*(n+1)*(2*n+1))/6;

ta4:expand(coef_j2*sum_j2+coef_j*sum_j+coef_j0);

ta5:int(func,x,aa,bb);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\displaystyle \int_{@a@}^{@b@}  {@func@} = {@int(func,x,aa,bb)@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Juma Zevick</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Dan Kelly</p>
<p id="description">This question involves evaluating a definite integral using the right-endpoint approximation and the concept of Riemann sum. The student is asked to calculate the integral of a given function over a specified interval using the method of approximating the area under the curve using the right endpoints of subintervals.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node2:The student forgets to subtract a from b in the formula*/
sans_error1:(aa+bb)/n;
/*Node3: The student confuses the order of a and b*/
sans_error2:(aa-bb)/n;
/*Node4: The student mistakenly multiplies a and b instead of subtracting them.*/
sans_error3:(aa*bb)/n;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>The student forgets to subtract a from b in the formula</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>sans_error1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your calculation of \(\Delta x\) is not quite accurate. Remember to subtract \(a\) from \(b\) to find the interval length correctly: \(\Delta x = \frac{b - a}{n}\). Double-check your subtraction and the order of the terms.
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>The student confuses the order of a and b</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>sans_error2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It seems like there's an issue with the order of \(a\) and \(b\) in your calculation of \(\Delta x\). Make sure to subtract \(a\) from \(b\) in the correct order: : \(\Delta x = \frac{b - a}{n}\). </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>The student mistakenly multiplies a and b instead of subtracting them.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>sans_error3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p> Instead of multiplying  \(a\) and \(b\), remember that you need to subtract \(a\) from \(b\) to find the interval length: \(\Delta x = \frac{b - a}{n}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node2: The student forgets to multiply j by delta x  in the formula xj=a+jΔx.*/
sans_error21:aa+ta1;

/*Node3: The student confuses the order of terms in the formula xj=a-jΔx*/
sans_error22:aa-ta1*j;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>The student forgets to multiply j by delta x  in the formula xj=a+jΔx.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>sans_error21</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Make sure to use the formula \(\displaystyle x_j = a+ j\Delta x \) and ensure you include the multiplication by \(j\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>checks for second expected error </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>sans_error22</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p> Remember to use the formula \(displaystyle x_j = a- j\Delta x\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node2:student carries forward an incorrect answer from input ans2 to input ans3 and uses it correctly in input ans3*/
sans_error31:subst(ans2,x,func);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for first expected error</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>sans_error31</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You've correctly used the value of \(xj\) in this step, which shows that you understand its calculation.  However, remember that the value you used for \(x_j\)  is from your previous calculation, which might have been incorrect.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*Node2:Student carries forward incorrect value from previous steps (input 3 and 1) and uses it correctly in input ans4. */
f_xj_error:expand(ans3*ans1);
coef_j0_error:coeff(f_xj_error,j,0);
coef_j_error:coeff(f_xj_error,j);
coef_j2_error:coeff(f_xj_error,j^2);

sum_j:(n*(n+1))/2;
sum_j2:(n*(n+1)*(2*n+1))/6;

sans_error41:expand(coef_j2_error*sum_j2+coef_j_error*sum_j+coef_j0_error);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Student carries forward incorrect value from previous steps (input 3 and 1) and uses it correctly in input ans4. </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>sans_error41</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It's great to see you using the value of \(f(x_j)\) and \(\Delta x\)  in this step. However, please note that the value you used might have been based on an earlier incorrect calculation in input ans3. Double-check your calculation in both inputs above.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>404681051</deployedseed>
    <deployedseed>2098884755</deployedseed>
    <deployedseed>583985754</deployedseed>
    <deployedseed>1222997169</deployedseed>
    <deployedseed>1156014003</deployedseed>
    <deployedseed>1388205436</deployedseed>
    <deployedseed>79090377</deployedseed>
    <deployedseed>243698839</deployedseed>
    <deployedseed>1352377746</deployedseed>
    <deployedseed>1481795353</deployedseed>
    <deployedseed>2103627231</deployedseed>
    <deployedseed>243949089</deployedseed>
    <deployedseed>1414753417</deployedseed>
    <deployedseed>1011985641</deployedseed>
    <deployedseed>1052415593</deployedseed>
    <deployedseed>828378871</deployedseed>
    <deployedseed>1006525275</deployedseed>
    <deployedseed>1511884195</deployedseed>
    <deployedseed>498236379</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
