<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Using Geometric Formulas to Calculate Definite Integrals (Polygons)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Evaluate the integral to find the area under the curve: \(\displaystyle \int_{@lower@}^{@upper@}   {@func1@} dx\)</p>
<p><em>Hint: You will have to sketch the function.</em> </p>
<br>
<p>\( A = \)[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h2>Worked solution</h2>


<p>The graph of the function \(y=  {@func1@}\) over the interval \(\left( 0,{@upper@}\right)\) is shown below.</p>
<br>
[[jsxgraph]]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [0, 10, 8, -2],
    keepaspectratio: true,
    showCopyright: false,
    axis: true,
    showNavigation: true
});
// Define the function y = mx + c
function f(x) {
    var m = {#grad#};
    var c = {#intercept#};
    return m * x + c;
}
// Create the curve for y = mx + c
var curve = board.create('functiongraph', [f, -1, 7]);

// Create the point
var point1 = board.create('point', [0, 0], { name: ''} );
var point2 = board.create('point', [{#upper#},0], { name: '' });
var point3 = board.create('point', [{#upper#}, {#height#}], { name: 'y={#func1#}' });

// Create the polygon using points p1, p2, and p3
var polygon = board.create('polygon', [point1, point2, point3], { fillColor: 'rgba(0, 0, 255, 0.3)' });
[[/jsxgraph]]

<p></p>
<p>To calculate the area under curve we'll use the formula \(\displaystyle A= \frac{1}{2}b \times h\)</p>
<p>\(\displaystyle \frac{1}{2}\times {@base@} \times {@height@}= {@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>lower:0;
upper:rand_with_step(3,5,1);
grad:rand(2)+1;
intercept:0;
func1:grad*x+intercept;
base:upper;
height:subst(upper,x,func1);
ta:(1/2)*base*height;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\displaystyle \int_{@-radius@}^{@radius@} \sqrt{{@radius^2@}-\left({@var1-x^2@}\right)}dx = {@ta1@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
