<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>GQ - Natural log improper Integral</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Evaluate the integral \(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx\) and determine if it converges or diverges. 
<p> First this is improper integral because  [[input:ans1a]] [[validation:ans1a]][[feedback:prt1a]]</p>
<p> So we express the integral as a limit. That is  \(\displaystyle\lim_{t\to{@const1@}} \int_{t}^{{@upper_limit@}} {@fx@}\,dx\). </p>
<p>To evaluate the integral, we will need to apply integration by parts. Recall the formula for integration by parts \[\int u.dv=u.v-\int v.du\] </p>
<p> You need to be careful with your choices of \(u\) and \(dv\) that will simplify the integral. So the correct choices are :
<ul>
  <li>\(u\) =[[input:ans1b]] so that \(du=\)[[input:ans1c]][[validation:ans1b]][[validation:ans1c]] </li>
 <li>\(dv\) =[[input:ans1d]] so that \(v=\)[[input:ans1e]][[validation:ans1d]][[validation:ans1e]] </li>
</ul></p>
<p> Applying these substitutions in the formula for integration by parts, we obtain the antiderivative</p>
<p>\(\displaystyle \int {@fx@}\,dx=\) [[input:ans1f]]-\(\displaystyle \int \) [[input:ans1g]]\(dx\) [[validation:ans1f]][[validation:ans1g]]</p>
<p> and the value of the antiderivative is </p>
<p>\(\displaystyle \int {@fx@}\,dx=\) [[input:ans1h]][[validation:ans1h]]</p>
<p> Evaluating the antiderivative at the limits of integration gives </p>
<p>\(\displaystyle \int_{t}^{{@upper_limit@}}{@fx@}\,dx=\)[[input:ans1i]] [[validation:ans1i]][[feedback:prt1b]]</p> 
<p> Finally evaluating the limit as \(t\longrightarrow 0\) of your answer above gives the value of the integral. (Note: \( \infty\) is entered as <b>inf</b> and \( -\infty\) is entered as <b>minf</b>). </p>
<p>\(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx=\lim_{t\to{@const1@}} \int_{t}^{{@upper_limit@}}{@fx@}\,dx=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>Therefore, the integral is [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Worked solution</b></p>
<p>To evaluate the integral \(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx\) and determine if it converges or diverges, we proceed as follows;</p>
<p> First this is improper integral because the limits of integration contain a point of discontinuity, the integrand is not continuous at a point in the interval of integration and the integrand is not defined a point in the interval of integration</p>
<p> So we express the integral as a limit. That is  \(\displaystyle\lim_{t\to{@const1@}} \int_{t}^{{@upper_limit@}} {@fx@}\,dx\). </p>
<p>To evaluate the integral, we will need to apply integration by parts. Recall the formula for integration by parts \[\int u.dv=u.v-\int v.du\] </p>
<p> You need to be careful with your choices of \(u\) and \(dv\) that will simplify the integral. So the correct choices are :
<ul>
  <li>\(u= {@ta_1b@}\) so that \(du={@ta_1c@}\)</li>
 <li>\(dv={@ta_1d@}\) so that \(v={@ta_1e@}\)</li>
</ul></p>
<p> Applying these substitutions in the formula for integration by parts, we obtain the antiderivative</p>
<p>\(\displaystyle \int {@fx@}\,dx={@ta_1f@}-\displaystyle \int {@ta_1g@}dx\) </p>
<p> and the value of the antiderivative is </p>
<p>\(\displaystyle \int {@fx@}\,dx={@ta_1h@}\)</p>
<p> Evaluating the antiderivative at the limits of integration gives </p>
<p>\(\displaystyle \int_{t}^{{@upper_limit@}}{@fx@}\,dx={@ta_1i@}\)</p> 
<p> Finally evaluating the limit as \(t\longrightarrow 0\) of your answer above gives the value of the integral. (Note: \( \infty\) is entered as <b>inf</b> and \( -\infty\) is entered as <b>minf</b>). </p>
<p>\(
  \begin{align}
\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx&=\lim_{t\to{@const1@}} \int_{t}^{{@upper_limit@}}{@fx@}\,dx\\
  \displaystyle\int_{t}^{{@upper_limit@}}{@expr_udv@}\,dx &=\left[{@expr_uv@} \right]_{t}^{{@upper_limit@}}+ \int_{t}^{{@upper_limit@}} {@expr_vdu@}\,dx \\
  & ={@expr_inta@} -{@expr_intb@} \\
&= {@expr_int_t@}
  \end{align}
  \)</p>
<p>We now need to apply L'Hospitals rule to evaluate the limit as \(t\to{@const1@}\):</p> 
<p> \(\displaystyle \lim_{t \to 0}\, t^2 ln(a*t)=0\)</p>
<p>Applying this rule into the integral, we get:
</p>
<p>    \(\begin{align}
  \displaystyle\int_{{@lower_limit@}}^{{@upper_limit@}}{@expr_udv@}\,dx &=\lim_{t \to 0}{@expr_int_t@}\\
       &={@ta1@}\\
    \end{align}\)</p>
<p>Therefore, the integral converges </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*new improper questions*/
coeff1: rand(9)+1;
coeff2:rand(9)+1;
const1:0;
lower_limit:const1;
upper_limit:rand(5)+1+lower_limit;
fx:(coeff1*x)*log(coeff2*x);

	
/*worked out solution variables*/
expr_u:log(coeff2*x-const1*coeff2);
expr_du:diff(expr_u,x);
expr_dv:(coeff1*x);
expr_v:int(expr_dv,x);
expr_udv: expr_u*expr_dv;
expr_uv: expr_u*expr_v;
expr_vdu: expr_du*expr_v;
antider_fx: int(fx,x);
expr_inta: ev(expr_uv, x=upper_limit)-ev(expr_uv, x=t), simp;
expr_intb: int(expr_vdu,x,t,upper_limit);
expr_int_t: expand(expr_inta-expr_intb), simp;
int_fx: int(fx,x, lower_limit, upper_limit);

ta_1a:random_permutation([["The limits of integration include infinity", false],["The limits of integration contain a point of discontinuity", true],["The integrand is not continuous at a point in the interval of integration", true],["The integrand is not defined a point in the interval of integration", true]]);

ta_1b: expr_u;
ta_1c: expr_du;
ta_1d: expr_dv;
ta_1e: expr_v;
ta_1f: expr_uv;
ta_1g: expr_vdu;
ta_1h: antider_fx;
ta_1i: expr_int_t;

ta: limit(expr_int_t, t, const1);
ta1: simplify(ta);
ta2: random_permutation([["converges", true],["diverges", false],["converges and diverges", false]]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx={@ta1@}\)</p>
<p>Therefore, the integral is {@mcq_correct(ta2)[1]@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Mike</p>
<p id="description">The question focuses on evaluating an improper Integral  by parts</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1a</name>
      <type>checkbox</type>
      <tans>ta_1a</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1b</name>
      <type>algebraic</type>
      <tans>ta_1b</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1c</name>
      <type>algebraic</type>
      <tans>ta_1c</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1d</name>
      <type>algebraic</type>
      <tans>ta_1d</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1e</name>
      <type>algebraic</type>
      <tans>ta_1e</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1f</name>
      <type>algebraic</type>
      <tans>ta_1f</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1g</name>
      <type>algebraic</type>
      <tans>ta_1g</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1h</name>
      <type>algebraic</type>
      <tans>ta_1h</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1i</name>
      <type>algebraic</type>
      <tans>ta_1i</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the antiderivative</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt1a</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct choices of why it is an improper integral</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1a</sans>
        <tans>mcq_correct(ta_1a)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1a-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1a-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt1b</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks the correct choice of u</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1b</sans>
        <tans>ta_1b</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1b-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks the correct evaluation of du</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1c</sans>
        <tans>ta_1c</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1b-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks for the correct choice of dv</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1d</sans>
        <tans>ta_1d</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1b-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Checks for the correct evaluation of v</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1e</sans>
        <tans>ta_1e</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1b-4-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Checks for the correct evaluation of \(u*v\)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1f</sans>
        <tans>ta_1f</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1b-5-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description>Checks for the correct evaluation of \(v*du\)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1g</sans>
        <tans>ta_1g</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt1b-6-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-6-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description>Checks for the correct antiderivative</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1h</sans>
        <tans>ta_1h</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty/>
        <truenextnode>7</truenextnode>
        <trueanswernote>prt1b-7-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-7-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <description>Checks for the correct antiderivative with correct limits</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1i</sans>
        <tans>ta_1i</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1b-8-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1b-8-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta2)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1276415170</deployedseed>
    <deployedseed>1940473086</deployedseed>
    <deployedseed>1324912757</deployedseed>
    <deployedseed>1976971798</deployedseed>
    <deployedseed>690124662</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta_4))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
