<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>MQH - Natural log improper Integral</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Calculate the integral \(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx\)</p>
<p>\(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx=\)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>Therefore, the integral is [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p><b>Worked solution</b></p>
<p>First, we need to recall that we can’t do the integral as long as there is a point of discontinuity for the integrand between the limits of integration. Therefore, we’ll need to replace the point of discontinuity with a parameter \(t\) as follows,
</p>
<p>\[\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@expr_udv@}\,dx =\lim_{t \to\{@const1@}} \displaystyle \int_{t}^{{@upper_limit@}}{@expr_udv@}\,dx\]</p>
<p> </p>
<p>We now need to use the integration by parts technique to evaluate the integral. The formula for integration by parts is given by \[\displaystyle \int udv=u\cdot v-\int vdu\] </p>
<p> We now need to carefully consider the integrand and make a choice for \(u\) and \(dv\) that will simplify the integral.</p>
<p>Let \(u={@expr_u@}\) and \(dv={@expr_dv@}\,dx\)</p>
<p>\(u={@expr_u@} \, \Rightarrow du={@expr_du@}\,dx\)</p>
<p>\(dv={@expr_dv@}\,dx \, \Leftarrow v={@expr_v@}\)</p>
<p>\[
  \begin{align}
  \displaystyle\int_{t}^{{@upper_limit@}}{@expr_udv@}\,dx &=\left[{@expr_uv@} \right]_{t}^{{@upper_limit@}}+ \int_{t}^{{@upper_limit@}} {@expr_vdu@}\,dx \\
  & ={@expr_inta@} -{@expr_intb@} \\
&= {@expr_int_t@}
  \end{align}
  \]</p>
<p>We now need to apply L'Hospitals rule to evaluate the limit as \(t\to{@const1@}\):</p> 
<p> \[\displaystyle \lim_{t \to 0}\, t^2 ln(a*t)=0\]</p>
<p>Applying this rule into the integral, we get:
</p>


<p>    \[    \begin{align}
  \displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@expr_udv@}\,dx &=\lim_{t \to \infty}{@expr_int_t@}\\
       &={@ta1@}\\
    \end{align}    \]</p>
<p>Therefore, the integral converges </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/*new improper questions*/
coeff1: rand(9)+1;
coeff2:rand(9)+1;
const1:0;
lower_limit:const1;
upper_limit:rand(5)+1+lower_limit;
fx:(coeff1*x)*log(coeff2*x);

	
/*worked out solution variables*/
expr_u:log(coeff2*x-const1*coeff2);
expr_du:diff(expr_u,x);
expr_dv:(coeff1*x);
expr_v:int(expr_dv,x);
expr_udv: expr_u*expr_dv;
expr_uv: expr_u*expr_v;
expr_vdu: expr_du*expr_v;
antider_fx: int(fx,x);
expr_inta: ev(expr_uv, x=upper_limit)-ev(expr_uv, x=t), simp;
expr_intb: int(expr_vdu,x,t,upper_limit);
expr_int_t: expand(expr_inta-expr_intb), simp;
int_fx: int(fx,x, lower_limit, upper_limit);

ta_1a:random_permutation([["The limits of integration include infinity", false],["The limits of integration contain a point of discontinuity", true],["The integrand is not continuous at a point in the interval of integration", true],["The integrand is not defined a point in the interval of integration", true]]);

ta_1b: expr_u;
ta_1c: expr_du;
ta_1d: expr_dv;
ta_1e: expr_v;
ta_1f: expr_uv;
ta_1g: expr_vdu;
ta_1h: antider_fx;
ta_1i: expr_int_t;

ta: limit(expr_int_t, t, const1);
ta1: simplify(ta);
ta2: random_permutation([["converges", true],["diverges", false],["converges and diverges", false]]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\displaystyle \int_{{@lower_limit@}}^{{@upper_limit@}}{@fx@}\,dx={@ta1@}\)</p>
<p>Therefore, the integral is {@mcq_correct(ta2)[1]@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Mike</p>
<p id="description">The question focuses on evaluating an improper Integral  by parts</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the antiderivative</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta2)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1276415170</deployedseed>
    <deployedseed>1940473086</deployedseed>
    <deployedseed>1324912757</deployedseed>
    <deployedseed>1976971798</deployedseed>
    <deployedseed>690124662</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta_4))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
