<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>MQH - Evaluating the Convergence and Value of the Improper Integral_Research</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Evaluate the integral below and determine if it converges or diverges. 
</p>
<p>Note: \( \infty\) is entered as <b>inf</b> and \( -\infty\) is entered as <b>minf</b>.</p>
<p>\(\displaystyle \int_{{@lower_lm@}}^{{@upper_lm@}}{@fx@}\,dx=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>Therefore, the integral is [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Recall that we can’t evaluate the integral as long as there is an infinity in one of the limits. Therefore, we’ll need to replace infinity with a parameter \(t\) as follows,
</p>
<p>The given improper integral is</p>

<p>
\(\displaystyle  I = \int_{@lower_lm@}^{{@upper_lm@}}{@fx@} \, dx. \)
</p>

<p>We will first evaluate whether the integral converges or diverges by finding the limit of the integral as \( t \to \infty \):</p>

<p>
\( \displaystyle I =  \int_{@lower_lm@}^{{@t@}}{@fx@} \, dx. \)
</p>

<p>The antiderivative of \(\displaystyle {@fx@}  \) is:</p>

<p>
\(\displaystyle  \int {@fx@} \, dx = {@indef_int@} \)
</p>

<p>Now, evaluate the integral from \( {@lower_lm@} \) to \( t\):</p>

<p>
\(\displaystyle  I(t) = \left[  {@indef_int@} \right]_ {@lower_lm@}^t = {@indef_val@}\)
</p>

<p>Take the limit as \( t \to \infty \)</p>

<p>
\(\displaystyle  I = \lim_{t \to \infty} \left( {@indef_val@}\right) = {@def_integral@} \)
</p>

<p>The improper integral is {@mcq_correct(ta_4)[1]@}, and its exact value is \( \boxed{{@def_integral@}} \).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Define the constants and variables */
num:rand(10)+1;
denom_const:rand(4)+4;
power:rand(5)+2/(rand(4)+2);
upper_lm:inf;
lower_lm:rand(5)+1;
limit_pt:t;
den:(x+denom_const)^power;
fx : num/den;

/* Compute the indefinite integral */
indef_int: int(fx, x);
/* Compute the definite integral from 1 to t */
upper:at(indef_int,x=limit_pt);
lower:at(indef_int,x=lower_lm);
indef_val: upper-lower;

/* Compute the definite integral from 1 to infinity */
def_integral: limit(indef_val, limit_pt, inf);
/*Model answers*/
ta_1: limit(indef_val, limit_pt, inf);
ta_2:if ta_1=upper_lm then "Divergent" elseif ta_1=minf then "Divergent" else"Convergent";
ta_3:if ta_2="Divergent" then "Convergent" else "Divergent";
ta_4:random_permutation([[ta_2, true],[ta_3, false],["Divergent and convergent", false]]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Evaluate the integral below and determine if it converges or diverges.</p>
<p>\(\displaystyle\int_{{@lower_lm@}}^{{@upper_lm@}}{@fx@}\,dx={@ta_1@}\)</p>
<p>Therefore, the integral is {@mcq_correct(ta_4)[1]@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Mary Sayuni</p>
<p id="concept">N/A</p>
<p id="reviewer">1</p>
<p id="reviewer">2</p>
<p id="description">The question focuses on evaluating an Integral  into Inverse Trigonometric function by substitution by a defined formula</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta_4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/*node 1:checks for correct answers*/

/*node 2:Calculating the integral as F(a)-F(b) instead of F(b)-F(a)*/
mistake_1:int(fx,upper_lm, lower_lm) ;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>checks for the first mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You calculated the value of the integral in a wrong format \(F(a)-F(b)\),instead of \(F(b)-F(a)\)</p>
<p>Hint: Check the worked out solution</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_4)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2053929365</deployedseed>
    <deployedseed>1713077684</deployedseed>
    <deployedseed>1135097558</deployedseed>
    <deployedseed>1085257259</deployedseed>
    <deployedseed>336429531</deployedseed>
    <deployedseed>1735702239</deployedseed>
    <deployedseed>1382360763</deployedseed>
    <deployedseed>1184361143</deployedseed>
    <deployedseed>1897229937</deployedseed>
    <deployedseed>84512596</deployedseed>
    <deployedseed>320197684</deployedseed>
    <deployedseed>474741903</deployedseed>
    <deployedseed>1611767008</deployedseed>
    <deployedseed>1410145781</deployedseed>
    <deployedseed>1544667461</deployedseed>
    <deployedseed>1666898777</deployedseed>
    <deployedseed>324379575</deployedseed>
    <deployedseed>565780063</deployedseed>
    <deployedseed>1252877097</deployedseed>
    <deployedseed>1847195690</deployedseed>
    <deployedseed>2050226279</deployedseed>
    <deployedseed>1553150849</deployedseed>
    <deployedseed>784009488</deployedseed>
    <deployedseed>1267603643</deployedseed>
    <deployedseed>148913360</deployedseed>
    <deployedseed>202939344</deployedseed>
    <deployedseed>900855818</deployedseed>
    <deployedseed>1420873199</deployedseed>
    <deployedseed>584399493</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta_4))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <hint format="html">
      <text><![CDATA[<p>This is an improper integral because the upper limit is \(\infty \). So, you need to rewrite the integral in terms of a limit:</p>
<p>\(\displaystyle \int_{{@lower_lm@}}^{{@upper_lm@}}{@fx@}\,dx=\lim_{b \to \infty } \displaystyle \int_{{@lower_lm@}}^{{@upper_lm@}}{@fx@}\,dx\)</p>
<p>Solve the integral taking the upper limit  as finite.</p>
<p>The integrand, \(\displaystyle {@fx@}\,dx\) is a rational function of \(x\), which suggests a standard power rule for integration.</p>
<p>\(\displaystyle \int\frac{1}{(x+a)^n}\,dx=\displaystyle \frac{(x+a)^{(1-n)}}{1-n}\) for \(n\neq 1\)</p>]]></text>
    </hint>
    <hint format="html">
      <text><![CDATA[<p>Start by recognizing that the integrand is a rational function with the form \(\displaystyle  \frac{{@num@}}{({@(x+denom_const)@})^p} \), where \(\displaystyle  p ={@power@} \). </p>
<p>Recall that for Integrals of the form \( \int \frac{1}{(x+c)^p} \, dx \), the antiderivative is:</p>
<p>
\(\displaystyle 
\int \frac{1}{(x+c)^p} \, dx = \frac{(x+c)^{1-p}}{1-p} + C
\)</p>
<p>
Here, the exponent is \( p ={@power@} \), so the antiderivative can be computed by applying this formula.</p>]]></text>
    </hint>
  </question>
</quiz>
